<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Range — a parabola in vertex form</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the function \( f(x) = {@a@}(x-{@h@})^2 + {@k@} \).</p>
<p>State the range of \(f\) as an inequality in \(y\).</p>
<div>\(y\) [[input:ans_sym]] [[validation:ans_sym]] [[input:ans_val]] [[validation:ans_val]]</div>
<div>[[feedback:prt_sym]] [[feedback:prt_val]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>In vertex form \(f(x)={@a@}(x-{@h@})^2+{@k@}\), the point \(({@h@},{@k@})\) is the vertex. Since \((x-{@h@})^2 \geq 0\) always,</p>
<ul>
<li>if \(a &gt; 0\), the parabola opens <strong>upward</strong> and the vertex is a <strong>minimum</strong>: range is \(y \geq {@k@}\).</li>
<li>if \(a &lt; 0\), the parabola opens <strong>downward</strong> and the vertex is a <strong>maximum</strong>: range is \(y \leq {@k@}\).</li>
</ul>
<p>Here \(a={@a@}\), so the range is \(y\) {@lorsymtext@} \({@k@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a : rand_with_prohib(-5,5,[0]);
h : rand_with_prohib(-6,6,[]);
k : rand_with_prohib(-6,6,[]);
opens_up : is(a>0);
ge_opt : "ge"; le_opt : "le";
opts : [[ge_opt, opens_up, "\\geq"],[le_opt, not opens_up, "\\leq"]];
ta_sym : if opens_up then ge_opt else le_opt;
ta_val : k;
wrong_val : h;                          /* used the x-coordinate of the vertex instead of the y-value */
lorsymtext : if opens_up then "at least" else "at most";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.4 (Range). Range of a parabola in vertex form; catches getting the inequality direction wrong for a<0, and confusing the vertex's x-coordinate with its y-coordinate.</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.4 (Range). Range of a parabola in vertex form; catches getting the inequality direction wrong for a<0, and confusing the vertex's x-coordinate with its y-coordinate.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans_sym</name>
      <type>dropdown</type>
      <tans>opts</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans_val</name>
      <type>algebraic</type>
      <tans>ta_val</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt_sym</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct direction of inequality (matches whether parabola opens up or down)?</description>
        <answertest>String</answertest>
        <sans>ans_sym</sans>
        <tans>ta_sym</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_sym-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_sym-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[Check the sign of \(a={@a@})\): if \(a&gt;0\) the parabola opens upward (range is \(\geq\) the minimum); if \(a&lt;0\) it opens downward (range is \(\leq\) the maximum).]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_val</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct bound value (the y-coordinate of the vertex)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_val</sans>
        <tans>ta_val</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_val-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_val-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Used the x-coordinate of the vertex instead of the y-coordinate?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_val</sans>
        <tans>wrong_val</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_val-usedh</trueanswernote>
        <truefeedback format="html">
          <text>That is \(h\), the \(x\)-coordinate of the vertex. The range bound is the \(y\)-coordinate of the vertex, \(k\).</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_val-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>That is not the y-coordinate of the vertex.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Fully correct</description>
      <testinput>
        <name>ans_sym</name>
        <value>ta_sym</value>
      </testinput>
      <testinput>
        <name>ans_val</name>
        <value>ta_val</value>
      </testinput>
      <expected>
        <name>prt_sym</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_sym-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt_val</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_val-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Used vertex x-coordinate instead of y</description>
      <testinput>
        <name>ans_sym</name>
        <value>ta_sym</value>
      </testinput>
      <testinput>
        <name>ans_val</name>
        <value>wrong_val</value>
      </testinput>
      <expected>
        <name>prt_sym</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_sym-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt_val</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt_val-usedh</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
