<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding Garden area from a slider and using it to form and solve the quadratic</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The diagram shows a rectangular garden whose length is \({@dd@}\) metres longer than its width. Let the width be \(x\) metres.</p>
<p>Drag the slider to choose a garden, and read its area from the diagram. You will use that area in part (a).</p>
[[jsxgraph input-ref-ans0="ref_x" width="470px" height="470px"]]
var DD = {#dd#};
var brd = JXG.JSXGraph.initBoard(divid, {boundingbox:[-3, 17, 17, -3], keepaspectratio:true, axis:false, showCopyright:false, showNavigation:false, grid:true});
var sl = brd.create('slider', [[1,-2],[13,-2],[0, 0, 10]], {name:'x', snapWidth:1, precision:0});
var wf = function(){ return Math.round(sl.Value()); };
var lf = function(){ return Math.round(sl.Value()) + DD; };
var af = function(){ return wf()*lf(); };
var p1 = brd.create('point',[0,0],{visible:false,fixed:true});
var p2 = brd.create('point',[wf,0],{visible:false,fixed:true});
var p3 = brd.create('point',[wf,lf],{visible:false,fixed:true});
var p4 = brd.create('point',[0,lf],{visible:false,fixed:true});
var poly = brd.create('polygon',[p1,p2,p3,p4],{fillColor:'#8fce8f',fillOpacity:0.5,borders:{strokeColor:'#2f6f2f',strokeWidth:2}});
brd.create('text',[function(){return wf()/2;}, -0.7, 'x'],{fixed:true,anchorX:'middle',fontSize:18,cssStyle:'font-style:italic'});
brd.create('text',[-1, 16.2, function(){ return 'area = '+af()+' m²'; }],{fixed:true,anchorX:'left',fontSize:16});
stack_jxg.bind_slider(ref_x, sl);
[[/jsxgraph]]
<div style="display:none">[[input:ans0]] [[validation:ans0]]</div>
<p><strong>a)</strong> Using the area you found from the slider, form a quadratic equation to represent the area of the garden:</p>
<div>Equation: \( \) [[input:ans1]] \( = 0 \) [[validation:ans1]] [[feedback:prt1]]</div>
<p><strong>b)</strong> Solve your quadratic equation to find the dimensions of the garden (in metres).</p>
<div>Width \( x = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</div>
<div>Length \( = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Whatever garden you choose, if the width is \(x\) then the length is \(x+{@dd@}\), so \[ \text{area} = x\,(x+{@dd@}). \]</p>
<p>Setting this equal to the area \(A\) you read from the slider gives \[ x(x+{@dd@}) = A \quad\Longrightarrow\quad x^2 + {@dd@}x - A = 0. \]</p>
<p>Factorising (or using the quadratic formula) gives one positive root — that is the width — and the length is \({@dd@}\) m more.</p>
<p><em>Example:</em> a width of \({@wex@}\) m gives length \({@wex+dd@}\) m and area \({@Aex@}\) m\(^2\), and \[ x^2+{@dd@}x-{@Aex@}=(x-{@wex@})(x+{@wex+dd@})=0, \] so \(x={@wex@}\) (width) and the length is \({@wex+dd@}\) m.</p>]]></text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>dd : rand_with_prohib(2,5,[]);   /* length is dd m longer than width */
xstart : 4;                      /* initial slider width (an example, not an answer) */
wex : xstart;                    /* example width for the displayed worked solution */
Aex : wex*(wex+dd);              /* example area */
/* Example model answers (display only; real marking uses the slider value ans0). */
tans1 : x^2 + dd*x - Aex;
tans2 : wex;
tans3 : wex + dd;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>length is \({@dd@}\) m more than width; area chosen from slider</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans0</name>
      <type>algebraic</type>
      <tans>wex</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>tans1</tans>
      <boxsize>16</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>tans2</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>tans3</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options>hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>xw : round(ans0);
Achosen : xw*(xw + dd);
target1 : x^2 + dd*x - Achosen;
implied_area : -ev(ans1, x=0, simp);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Quadratic matches the area chosen on the slider?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>target1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Correct structure (x^2 + d x - K) but a different area K?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>x^2 + dd*x - implied_area</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-wrongarea</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You set the width to \({@xw@}\) m on the slider, so the area you found is \({@Achosen@}\text{ m}^2\). But your equation represents an area of \({@implied_area@}\text{ m}^2\). Rewrite it using the area you found from the slider.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The area is width \(\times\) length \(= x(x+{@dd@})\). Set this equal to the area you found on the slider, \({@Achosen@}\text{ m}^2\), and rearrange to \(x^2+{@dd@}x-({@Achosen@})=0\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>xw : round(ans0);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct width (the slider value)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>xw</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Gave the length instead of the width?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>xw + dd</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-length</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the length (the larger dimension). The width is the positive root of your quadratic.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Gave the negative root?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>-(xw + dd)</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-neg</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>A width cannot be negative — reject the negative root and take the positive solution.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Solve your quadratic \(x^2+{@dd@}x-(\text{area})=0\); the width is the positive root.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>xw : round(ans0);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct length (width + d)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>xw + dd</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Gave the width instead of the length?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>xw</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-width</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the width. The length is \({@dd@}\) m more than the width.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The length is \({@dd@}\) m more than the width you found.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>8001</deployedseed>
    <deployedseed>8002</deployedseed>
    <deployedseed>8003</deployedseed>
    <deployedseed>8004</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All correct for a chosen width of 5</description>
      <testinput>
        <name>ans0</name>
        <value>5</value>
      </testinput>
      <testinput>
        <name>ans1</name>
        <value>x^2 + dd*x - 5*(5+dd)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>5</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>5+dd</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Used a different area than the slider (width 5, but area for 6)</description>
      <testinput>
        <name>ans0</name>
        <value>5</value>
      </testinput>
      <testinput>
        <name>ans1</name>
        <value>x^2 + dd*x - 6*(6+dd)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>5</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>5+dd</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-wrongarea</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Part (a) wrong structure (wrong middle coefficient)</description>
      <testinput>
        <name>ans0</name>
        <value>5</value>
      </testinput>
      <testinput>
        <name>ans1</name>
        <value>x^2 + (dd+1)*x - 5*(5+dd)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>5</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>5+dd</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-other</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Width: gave the length</description>
      <testinput>
        <name>ans0</name>
        <value>5</value>
      </testinput>
      <testinput>
        <name>ans1</name>
        <value>x^2 + dd*x - 5*(5+dd)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>5+dd</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>5+dd</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-length</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Width: gave the negative root</description>
      <testinput>
        <name>ans0</name>
        <value>5</value>
      </testinput>
      <testinput>
        <name>ans1</name>
        <value>x^2 + dd*x - 5*(5+dd)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>-(5+dd)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>5+dd</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-neg</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description>Length: gave the width</description>
      <testinput>
        <name>ans0</name>
        <value>5</value>
      </testinput>
      <testinput>
        <name>ans1</name>
        <value>x^2 + dd*x - 5*(5+dd)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>5</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>5</value>
      </testinput>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-width</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
