<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Finding the number of bags for zero profit</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>A small business sells handmade bags. Its weekly profit, \(P\), in Kenya shillings, is modelled by \[ P = -2x^2+2{@(r1+r2)@}x - 2{@r1 * r2@}, \] where \(x\) is the number of bags sold.</p>
<p><strong>a)</strong> Determine the number of bags that when the business sells it makes <em>neither a profit nor a loss</em>.</p>
<div>\( x = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>
<div>\( x = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The profit is zero when \(P=0\):</p>
\[
\begin{align*}
{@Pexpr@} &= 0\\
x^2 - {@r1+r2@}x + {@r1*r2@} &= 0 &&\text{(divide by } -2\text{)}\\
(x-{@r1@})(x-{@r2@}) &= 0
\end{align*}
\]
<p>so \(x={@r1@}\) or \(x={@r2@}\). </p>
<p>The business makes no profit and no loss when it sells \({@r1@}\) bags or \({@r2@}\) bags.</p>
<p> (Between these two values the profit is positive; outside them it is negative.)</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>r1 : 5*(rand(5)+2);              /* first root, 10..30 */
r2 : r1 + 5*(rand(4)+1);         /* second root, r1 + (5..20) */
Pexpr : -2*x^2 + 2*(r1+r2)*x - 2*r1*r2;   /* = -2(x-r1)(x-r2) */
rs : {r1, r2};</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( P={@Pexpr@}\); zero profit at \(x={@r1@}\) and \(x={@r2@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>r1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>r2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>ok1 : is(elementp(ans1, {r1, r2}));</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Is the first box one of the two solutions?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ok1</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>This value does not give zero profit. Set \(P=0\) and solve \({@Pexpr@}=0\) for \(x\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>ok2 : is(elementp(ans2, {r1, r2})) and is(ans2 # ans1);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Is the second box the OTHER solution (different from box 1)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ok2</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Did they repeat the first box's value?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ans1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-same</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The two boxes should hold the two <em>different</em> solutions — you repeated the value from the first box.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>This is not the other solution. Solving \({@Pexpr@}=0\) gives two values of \(x\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>10001</deployedseed>
    <deployedseed>10002</deployedseed>
    <deployedseed>10003</deployedseed>
    <deployedseed>10004</deployedseed>
    <deployedseed>10005</deployedseed>
    <deployedseed>10006</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Both correct, in order</description>
      <testinput>
        <name>ans1</name>
        <value>r1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>r2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Both correct, reversed order</description>
      <testinput>
        <name>ans1</name>
        <value>r2</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>r1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Repeated the same value</description>
      <testinput>
        <name>ans1</name>
        <value>r1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>r1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-same</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>First box wrong, second correct</description>
      <testinput>
        <name>ans1</name>
        <value>r1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>r2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Second box a wrong (non-root) value</description>
      <testinput>
        <name>ans1</name>
        <value>r1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>r2+1</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-wrong</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
