<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>GeoGebra Excircle Construction (opposite B) (INNODEMS)</text>
    </name>
    <questiontext format="plain_text">
      <text><![CDATA[<p>The triangle \(ABC\) below has vertices
\(A={@ntupleify(A)@}\), \(B={@ntupleify(B)@}\) and \(C={@ntupleify(C)@}\).</p>

<p>Using the construction tools provided, construct the <strong>escribed circle opposite \(B\)</strong> (the \(B\)-excircle).</p>

<p>GeoGebra displays the coordinates of the points you construct. Read off the position of the \(B\)-excentre \(J_B\) and the radius of the excircle (the exradius \(r_B\)), and enter them below, each rounded to <strong>3 decimal places</strong>. </p>

[[geogebra set="A__fixed,B__fixed,C__fixed" width="1200px" height="600px"]]
params["material_id"]="vdqtbktj";
/* Restricted compass-and-straightedge toolbar. customToolBar OVERRIDES the file's
   toolbar; numbers are GeoGebra tool IDs, | starts a new menu.
   0 Move, 6 Delete | 5 Intersect | 2 Line | 10 Circle(centre,point) |
   9 Angle-bisector, 4 Perpendicular line | 38 Distance or Length.
   Provides 9 Angle-bisector and 4 Perpendicular line (perpendicular to the
   internal bisector gives the EXTERNAL bisector needed for an excentre), plus
   38 to measure the radius. */
/* params["appName"]="geometry"; */
params["showToolBar"]=true;
params["customToolBar"]="0 | 6 | 1 |  5 | 2 | 10 | 9 | 4 | 8 | 38";
params["showToolBarHelp"]=false;
/* Hide the algebra input bar and menu bar. NOTE: the algebra VIEW panel is
   governed by the saved file (perspective is ignored with material_id), so ALSO
   close the Algebra View in the GeoGebra material and re-save. */
params["showAlgebraInput"]=true;
params["algebraInputPosition"]="bottom";
params["showMenuBar"]=false;
params["showResetIcon"]=true;
params["showFullscreenButton"]=true;
[[/geogebra]]

<div>\(B\)-excentre \(J_B = (\,\) [[input:xc]] \(,\) [[input:yc]] \(\,)\) \( (\,\)[[validation:xc]], [[validation:yc]] \(\,)\)</div>

<div>[[feedback:prt1]]</div>

<div>Exradius \(r_B = \) [[input:r]] [[validation:r]]</div>

<div>[[feedback:prt2]]</div>]]></text>
    </questiontext>
    <generalfeedback format="plain_text">
      <text><![CDATA[<p>The centre of the escribed circle opposite \(B\) is the <em>\(B\)-excentre</em>: the point where the <strong>external</strong> bisectors of angles \(A\) and \(C\) meet (it also lies on the <strong>internal</strong> bisector of angle \(B\)). It is equidistant from line \(AC\) and from the extensions of the other two sides, so the excircle is tangent to side \(AC\) and to the extensions of \(BA\) and \(BC\). That common distance is the exradius \(r_B\).</p>

<p><strong>Construction.</strong> Extend sides \(BC\) and \(BA\) and construct the external angle bisectors at \(B\) and \(C\). Their intersection \(J_B\) is the \(B\)-excentre. Dropping a perpendicular from \(J\) to extended side \(AC\) to the point of tangency) gives the exradius.</p>

[[geogebra set="A__fixed,B__fixed,C__fixed" width="1200px" height="600px"]]
params["material_id"]="xpks6xhd";
/* Restricted compass-and-straightedge toolbar. customToolBar OVERRIDES the file's
   toolbar; numbers are GeoGebra tool IDs, | starts a new menu.
   0 Move, 6 Delete | 5 Intersect | 2 Line | 10 Circle(centre,point) |
   9 Angle-bisector, 4 Perpendicular line | 38 Distance or Length.
   Provides 9 Angle-bisector and 4 Perpendicular line (perpendicular to the
   internal bisector gives the EXTERNAL bisector needed for an excentre), plus
   38 to measure the radius. */
/* params["appName"]="geometry"; */
params["showToolBar"]=true;
params["customToolBar"]="0 | 6 | 1 |  5 | 2 | 10 | 9 | 4 | 8 | 38";
params["showToolBarHelp"]=false;
/* Hide the algebra input bar and menu bar. NOTE: the algebra VIEW panel is
   governed by the saved file (perspective is ignored with material_id), so ALSO
   close the Algebra View in the GeoGebra material and re-save. */
params["showAlgebraInput"]=true;
params["algebraInputPosition"]="bottom";
params["showMenuBar"]=false;
params["showResetIcon"]=true;
params["showFullscreenButton"]=true;
[[/geogebra]]

<p>For this triangle the \(B\)-excentre is
\(J_B \approx ({@dispdp(ta_EB[1],3)@},\ {@dispdp(ta_EB[2],3)@})\) and the exradius is
\(r_B \approx {@dispdp(ta_RB,3)@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>tris : [ [[0,0],[11,0],[-1,4]], [[0,0],[11,0],[0,5]], [[0,0],[10,0],[-1,5]], [[0,0],[9,0],[-1,4]], [[0,0],[11,0],[2,4]], [[0,0],[11,0],[-1,6]], [[0,0],[10,0],[0,4]], [[0,0],[11,0],[1,6]], [[0,0],[10,0],[2,4]], [[0,0],[11,0],[2,5]] ];
tri : tris[rand(length(tris))+1];
A : tri[1] + [-1,-1];
B : tri[2] + [-1,-1];
C : tri[3] + [-1,-1];
/* side lengths, each opposite its like-named vertex */
la : Distance(B,C);  /* a = |BC|, opposite A */
lb : Distance(A,C);  /* b = |CA|, opposite B */
lc : Distance(A,B);  /* c = |AB|, opposite C */
perim : la+lb+lc;
ssemi : perim/2;
area : abs(A[1]*(B[2]-C[2])+B[1]*(C[2]-A[2])+C[1]*(A[2]-B[2]))/2;
/* B-excentre = (a*A - b*B + c*C)/(a - b + c)  (target answer) */
denB : la-lb+lc;
ta_EB : [(la*A[1]-lb*B[1]+lc*C[1])/denB, (la*A[2]-lb*B[2]+lc*C[2])/denB];
/* exradius opposite B: r_B = area/(s - b) */
ta_RB : area/(ssemi-lb);
/* distractor centres used only for misconception feedback */
ta_X : [(la*A[1]+lb*B[1]+lc*C[1])/perim, (la*A[2]+lb*B[2]+lc*C[2])/perim];   /* incentre: all internal bisectors */
denA : -la+lb+lc;
ta_EA : [(-la*A[1]+lb*B[1]+lc*C[1])/denA, (-la*A[2]+lb*B[2]+lc*C[2])/denA];   /* excentre opposite A */
denC : la+lb-lc;
ta_EC : [(la*A[1]+lb*B[1]-lc*C[1])/denC, (la*A[2]+lb*B[2]-lc*C[2])/denC];   /* excentre opposite C */
/* inradius, used only to detect the wrong radius in feedback */
ta_inr : area/ssemi;
/* circumcentre and circumradius, used only to detect this misconception */
ddet : 2*(A[1]*(B[2]-C[2])+B[1]*(C[2]-A[2])+C[1]*(A[2]-B[2]));
ta_O : [((A[1]^2+A[2]^2)*(B[2]-C[2])+(B[1]^2+B[2]^2)*(C[2]-A[2])+(C[1]^2+C[2]^2)*(A[2]-B[2]))/ddet, ((A[1]^2+A[2]^2)*(C[1]-B[1])+(B[1]^2+B[2]^2)*(A[1]-C[1])+(C[1]^2+C[2]^2)*(B[1]-A[1]))/ddet];
ta_Rc : sqrt((ta_O[1]-A[1])^2+(ta_O[2]-A[2])^2);
/* rounded model answers (to 3 d.p.), used for the inputs and the question tests */
rnd3(z) := float(round(float(z)*1000)/1000);
ans_xc : rnd3(ta_EB[1]);
ans_yc : rnd3(ta_EB[2]);
ans_r  : rnd3(ta_RB);
ans_Xx : rnd3(ta_X[1]);
ans_Xy : rnd3(ta_X[2]);
ans_EAx : rnd3(ta_EA[1]);
ans_EAy : rnd3(ta_EA[2]);
ans_ECx : rnd3(ta_EC[1]);
ans_ECy : rnd3(ta_EC[2]);
ans_ir : rnd3(ta_inr);
ans_Ox : rnd3(ta_O[1]);
ans_Oy : rnd3(ta_O[2]);
ans_Rc : rnd3(ta_Rc);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(A={@ntupleify(A)@},\ B={@ntupleify(B)@},\ C={@ntupleify(C)@},\ J \approx ({@dispdp(ta_EB[1],3)@},\ {@dispdp(ta_EB[2],3)@}),\ r_B \approx {@dispdp(ta_RB,3)@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>r</name>
      <type>numerical</type>
      <tans>ans_r</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>xc</name>
      <type>numerical</type>
      <tans>ans_xc</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>yc</name>
      <type>numerical</type>
      <tans>ans_yc</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* largest single-coordinate error, so x AND y must each be correct to 3 d.p. */
errEB  : max(abs(xc-ta_EB[1]), abs(yc-ta_EB[2]));
errInc : max(abs(xc-ta_X[1]), abs(yc-ta_X[2]));
errEA  : max(abs(xc-ta_EA[1]), abs(yc-ta_EA[2]));
errEC  : max(abs(xc-ta_EC[1]), abs(yc-ta_EC[2]));
errCirc: max(abs(xc-ta_O[1]), abs(yc-ta_O[2]));</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Both B-excentre coordinates correct to 3 d.p.?</description>
        <answertest>NumAbsolute</answertest>
        <sans>errEB</sans>
        <tans>0</tans>
        <testoptions>0.0005</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-excentre-B</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The centre is correctly placed at the \(B\)-excentre.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Did they find the incentre instead?</description>
        <answertest>NumAbsolute</answertest>
        <sans>errInc</sans>
        <tans>0</tans>
        <testoptions>0.05</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-incentre</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That point is the <em>incentre</em> &mdash; you used the <strong>internal</strong> bisectors at every vertex. For an excentre you need the <strong>external</strong> bisectors at two of the vertices. The \(B\)-excentre uses the external bisectors of \(A\) and \(C\) (each perpendicular to the internal bisector there).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Did they find the excentre opposite A instead?</description>
        <answertest>NumAbsolute</answertest>
        <sans>errEA</sans>
        <tans>0</tans>
        <testoptions>0.05</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-excentre-A</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the excentre opposite \(A\) &mdash; the centre of the excircle tangent to side \(BC\). You took external bisectors at the wrong pair of vertices. For the \(B\)-excircle (tangent to side \(AC\)) use the external bisectors at \(A\) and \(C\), with the internal bisector at \(B\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Did they find the excentre opposite C instead?</description>
        <answertest>NumAbsolute</answertest>
        <sans>errEC</sans>
        <tans>0</tans>
        <testoptions>0.05</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-excentre-C</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the excentre opposite \(C\) &mdash; the centre of the excircle tangent to side \(AB\). You took external bisectors at the wrong pair of vertices. For the \(B\)-excircle (tangent to side \(AC\)) use the external bisectors at \(A\) and \(C\), with the internal bisector at \(B\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>4</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>4</name>
        <description>Did they find the circumcentre instead?</description>
        <answertest>NumAbsolute</answertest>
        <sans>errCirc</sans>
        <tans>0</tans>
        <testoptions>0.05</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-circumcentre</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That point is the <em>circumcentre</em> &mdash; where the perpendicular bisectors of the <strong>sides</strong> meet (equidistant from the three vertices). The \(B\)-excentre instead comes from the <strong>external angle bisectors</strong> at \(A\) and \(C\), and lies outside the triangle on the far side of \(AC\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your point \(J\) is not the \(B\)-excentre. It is the meeting point of the external bisectors of angles \(A\) and \(C\) &mdash; check that each external bisector is perpendicular to the internal bisector at that vertex, and that it lies outside the triangle on the far side of \(AC\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>errInr : abs(r-ta_inr);
errRc  : abs(r-ta_Rc);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Exradius correct to 3 d.p.?</description>
        <answertest>NumAbsolute</answertest>
        <sans>r</sans>
        <tans>ta_RB</tans>
        <testoptions>0.0005</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-radius-ok</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The exradius is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Did they give the inradius instead of the exradius?</description>
        <answertest>NumAbsolute</answertest>
        <sans>errInr</sans>
        <tans>0</tans>
        <testoptions>0.005</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-inradius</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the <em>inradius</em> &mdash; the radius of the <strong>inscribed</strong> circle inside the triangle. Here you want the <strong>exradius</strong> \(r_B\) of the escribed circle opposite \(B\), which is larger: it is the distance from the \(B\)-excentre to side \(AC\), equal to \(\dfrac{\text{area}}{s-b}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Did they give the circumradius instead of the exradius?</description>
        <answertest>NumAbsolute</answertest>
        <sans>errRc</sans>
        <tans>0</tans>
        <testoptions>0.005</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-circumradius</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the <em>circumradius</em> &mdash; the radius of the circle through the three vertices. Here you want the <strong>exradius</strong> \(r_B\) of the escribed circle opposite \(B\): the distance from the \(B\)-excentre to side \(AC\), equal to \(\dfrac{\text{area}}{s-b}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-radius-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your exradius is not correct (to 3 decimal places). The exradius \(r_B\) is the distance from the \(B\)-excentre to side \(AC\) (the perpendicular distance to the line it is tangent to), equal to \(\dfrac{\text{area}}{s-b}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1001</deployedseed>
    <deployedseed>1002</deployedseed>
    <deployedseed>1003</deployedseed>
    <deployedseed>1004</deployedseed>
    <deployedseed>1005</deployedseed>
    <deployedseed>1006</deployedseed>
    <deployedseed>1007</deployedseed>
    <deployedseed>1008</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct: B-excentre and exradius both correct</description>
      <testinput>
        <name>r</name>
        <value>ans_r</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_xc</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_yc</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-excentre-B</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-radius-ok</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Centre correct but exradius wrong</description>
      <testinput>
        <name>r</name>
        <value>ans_r+5</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_xc</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_yc</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-excentre-B</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-F | prt2-1-F | prt2-radius-no</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Used the incentre instead (all internal bisectors)</description>
      <testinput>
        <name>r</name>
        <value>ans_r+5</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_Xx</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_Xy</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-F | prt1-incentre</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-F | prt2-1-F | prt2-radius-no</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Used the excentre opposite A instead</description>
      <testinput>
        <name>r</name>
        <value>ans_r+5</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_EAx</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_EAy</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-F | prt1-1-F | prt1-excentre-A</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-F | prt2-1-F | prt2-radius-no</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Used the excentre opposite C instead</description>
      <testinput>
        <name>r</name>
        <value>ans_r+5</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_ECx</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_ECy</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-F | prt1-1-F | prt1-2-F | prt1-excentre-C</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-F | prt2-1-F | prt2-radius-no</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description>Neither a recognised centre nor exradius</description>
      <testinput>
        <name>r</name>
        <value>ans_r+5</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>100</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>100</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-F | prt1-1-F | prt1-2-F | prt1-3-F | prt1-other</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-F | prt2-1-F | prt2-radius-no</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>7</testcase>
      <description>Centre correct but gave the inradius instead of the exradius</description>
      <testinput>
        <name>r</name>
        <value>ans_ir</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_xc</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_yc</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-excentre-B</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-F | prt2-inradius</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>8</testcase>
      <description>Used the circumcentre and circumradius instead</description>
      <testinput>
        <name>r</name>
        <value>ans_Rc</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_Ox</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_Oy</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-F | prt1-1-F | prt1-2-F | prt1-3-F | prt1-circumcentre</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-0-F | prt2-1-F | prt2-circumradius</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
