<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Determinant Counterexamples (INNODEMS)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Throughout, write \(A=\begin{bmatrix}a&b\\c&d\end{bmatrix}\) and, where needed, \(B=\begin{bmatrix}p&q\\r&s\end{bmatrix}\). For each statement, compute the two sides step by step, then decide whether they are equal for <em>all</em> such matrices.</p>

<p><b>(a)</b> Claim: \(\det(A+B)=\det A+\det B\).</p>
<p>\(A+B=\) [[input:asum]] [[validation:asum]] [[feedback:prtasum]]</p>
<p>\(\det(A+B)=\) [[input:adetsum]] [[validation:adetsum]] [[feedback:prtadetsum]]</p>
<div>\(\det A+\det B=\) [[input:asumdet]] [[validation:asumdet]] [[feedback:prtasumdet]]</div>
<div>Are these two quantities equal for all \(A,B\)? [[input:aeq]] [[validation:aeq]] [[feedback:prtaeq]]</div>
[[reveal input="aeq" value="2"]]
<div><b>Follow-up.</b> Give a counterexample: specific matrices \(A\) and \(B\) with whole-number entries for which \(\det(A+B)\neq\det A+\det B\).</div>
<div>\(A=\) [[input:ace1]] [[validation:ace1]] \(\qquad B=\) [[input:ace2]] [[validation:ace2]]</div>
<div>[[feedback:prtace]]</div>
<div>The claim fails <em>in general</em>. Can particular choices of \(A\) and \(B\) still make the two sides equal? [[input:asp]] [[validation:asp]]</div>
[[reveal input="asp" value="1"]]
<div>Give such a pair: \(A=\) [[input:asp1]] [[validation:asp1]] \(\qquad B=\) [[input:asp2]] [[validation:asp2]]</div>
[[/reveal]]
<div>[[feedback:prtasp]]</div>
[[/reveal]]

<p><b>(b)</b> Claim: \(\det({@kval@}A)={@kval^2@}\det A\).</p>
<div>\({@kval@}A=\) [[input:bscal]] [[validation:bscal]] [[feedback:prtbscal]]</div>
<div>\(\det({@kval@}A)=\) [[input:bdetscal]] [[validation:bdetscal]] [[feedback:prtbdetscal]]</div>
<div>\({@kval^2@}\det A=\) [[input:brhs]] [[validation:brhs]] [[feedback:prtbrhs]]</div>
<div>Are these two quantities equal for all \(A\)? [[input:beq]] [[validation:beq]] [[feedback:prtbeq]]</div>
[[reveal input="beq" value="2"]]
<div><b>Follow-up.</b> Give a counterexample: a specific matrix \(A\) with whole-number entries for which \(\det({@kval@}A)\neq{@kval^2@}\det A\).</div>
<div>\(A=\) [[input:bce1]] [[validation:bce1]]</div>
<div>[[feedback:prtbce]]</div>
[[/reveal]]

<p><b>(c)</b> Claim: \(\det(A^2)=2\det A\).</p>
<div>\(A^2=\) [[input:csq]] [[validation:csq]] [[feedback:prtcsq]]</div>
<div>\(\det(A^2)=\) [[input:cdetsq]] [[validation:cdetsq]] [[feedback:prtcdetsq]]</div>
<div>\(2\det A=\) [[input:crhs]] [[validation:crhs]] [[feedback:prtcrhs]]</div>
<div>Are these two quantities equal for all \(A\)? [[input:ceq]] [[validation:ceq]] [[feedback:prtceq]]</div>
[[reveal input="ceq" value="2"]]
<div><b>Follow-up.</b> Give a counterexample: a specific matrix \(A\) with whole-number entries for which \(\det(A^2)\neq 2\det A\).</div>
<div>\(A=\) [[input:cce1]] [[validation:cce1]]</div>
<div>[[feedback:prtcce]]</div>
<div>The claim fails <em>in general</em>. Can a particular matrix \(A\) still make the two sides equal? [[input:csp]] [[validation:csp]]</div>
[[reveal input="csp" value="1"]]
<div>Give such a matrix: \(A=\) [[input:csp1]] [[validation:csp1]]</div>
[[/reveal]]
<div>[[feedback:prtcsp]]</div>
[[/reveal]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p><b>(a) Not equal &mdash; the claim is false.</b> \(A+B=\begin{bmatrix}a+p&b+q\\c+r&d+s\end{bmatrix}\), so
\[\det(A+B)=(a+p)(d+s)-(b+q)(c+r)=ad+as+dp+ps-bc-br-cq-qr.\]
But \(\det A+\det B=(ad-bc)+(ps-qr)\). These differ by the cross terms \(as+dp-br-cq\), so they are not equal in general. In words: a determinant is found by <em>multiplying</em> entries together, and multiplying mixes the entries of \(A\) and \(B\) &mdash; so the determinant of a sum is not the sum of the determinants. For a counterexample take \(A=\begin{bmatrix}1&0\\0&1\end{bmatrix}\) and \(B=\begin{bmatrix}2&0\\0&2\end{bmatrix}\): then \(A+B=\begin{bmatrix}3&0\\0&3\end{bmatrix}\), so \(\det(A+B)=9\), but \(\det A+\det B=1+4=5\). The two sides <em>can</em> agree for special choices &mdash; whenever the cross terms vanish, e.g. \(B=\begin{bmatrix}0&0\\0&0\end{bmatrix}\), where both sides equal \(\det A\). So a claim can fail in general yet hold in special cases.</p>

<p><b>(b) Equal &mdash; the claim is true.</b> \({@kval@}A={@b_scal_ta@}\), so
\[\det({@kval@}A)=({@kval@}a)({@kval@}d)-({@kval@}b)({@kval@}c)={@kval^2@}(ad-bc)={@kval^2@}\det A.\]
Each of the two rows carries a factor \({@kval@}\), and each row contributes its own factor to the determinant, giving \({@kval@}\times{@kval@}={@kval^2@}\) &mdash; not \({@kval@}\). The same reasoning shows \(\det(kA)=k^2\det A\) for any scalar \(k\) and \(2\times2\) matrix \(A\)</p>

<p><b>(c) Not equal &mdash; the claim is false.</b> \(A^2=\begin{bmatrix}a^2+bc&ab+bd\\ac+cd&bc+d^2\end{bmatrix}\), so
\[\det(A^2)=(a^2+bc)(bc+d^2)-(ab+bd)(ac+cd)=(ad-bc)^2.\]
But \(2\det A=2(ad-bc)\). Since \((ad-bc)^2\neq 2(ad-bc)\) in general, they are not equal. (In fact \(\det(A^2)=(\det A)^2\), the correct identity.) Any \(A\) whose determinant is neither \(0\) nor \(2\) is a counterexample, e.g. \(A=\begin{bmatrix}1&0\\0&1\end{bmatrix}\) gives \(A^2=\begin{bmatrix}1&0\\0&1\end{bmatrix}\), \(\det(A^2)=1\) but \(2\det A=2\). The two sides agree exactly when \((\det A)^2=2\det A\), i.e. \(\det A(\det A-2)=0\): matrices with \(\det A=0\) or \(\det A=2\), such as \(A=\begin{bmatrix}2&0\\0&1\end{bmatrix}\), are special cases where the false general claim happens to hold.</p>]]></text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Part (a): det(A+B) vs det A + det B  (NOT equal) */
a_sum_ta    : matrix([a+p,b+q],[c+r,d+s]);
a_detsum_ta : (a+p)*(d+s)-(b+q)*(c+r);
a_sumdet_ta : (a*d-b*c)+(p*s-q*r);
a_eq_ta     : [[true,false,"Yes: equal for all such matrices"],[false,true,"No: not equal in general"]];

/* Part (b): det(kval*A) vs kval^2 det A  (equal). kval randomised;
   2..6 avoids the degenerate scalars 0 and 1. */
kval : rand([2,3,4,5,6]);
b_scal_ta    : matrix([kval*a,kval*b],[kval*c,kval*d]);
b_detscal_ta : (kval*a)*(kval*d)-(kval*b)*(kval*c);
b_rhs_ta     : kval^2*(a*d-b*c);
b_eq_ta      : [[true,true,"Yes: equal for all such matrices"],[false,false,"No: not equal in general"]];

/* Part (c): det(A^2) vs 2 det A  (NOT equal) */
c_sq_ta    : matrix([a^2+b*c, a*b+b*d],[a*c+c*d, b*c+d^2]);
c_detsq_ta : (a*d-b*c)^2;
c_rhs_ta   : 2*(a*d-b*c);
c_eq_ta    : [[true,false,"Yes: equal for all such matrices"],[false,true,"No: not equal in general"]];

/* Follow-up boxes on the false parts (a),(c), revealed on "No: not equal in general".
   Marking is by property, so these are only sample instances. */
a_ce1_ta : matrix([1,0],[0,1]);
a_ce2_ta : matrix([2,0],[0,2]);
a_sp_ta  : [[true,true,"Yes, for some special choices"],[false,false,"No, the two sides are never equal"]];
a_sp1_ta : matrix([1,0],[0,1]);
a_sp2_ta : matrix([0,0],[0,0]);
/* (b) trap: the claim is true, so no counterexample exists; dummy model answer. */
bce1_ta  : matrix([1,0],[0,1]);
c_ce1_ta : matrix([1,0],[0,1]);
c_sp_ta  : [[true,true,"Yes, for some special choices"],[false,false,"No, the two sides are never equal"]];
c_sp1_ta : matrix([2,0],[0,1]);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>Determinant properties by computation: (a) det(A+B) [not equal], (b) det({@kval@}A)={@kval^2@} detA [equal], (c) det(A^2)=2 detA [not equal].</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ace1</name>
      <type>matrix</type>
      <tans>a_ce1_ta</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options>allowempty</options>
    </input>
    <input>
      <name>ace2</name>
      <type>matrix</type>
      <tans>a_ce2_ta</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options>allowempty</options>
    </input>
    <input>
      <name>adetsum</name>
      <type>algebraic</type>
      <tans>a_detsum_ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>aeq</name>
      <type>dropdown</type>
      <tans>a_eq_ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>asp</name>
      <type>dropdown</type>
      <tans>a_sp_ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>asp1</name>
      <type>matrix</type>
      <tans>a_sp1_ta</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options>allowempty</options>
    </input>
    <input>
      <name>asp2</name>
      <type>matrix</type>
      <tans>a_sp2_ta</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options>allowempty</options>
    </input>
    <input>
      <name>asum</name>
      <type>matrix</type>
      <tans>a_sum_ta</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>asumdet</name>
      <type>algebraic</type>
      <tans>a_sumdet_ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>bce1</name>
      <type>matrix</type>
      <tans>bce1_ta</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options>allowempty</options>
    </input>
    <input>
      <name>bdetscal</name>
      <type>algebraic</type>
      <tans>b_detscal_ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>beq</name>
      <type>dropdown</type>
      <tans>b_eq_ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>brhs</name>
      <type>algebraic</type>
      <tans>b_rhs_ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>bscal</name>
      <type>matrix</type>
      <tans>b_scal_ta</tans>
      <boxsize>3</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>cce1</name>
      <type>matrix</type>
      <tans>c_ce1_ta</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options>allowempty</options>
    </input>
    <input>
      <name>cdetsq</name>
      <type>algebraic</type>
      <tans>c_detsq_ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ceq</name>
      <type>dropdown</type>
      <tans>c_eq_ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>crhs</name>
      <type>algebraic</type>
      <tans>c_rhs_ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>csp</name>
      <type>dropdown</type>
      <tans>c_sp_ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>csp1</name>
      <type>matrix</type>
      <tans>c_sp1_ta</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options>allowempty</options>
    </input>
    <input>
      <name>csq</name>
      <type>matrix</type>
      <tans>c_sq_ta</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prtace</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>filled_ace : matrixp(ace1) and matrixp(ace2);
num_ace : if filled_ace then every(numberp, append(flatten(args(ace1)), flatten(args(ace2)))) else false;
gap_ace : if num_ace then determinant(ace1+ace2) - (determinant(ace1)+determinant(ace2)) else 1;
/* guarded copies for display; zero matrices when empty/non-numeric */
mA_ace : if num_ace then ace1 else matrix([0,0],[0,0]);
mB_ace : if num_ace then ace2 else matrix([0,0],[0,0]);
sum_ace : mA_ace + mB_ace;
s11 : sum_ace[1,1]; s12 : sum_ace[1,2]; s21 : sum_ace[2,1]; s22 : sum_ace[2,2];
p11 : mA_ace[1,1]; p12 : mA_ace[1,2]; p21 : mA_ace[2,1]; p22 : mA_ace[2,2];
q11 : mB_ace[1,1]; q12 : mB_ace[1,2]; q21 : mB_ace[2,1]; q22 : mB_ace[2,2];
lhs_ace : determinant(sum_ace);
d1_ace  : determinant(mA_ace);
d2_ace  : determinant(mB_ace);
rhs_ace : d1_ace + d2_ace;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Both matrices filled with specific numbers?</description>
        <answertest>AlgEquiv</answertest>
        <sans>num_ace</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtace-numok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtace-notnumeric</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Enter a specific matrix for both \(A\) and \(B\), every entry a whole number.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Do their matrices break the identity?</description>
        <answertest>AlgEquiv</answertest>
        <sans>gap_ace</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtace-satisfies</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Let's check your pair. \(A+B={@sum_ace@}\), so \[\det(A+B)=({@s11@})({@s22@})-({@s12@})({@s21@})={@lhs_ace@}.\] Also \[\det A=({@p11@})({@p22@})-({@p12@})({@p21@})={@d1_ace@},\qquad \det B=({@q11@})({@q22@})-({@q12@})({@q21@})={@d2_ace@},\] so \(\det A+\det B={@d1_ace@}+{@d2_ace@}={@rhs_ace@}\). Both sides equal \({@lhs_ace@}\), so this pair is <em>not</em> a counterexample. Try a different pair.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtace-correct</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Correct. \(A+B={@sum_ace@}\), so \[\det(A+B)=({@s11@})({@s22@})-({@s12@})({@s21@})={@lhs_ace@}.\] But \[\det A=({@p11@})({@p22@})-({@p12@})({@p21@})={@d1_ace@},\qquad \det B=({@q11@})({@q22@})-({@q12@})({@q21@})={@d2_ace@},\] so \(\det A+\det B={@d1_ace@}+{@d2_ace@}={@rhs_ace@}\). Since \({@lhs_ace@}\neq{@rhs_ace@}\), your matrices are a genuine counterexample.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtadetsum</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* determinant of the student's own A+B, for follow-through marking */
ft_adetsum : determinant(asum);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>det(A+B) correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>adetsum</sans>
        <tans>a_detsum_ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtadetsum-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtadetsum-notmodel</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>det of their own A+B? (follow-through)</description>
        <answertest>AlgEquiv</answertest>
        <sans>adetsum</sans>
        <tans>ft_adetsum</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtadetsum-ft</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is the correct determinant of <em>your</em> matrix \(A+B\) in the box above, so you have full marks here. But that matrix is not right &mdash; revisit the box above, then recompute this determinant.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtadetsum-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Apply \(\det\begin{bmatrix}w&x\\y&z\end{bmatrix}=wz-xy\) to your matrix \(A+B\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtaeq</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct equality verdict (not equal)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>aeq</sans>
        <tans>false</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtaeq-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtaeq-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Compare your two expressions term by term. Do all the terms match for every choice of entries?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtasp</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>filled_asp : matrixp(asp1) and matrixp(asp2);
num_asp : if filled_asp then every(numberp, append(flatten(args(asp1)), flatten(args(asp2)))) else false;
gap_asp : if num_asp then determinant(asp1+asp2) - (determinant(asp1)+determinant(asp2)) else 1;
/* guarded copies for display; zero matrices when empty/non-numeric */
mA_asp : if num_asp then asp1 else matrix([0,0],[0,0]);
mB_asp : if num_asp then asp2 else matrix([0,0],[0,0]);
sum_asp : mA_asp + mB_asp;
t11 : sum_asp[1,1]; t12 : sum_asp[1,2]; t21 : sum_asp[2,1]; t22 : sum_asp[2,2];
u11 : mA_asp[1,1]; u12 : mA_asp[1,2]; u21 : mA_asp[2,1]; u22 : mA_asp[2,2];
v11 : mB_asp[1,1]; v12 : mB_asp[1,2]; v21 : mB_asp[2,1]; v22 : mB_asp[2,2];
lhs_asp : determinant(sum_asp);
d1_asp  : determinant(mA_asp);
d2_asp  : determinant(mB_asp);
rhs_asp : d1_asp + d2_asp;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Do they say special cases exist?</description>
        <answertest>AlgEquiv</answertest>
        <sans>asp</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtasp-saysyes</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtasp-saysnever</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Don't rule it out without experimenting. Try making one of the two matrices as simple as you possibly can, and compute both sides.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Both matrices filled with specific numbers?</description>
        <answertest>AlgEquiv</answertest>
        <sans>num_asp</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtasp-numok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtasp-notnumeric</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Enter a specific matrix for both \(A\) and \(B\), every entry a whole number.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Do their matrices make the sides equal?</description>
        <answertest>AlgEquiv</answertest>
        <sans>gap_asp</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtasp-correct</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Indeed. \(A+B={@sum_asp@}\), so \[\det(A+B)=({@t11@})({@t22@})-({@t12@})({@t21@})={@lhs_asp@},\] while \[\det A=({@u11@})({@u22@})-({@u12@})({@u21@})={@d1_asp@},\qquad \det B=({@v11@})({@v22@})-({@v12@})({@v21@})={@d2_asp@},\] so \(\det A+\det B={@rhs_asp@}\). Both sides equal \({@lhs_asp@}\) &mdash; a special pair where the claim holds. A claim can fail in general and yet still hold for special choices; that is exactly why one counterexample disproves a "for all" statement, but one confirming example proves nothing.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtasp-notequal</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Let's check your pair. \(A+B={@sum_asp@}\), so \[\det(A+B)=({@t11@})({@t22@})-({@t12@})({@t21@})={@lhs_asp@},\] while \[\det A=({@u11@})({@u22@})-({@u12@})({@u21@})={@d1_asp@},\qquad \det B=({@v11@})({@v22@})-({@v12@})({@v21@})={@d2_asp@},\] so \(\det A+\det B={@rhs_asp@}\). Since \({@lhs_asp@}\neq{@rhs_asp@}\), this pair is not one of the special cases. Look at your expansion of \(\det(A+B)\) in the boxes above: which extra terms would need to vanish?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtasum</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>A+B correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>asum</sans>
        <tans>a_sum_ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtasum-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtasum-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Add \(A\) and \(B\) entry by entry to get \(A+B\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtasumdet</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>det A + det B correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>asumdet</sans>
        <tans>a_sumdet_ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtasumdet-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtasumdet-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Find \(\det A\) and \(\det B\) separately, then add them.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtbce</name>
      <value>0.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>filled_bce : matrixp(bce1);
num_bce : if filled_bce then every(numberp, flatten(args(bce1))) else false;
/* guarded copies for display; zero matrices when empty/non-numeric */
mA_bce : if num_bce then bce1 else matrix([0,0],[0,0]);
mk_bce : kval*mA_bce;
m11 : mk_bce[1,1]; m12 : mk_bce[1,2]; m21 : mk_bce[2,1]; m22 : mk_bce[2,2];
n11 : mA_bce[1,1]; n12 : mA_bce[1,2]; n21 : mA_bce[2,1]; n22 : mA_bce[2,2];
lhs_bce : determinant(mk_bce);
d_bce   : determinant(mA_bce);
rhs_bce : kval^2*d_bce;
gap_bce : if num_bce then lhs_bce - rhs_bce else 1;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Matrix filled with whole numbers?</description>
        <answertest>AlgEquiv</answertest>
        <sans>num_bce</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtbce-numok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtbce-notnumeric</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Enter a specific matrix \(A\), every entry a whole number.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Their matrix always satisfies the (true) claim</description>
        <answertest>AlgEquiv</answertest>
        <sans>gap_bce</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtbce-satisfies</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Let's check your matrix. \({@kval@}A={@mk_bce@}\), so \[\det({@kval@}A)=({@m11@})({@m22@})-({@m12@})({@m21@})={@lhs_bce@}.\] Also \[\det A=({@n11@})({@n22@})-({@n12@})({@n21@})={@d_bce@},\qquad\text{so}\qquad {@kval^2@}\det A={@kval^2@}\cdot{@d_bce@}={@rhs_bce@}.\] Both sides equal \({@lhs_bce@}\), so your matrix is <em>not</em> a counterexample. Try more matrices&mdash; can you find one that breaks the claim? If not, what does that suggest?</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtbce-check</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Check your entries and try again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtbdetscal</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* determinant of the student's own kval*A, for follow-through marking */
ft_bdetscal : determinant(bscal);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>det(kval*A) correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>bdetscal</sans>
        <tans>b_detscal_ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtbdetscal-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtbdetscal-notmodel</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>det of their own kval*A? (follow-through)</description>
        <answertest>AlgEquiv</answertest>
        <sans>bdetscal</sans>
        <tans>ft_bdetscal</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtbdetscal-ft</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is the correct determinant of <em>your</em> matrix \({@kval@}A\) in the box above, so you have full marks here. But that matrix is not right &mdash; revisit the box above, then recompute this determinant.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtbdetscal-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Apply \(\det\begin{bmatrix}w&x\\y&z\end{bmatrix}=wz-xy\) to your matrix \({@kval@}A\), and simplify each product carefully.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtbeq</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct equality verdict (equal)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>beq</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtbeq-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtbeq-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Compare your two expressions term by term. Do they match for every choice of entries?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtbrhs</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>9 det A correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>brhs</sans>
        <tans>b_rhs_ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtbrhs-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtbrhs-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>This is \({@kval^2@}\) times \(\det A=ad-bc\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtbscal</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>3A correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>bscal</sans>
        <tans>b_scal_ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtbscal-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtbscal-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Multiply every entry of \(A\) by \({@kval@}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtcce</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>filled_cce : matrixp(cce1);
num_cce : if filled_cce then every(numberp, flatten(args(cce1))) else false;
gap_cce : if num_cce then determinant(cce1^^2) - 2*determinant(cce1) else 1;
/* guarded copies for display; zero matrices when empty/non-numeric */
mA_cce : if num_cce then cce1 else matrix([0,0],[0,0]);
sq_cce : mA_cce^^2;
w11 : sq_cce[1,1]; w12 : sq_cce[1,2]; w21 : sq_cce[2,1]; w22 : sq_cce[2,2];
r11 : mA_cce[1,1]; r12 : mA_cce[1,2]; r21 : mA_cce[2,1]; r22 : mA_cce[2,2];
lhs_cce : determinant(sq_cce);
d_cce   : determinant(mA_cce);
rhs_cce : 2*d_cce;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Matrix filled with specific numbers?</description>
        <answertest>AlgEquiv</answertest>
        <sans>num_cce</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtcce-numok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtcce-notnumeric</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Enter a specific matrix \(A\), every entry a whole number.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Does their matrix break the identity?</description>
        <answertest>AlgEquiv</answertest>
        <sans>gap_cce</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtcce-satisfies</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Let's check your matrix. \(A^2={@sq_cce@}\), so \[\det(A^2)=({@w11@})({@w22@})-({@w12@})({@w21@})={@lhs_cce@}.\] Also \[\det A=({@r11@})({@r22@})-({@r12@})({@r21@})={@d_cce@},\qquad\text{so}\qquad 2\det A=2\cdot{@d_cce@}={@rhs_cce@}.\] Both sides equal \({@lhs_cce@}\), so this matrix is <em>not</em> a counterexample. Try a matrix with a different determinant.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtcce-correct</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Correct. \(A^2={@sq_cce@}\), so \[\det(A^2)=({@w11@})({@w22@})-({@w12@})({@w21@})={@lhs_cce@}.\] But \[\det A=({@r11@})({@r22@})-({@r12@})({@r21@})={@d_cce@},\qquad\text{so}\qquad 2\det A=2\cdot{@d_cce@}={@rhs_cce@}.\] Since \({@lhs_cce@}\neq{@rhs_cce@}\), your matrix is a genuine counterexample.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtcdetsq</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* determinant of the student's own A^2, for follow-through marking */
ft_cdetsq : determinant(csq);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>det(A^2) correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>cdetsq</sans>
        <tans>c_detsq_ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtcdetsq-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtcdetsq-notmodel</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>det of their own A^2? (follow-through)</description>
        <answertest>AlgEquiv</answertest>
        <sans>cdetsq</sans>
        <tans>ft_cdetsq</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtcdetsq-ft</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is the correct determinant of <em>your</em> matrix \(A^2\) in the box above, so you have full marks here. But that matrix is not right &mdash; revisit the box above, then recompute this determinant.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtcdetsq-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Apply \(\det\begin{bmatrix}w&x\\y&z\end{bmatrix}=wz-xy\) to your matrix \(A^2\). It should factor neatly.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtceq</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct equality verdict (not equal)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ceq</sans>
        <tans>false</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtceq-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtceq-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Compare your two expressions. Is a square equal to twice the same quantity for all values?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtcrhs</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>2 det A correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>crhs</sans>
        <tans>c_rhs_ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtcrhs-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtcrhs-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>This is \(2\) times \(\det A=ad-bc\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtcsp</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>filled_csp : matrixp(csp1);
num_csp : if filled_csp then every(numberp, flatten(args(csp1))) else false;
gap_csp : if num_csp then determinant(csp1^^2) - 2*determinant(csp1) else 1;
/* guarded copies for display; zero matrices when empty/non-numeric */
mA_csp : if num_csp then csp1 else matrix([0,0],[0,0]);
sq_csp : mA_csp^^2;
y11 : sq_csp[1,1]; y12 : sq_csp[1,2]; y21 : sq_csp[2,1]; y22 : sq_csp[2,2];
z11 : mA_csp[1,1]; z12 : mA_csp[1,2]; z21 : mA_csp[2,1]; z22 : mA_csp[2,2];
lhs_csp : determinant(sq_csp);
d_csp   : determinant(mA_csp);
rhs_csp : 2*d_csp;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Do they say special cases exist?</description>
        <answertest>AlgEquiv</answertest>
        <sans>csp</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prtcsp-saysyes</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtcsp-saysnever</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Look back at your two boxes: both sides depend only on \(\det A\). Treat \(\det A\) as an unknown and ask for which values the two sides agree.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Matrix filled with specific numbers?</description>
        <answertest>AlgEquiv</answertest>
        <sans>num_csp</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prtcsp-numok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtcsp-notnumeric</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Enter a specific matrix \(A\), every entry a whole number.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Does their matrix make the sides equal?</description>
        <answertest>AlgEquiv</answertest>
        <sans>gap_csp</sans>
        <tans>0</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtcsp-correct</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Indeed. \(A^2={@sq_csp@}\), so \[\det(A^2)=({@y11@})({@y22@})-({@y12@})({@y21@})={@lhs_csp@},\] while \[\det A=({@z11@})({@z22@})-({@z12@})({@z21@})={@d_csp@},\qquad\text{so}\qquad 2\det A=2\cdot{@d_csp@}={@rhs_csp@}.\] Both sides equal \({@lhs_csp@}\): your matrix has \(\det A={@d_csp@}\), and \(({@d_csp@})^2=2\cdot{@d_csp@}\). The special cases are exactly the matrices whose determinant \(x\) satisfies \(x^2=2x\) &mdash; a claim can fail in general and yet hold for special choices.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtcsp-notequal</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Let's check your matrix. \(A^2={@sq_csp@}\), so \[\det(A^2)=({@y11@})({@y22@})-({@y12@})({@y21@})={@lhs_csp@},\] while \[\det A=({@z11@})({@z22@})-({@z12@})({@z21@})={@d_csp@},\qquad\text{so}\qquad 2\det A=2\cdot{@d_csp@}={@rhs_csp@}.\] Since \({@lhs_csp@}\neq{@rhs_csp@}\), this is not one of the special cases. For which values of \(x\) does \(x^2=2x\) hold?</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtcsq</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>A^2 correct?</description>
        <answertest>AlgEquiv</answertest>
        <sans>csq</sans>
        <tans>c_sq_ta</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtcsq-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtcsq-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Compute \(A^2=A\cdot A\) by matrix multiplication (not by squaring each entry).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1001</deployedseed>
    <deployedseed>2002</deployedseed>
    <deployedseed>3003</deployedseed>
    <deployedseed>4004</deployedseed>
    <deployedseed>5005</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All boxes correct across all three parts</description>
      <testinput>
        <name>ace1</name>
        <value>matrix([1,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>ace2</name>
        <value>matrix([2,0],[0,2])</value>
      </testinput>
      <testinput>
        <name>adetsum</name>
        <value>(a+p)*(d+s)-(b+q)*(c+r)</value>
      </testinput>
      <testinput>
        <name>aeq</name>
        <value>false</value>
      </testinput>
      <testinput>
        <name>asp</name>
        <value>true</value>
      </testinput>
      <testinput>
        <name>asp1</name>
        <value>matrix([1,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>asp2</name>
        <value>matrix([0,0],[0,0])</value>
      </testinput>
      <testinput>
        <name>asum</name>
        <value>matrix([a+p,b+q],[c+r,d+s])</value>
      </testinput>
      <testinput>
        <name>asumdet</name>
        <value>(a*d-b*c)+(p*s-q*r)</value>
      </testinput>
      <testinput>
        <name>bdetscal</name>
        <value>(kval*a)*(kval*d)-(kval*b)*(kval*c)</value>
      </testinput>
      <testinput>
        <name>beq</name>
        <value>true</value>
      </testinput>
      <testinput>
        <name>brhs</name>
        <value>kval^2*(a*d-b*c)</value>
      </testinput>
      <testinput>
        <name>bscal</name>
        <value>matrix([kval*a,kval*b],[kval*c,kval*d])</value>
      </testinput>
      <testinput>
        <name>cce1</name>
        <value>matrix([1,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>cdetsq</name>
        <value>(a*d-b*c)^2</value>
      </testinput>
      <testinput>
        <name>ceq</name>
        <value>false</value>
      </testinput>
      <testinput>
        <name>crhs</name>
        <value>2*(a*d-b*c)</value>
      </testinput>
      <testinput>
        <name>csp</name>
        <value>true</value>
      </testinput>
      <testinput>
        <name>csp1</name>
        <value>matrix([2,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>csq</name>
        <value>matrix([a^2+b*c,a*b+b*d],[a*c+c*d,b*c+d^2])</value>
      </testinput>
      <expected>
        <name>prtace</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtace-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtadetsum</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtadetsum-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtaeq</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtaeq-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtasp</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtasp-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtasum</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtasum-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtasumdet</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtasumdet-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtbdetscal</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbdetscal-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtbeq</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbeq-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtbrhs</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbrhs-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtbscal</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbscal-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtcce</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcce-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtcdetsq</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcdetsq-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtceq</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtceq-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtcrhs</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcrhs-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtcsp</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcsp-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtcsq</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcsq-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description><![CDATA[(a) equality wrongly says "yes equal"; det(A+B) expanded form accepted]]></description>
      <testinput>
        <name>adetsum</name>
        <value>a*d+a*s+d*p+p*s-b*c-b*r-c*q-q*r</value>
      </testinput>
      <testinput>
        <name>aeq</name>
        <value>true</value>
      </testinput>
      <testinput>
        <name>asum</name>
        <value>matrix([a+p,b+q],[c+r,d+s])</value>
      </testinput>
      <expected>
        <name>prtadetsum</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtadetsum-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtaeq</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtaeq-wrong</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description><![CDATA[(b) equality wrongly says "no"; wrong det(kval*A) (used kval instead of kval^2); attempted counterexample shown to satisfy the claim]]></description>
      <testinput>
        <name>bce1</name>
        <value>matrix([1,2],[3,4])</value>
      </testinput>
      <testinput>
        <name>bdetscal</name>
        <value>kval*(a*d-b*c)</value>
      </testinput>
      <testinput>
        <name>beq</name>
        <value>false</value>
      </testinput>
      <testinput>
        <name>bscal</name>
        <value>matrix([kval*a,kval*b],[kval*c,kval*d])</value>
      </testinput>
      <expected>
        <name>prtbce</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbce-satisfies</expectedanswernote>
      </expected>
      <expected>
        <name>prtbdetscal</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbdetscal-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prtbeq</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbeq-wrong</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>(c) A^2 done by squaring entries (wrong); det(A^2) as (ad-bc)^2 accepted</description>
      <testinput>
        <name>cdetsq</name>
        <value>a^2*d^2-2*a*b*c*d+b^2*c^2</value>
      </testinput>
      <testinput>
        <name>ceq</name>
        <value>false</value>
      </testinput>
      <testinput>
        <name>csq</name>
        <value>matrix([a^2,b^2],[c^2,d^2])</value>
      </testinput>
      <expected>
        <name>prtcdetsq</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcdetsq-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtceq</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtceq-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtcsq</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcsq-wrong</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>(a) follow-up: pair satisfies the identity (not a counterexample); special-case pair that is not equal</description>
      <testinput>
        <name>ace1</name>
        <value>matrix([1,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>ace2</name>
        <value>matrix([0,0],[0,0])</value>
      </testinput>
      <testinput>
        <name>asp</name>
        <value>true</value>
      </testinput>
      <testinput>
        <name>asp1</name>
        <value>matrix([1,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>asp2</name>
        <value>matrix([2,0],[0,2])</value>
      </testinput>
      <expected>
        <name>prtace</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtace-satisfies</expectedanswernote>
      </expected>
      <expected>
        <name>prtasp</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtasp-notequal</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description>(a) follow-up: non-numeric counterexample entry; special case wrongly declared impossible</description>
      <testinput>
        <name>ace1</name>
        <value>matrix([x,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>ace2</name>
        <value>matrix([2,0],[0,2])</value>
      </testinput>
      <testinput>
        <name>asp</name>
        <value>false</value>
      </testinput>
      <expected>
        <name>prtace</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtace-notnumeric</expectedanswernote>
      </expected>
      <expected>
        <name>prtasp</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtasp-saysnever</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>7</testcase>
      <description>(c) follow-up: det=2 matrix satisfies the identity (not a counterexample); special-case matrix with det=1 not equal</description>
      <testinput>
        <name>cce1</name>
        <value>matrix([2,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>csp</name>
        <value>true</value>
      </testinput>
      <testinput>
        <name>csp1</name>
        <value>matrix([1,0],[0,1])</value>
      </testinput>
      <expected>
        <name>prtcce</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcce-satisfies</expectedanswernote>
      </expected>
      <expected>
        <name>prtcsp</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcsp-notequal</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>8</testcase>
      <description>(a) wrong computation boxes: entries multiplied not added; sign error in det; determinants multiplied not added</description>
      <testinput>
        <name>adetsum</name>
        <value>(a+p)*(d+s)+(b+q)*(c+r)</value>
      </testinput>
      <testinput>
        <name>asum</name>
        <value>matrix([a*p,b*q],[c*r,d*s])</value>
      </testinput>
      <testinput>
        <name>asumdet</name>
        <value>(a*d-b*c)*(p*s-q*r)</value>
      </testinput>
      <expected>
        <name>prtadetsum</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtadetsum-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prtasum</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtasum-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prtasumdet</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtasumdet-wrong</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>9</testcase>
      <description>(b) wrong computation boxes: only diagonal entries scaled; RHS uses kval instead of kval^2; non-numeric counterexample attempt</description>
      <testinput>
        <name>bce1</name>
        <value>matrix([z,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>brhs</name>
        <value>kval*(a*d-b*c)</value>
      </testinput>
      <testinput>
        <name>bscal</name>
        <value>matrix([kval*a,b],[c,kval*d])</value>
      </testinput>
      <expected>
        <name>prtbce</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbce-notnumeric</expectedanswernote>
      </expected>
      <expected>
        <name>prtbrhs</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbrhs-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prtbscal</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbscal-wrong</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>10</testcase>
      <description><![CDATA[(c) wrong det(A^2) and RHS; verdict wrongly "yes"; special-case matrix non-numeric]]></description>
      <testinput>
        <name>cdetsq</name>
        <value>a^2*d^2-b^2*c^2</value>
      </testinput>
      <testinput>
        <name>ceq</name>
        <value>true</value>
      </testinput>
      <testinput>
        <name>crhs</name>
        <value>2*a*d-b*c</value>
      </testinput>
      <testinput>
        <name>csp</name>
        <value>true</value>
      </testinput>
      <testinput>
        <name>csp1</name>
        <value>matrix([y,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>csq</name>
        <value>matrix([a^2+b*c,a*b+b*d],[a*c+c*d,b*c+d^2])</value>
      </testinput>
      <expected>
        <name>prtcdetsq</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcdetsq-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prtceq</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtceq-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prtcrhs</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcrhs-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prtcsp</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcsp-notnumeric</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>11</testcase>
      <description>Follow-ups: (a) special-case matrices non-numeric; (c) counterexample non-numeric and special case wrongly declared impossible</description>
      <testinput>
        <name>asp</name>
        <value>true</value>
      </testinput>
      <testinput>
        <name>asp1</name>
        <value>matrix([x,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>asp2</name>
        <value>matrix([1,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>cce1</name>
        <value>matrix([x,0],[0,1])</value>
      </testinput>
      <testinput>
        <name>csp</name>
        <value>false</value>
      </testinput>
      <expected>
        <name>prtasp</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtasp-notnumeric</expectedanswernote>
      </expected>
      <expected>
        <name>prtcce</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcce-notnumeric</expectedanswernote>
      </expected>
      <expected>
        <name>prtcsp</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcsp-saysnever</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>12</testcase>
      <description>Follow-through: wrong matrix in each part, but the determinant of THEIR matrix earns full marks</description>
      <testinput>
        <name>adetsum</name>
        <value>(a*p)*(d*s)-(b*q)*(c*r)</value>
      </testinput>
      <testinput>
        <name>asum</name>
        <value>matrix([a*p,b*q],[c*r,d*s])</value>
      </testinput>
      <testinput>
        <name>bdetscal</name>
        <value>kval^2*a*d-b*c</value>
      </testinput>
      <testinput>
        <name>bscal</name>
        <value>matrix([kval*a,b],[c,kval*d])</value>
      </testinput>
      <testinput>
        <name>cdetsq</name>
        <value>a^2*d^2-b^2*c^2</value>
      </testinput>
      <testinput>
        <name>csq</name>
        <value>matrix([a^2,b^2],[c^2,d^2])</value>
      </testinput>
      <expected>
        <name>prtadetsum</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtadetsum-ft</expectedanswernote>
      </expected>
      <expected>
        <name>prtasum</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtasum-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prtbdetscal</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbdetscal-ft</expectedanswernote>
      </expected>
      <expected>
        <name>prtbscal</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtbscal-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prtcdetsq</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcdetsq-ft</expectedanswernote>
      </expected>
      <expected>
        <name>prtcsq</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtcsq-wrong</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
