<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Opposite angles of a cyclic quadrilateral</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>\(ABCD\) is a cyclic quadrilateral. Angle \(ABC = {@x@}^\circ\) and angle \(DAB = {@y@}^\circ\).</p>

[[jsxgraph width="380px" height="380px"]]
var board = JXG.JSXGraph.initBoard(divid, {boundingbox:[-1.5,1.5,1.5,-1.5], axis:false, showCopyright:false, showNavigation:false, keepAspectRatio:true});
var r = Math.PI/180;
var a0 = 200*r;
var angs = [a0, a0+{#pp#}*r, a0+({#pp#}+{#qq#})*r, a0+({#pp#}+{#qq#}+{#rr#})*r];
var names = ['A','B','C','D'];
var P = [];
for (var i=0; i<4; i++) {
  P.push(board.create('point',[Math.cos(angs[i]),Math.sin(angs[i])],
    {name:names[i], fixed:true, size:2, label:{offset:[14*Math.cos(angs[i])-4, 14*Math.sin(angs[i])]}}));
}
board.create('circle',[[0,0],1],{fixed:true,strokeColor:'#333'});
board.create('polygon',P,{fixed:true,fillColor:'none',borders:{strokeColor:'#333'},vertices:{visible:true}});
[[/jsxgraph]]

<p>(a) Find angle \(ADC\). [[input:ans1]] [[validation:ans1]]</p>
<div>[[feedback:prt1]]</div>

<p>(b) Find angle \(BCD\). [[input:ans2]] [[validation:ans2]]</p>
<div>[[feedback:prt2]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h4>Worked solution</h4>
<p><strong>(a)</strong> Opposite angles of a cyclic quadrilateral add up to \(180^\circ\).
\(ADC = 180 - {@x@} = {@ta1@}^\circ\).</p>
<p><strong>(b)</strong> Angle \(BCD\) is opposite angle \(DAB\), so
\(BCD = 180 - {@y@} = {@ta2@}^\circ\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>x: rand_with_step(70,105,5);
y: rand_with_step(70,95,5);
rr: 60;
pp: 420-2*x-2*y;
qq: 2*y-rr;
ss: 2*x-rr;
ta1: 180-x;
ta2: 180-y;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@},{@ta2@},x = {@x@}, y = {@y@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>779211275</deployedseed>
    <deployedseed>1061353730</deployedseed>
    <deployedseed>1063216654</deployedseed>
    <deployedseed>2121852924</deployedseed>
    <deployedseed>2016149893</deployedseed>
    <deployedseed>122633640</deployedseed>
    <deployedseed>2082450259</deployedseed>
    <deployedseed>1317936243</deployedseed>
    <deployedseed>762313730</deployedseed>
    <deployedseed>649069191</deployedseed>
    <deployedseed>1107834585</deployedseed>
    <deployedseed>1391676507</deployedseed>
    <deployedseed>1596406228</deployedseed>
    <deployedseed>741430779</deployedseed>
    <deployedseed>1491355747</deployedseed>
    <deployedseed>1564764703</deployedseed>
    <deployedseed>1188345891</deployedseed>
    <deployedseed>861658825</deployedseed>
    <deployedseed>755874958</deployedseed>
    <deployedseed>987986495</deployedseed>
  </question>
</quiz>
