<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Graphing an inclusive linear inequality in standard form</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Solve and graph the following linear inequality in two unknowns:</p>
<p style="margin-left: 20px;">
\( \displaystyle {@A*x + B*y@} \le {@C@} \)
</p>

<p><b>a) </b> Find the axes intercepts for the boundary equation \( {@A*x + B*y@} = {@C@} \).</p>
<p><i>(Enter the numerical values for the intercepts below.)</i></p>

<p>\(y\)-intercept (when \(x = 0\)): \( y = \) [[input:ans1]]</p>
<p>\(x\)-intercept (when \(y = 0\)): \( x = \) [[input:ans2]]</p>

<div><b>\(y\)-intercept check:</b> [[validation:ans1]] [[feedback:prt1]]</div>
<div><b>\(x\)-intercept check:</b> [[validation:ans2]] [[feedback:prt2]]</div>

<p><b>b)</b> Graph the inequality on the Cartesian plane below:</p>
<ul>
<li>Drag the <b>two blue points</b> to set the boundary line. <i>(Notice the line is solid because the inequality is inclusive: \(\le\)).</i></li>
<li>Drag the <b>green Region Marker</b> into the area that satisfies the inequality \( {@A*x + B*y@} \le {@C@} \).</li>
</ul>

[[jsxgraph input-ref-ans3="rP1" input-ref-ans4="rP2" input-ref-ans5="rP3" width="600px" height="600px"]]
var brd = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-10.5, 10.5, 10.5, -10.5],
    axis: false, grid: false, showNavigation: false, showCopyright: false, pan: {enabled: false}, zoom: {enabled: false}
});

var xaxis = brd.create('axis', [[0, 0], [1, 0]], {
    name: 'x', withLabel: true, label: {position: 'rt', offset: [-15, 20]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-8, -15]}, drawZero: true, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

var yaxis = brd.create('axis', [[0, 0], [0, 1]], {
    name: 'y', withLabel: true, label: {position: 'rt', offset: [15, -15]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-25, -5]}, drawZero: false, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

/* Create two boundary points starting at default positions */
var p1 = brd.create('point', [0, 5], {
    name: '', size: 5, strokeColor: '#1565c0', fillColor: '#ffffff', strokeWidth: 3, snapToGrid: true, snapSizeX: 1, snapSizeY: 1
});
var p2 = brd.create('point', [-5, 0], {
    name: '', size: 5, strokeColor: '#1565c0', fillColor: '#ffffff', strokeWidth: 3, snapToGrid: true, snapSizeX: 1, snapSizeY: 1
});

/* Draw the boundary line (solid for inclusive inequality) */
var line = brd.create('line', [p1, p2], {
    strokeWidth: 3, strokeColor: '#1565c0', dash: 0, highlight: false
});

/* Create the Region Marker */
var p3 = brd.create('point', [4, -4], {
    name: 'Region', size: 6, strokeColor: '#2e7d32', fillColor: '#4caf50', strokeWidth: 2, snapToGrid: true, snapSizeX: 1, snapSizeY: 1
});

stack_jxg.bind_point(rP1, p1);
stack_jxg.bind_point(rP2, p2);
stack_jxg.bind_point(rP3, p3);
stack_jxg.starts_moved(p1);
stack_jxg.starts_moved(p2);
stack_jxg.starts_moved(p3);
[[/jsxgraph]]

<div style="display:none">
[[input:ans3]] [[validation:ans3]]
[[input:ans4]] [[validation:ans4]]
[[input:ans5]] [[validation:ans5]]
</div>
<div><b>Boundary Line Status:</b> [[feedback:prt3]]</div>
<div><b>Region Marker Status:</b> [[feedback:prt4]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked Solution:</b></p>

<p><b>Part a: The Algebraic Boundary</b></p>
<p>To plot the boundary line \( {@A*x + B*y@} = {@C@} \), we find the axes intercepts.</p>
<p>First, we find the \(y\)-intercept by substituting \( x = 0 \):</p>
<p style="margin-left: 20px;">
\( \displaystyle \begin{aligned}
{@A@}(0) + {@B@}y &= {@C@} \\
{@B@}y &= {@C@} \\
y &= {@ta1@}
\end{aligned} \)
</p>
<p>Next, we find the \(x\)-intercept by substituting \( y = 0 \):</p>
<p style="margin-left: 20px;">
\( \displaystyle \begin{aligned}
{@A@}x + {@B@}(0) &= {@C@} \\
{@A@}x &= {@C@} \\
x &= {@ta2@}
\end{aligned} \)
</p>
<p>The boundary line passes directly through \((0, {@ta1@})\) and \(({@ta2@}, 0)\).</p>

<p><b>Part b: Graphical Representation</b></p>
<p>We plot the two intercepts on the Cartesian plane and draw a <b>solid line</b> through them because the inequality is inclusive (\(\le\)).</p>
<p>Next, we determine which region to shade using a test point. The origin \((0,0)\) is the simplest test point. We substitute \(x = 0\) and \(y = 0\) into our inequality:</p>
<p style="margin-left: 20px;">
\( \displaystyle \begin{aligned}
{@A*x + B*y@} &\le {@C@} \\
{@A@}(0) + {@B@}(0) &\le {@C@} \\
0 &\le {@C@}
\end{aligned} \)
</p>
<p>This statement {@origin_statement@}. Therefore, we place our region marker on the corresponding side of the solid line to indicate the shaded area.</p>

[[jsxgraph width="400px" height="400px"]]
var a_val = {#A#};
var b_val = {#B#};
var c_val = {#C#};

var brd = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-10.5, 10.5, 10.5, -10.5],
    axis: false, grid: false, showNavigation: false, showCopyright: false, pan: {enabled: false}, zoom: {enabled: false}
});

var xaxis = brd.create('axis', [[0, 0], [1, 0]], {
    name: 'x', withLabel: true, label: {position: 'rt', offset: [-15, 20]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-8, -15]}, drawZero: true, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

var yaxis = brd.create('axis', [[0, 0], [0, 1]], {
    name: 'y', withLabel: true, label: {position: 'rt', offset: [15, -15]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-25, -5]}, drawZero: false, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

/* Draw correct boundary line */
var p1 = brd.create('point', [0, c_val / b_val], {visible: false});
var p2 = brd.create('point', [1, (c_val - a_val) / b_val], {visible: false});
var line = brd.create('line', [p1, p2], {strokeWidth: 3, strokeColor: '#1565c0', dash: 0});

/* Shade the correct region. If B > 0, shading y <= ... means shading below (right of the directed line). Inverse: true handles this. */
var ineq = brd.create('inequality', [line], {fillColor: '#4caf50', fillOpacity: 0.2, inverse: (b_val > 0)});

/* Show the marker in the origin if valid, else in a far quadrant */
var mx = (0 <= c_val) ? 0 : -8 * Math.sign(a_val);
var my = (0 <= c_val) ? 0 : -8 * Math.sign(b_val);
brd.create('point', [mx, my], {name: 'Solution Region', size: 6, strokeColor: '#2e7d32', fillColor: '#4caf50', fixed: true});
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>4</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Use predefined coprime tuples to ensure clean integer intercepts */
tuples : [[2, 3, 6], [3, 2, 6], [3, 4, 12], [4, 3, 12], [2, 5, 10], [5, 2, 10], [1, 2, 4], [2, 1, 4], [1, 3, 6], [3, 1, 6]];
tup : rand(tuples);
A_mag : tup[1];
B_mag : tup[2];
C_mag : tup[3];

sign_A : rand([1, -1]);
sign_B : rand([1, -1]);
sign_C : rand([1, -1]);

A : A_mag * sign_A;
B : B_mag * sign_B;
C : C_mag * sign_C;

/* Target Intercepts */
ta1 : C / B;
ta2 : C / A;

/* Graph Line Targets */
ta3 : [0, ta1];
ta4 : [ta2, 0];

/* Region Marker Targets - dynamically determined to guarantee safe placement */
ta5_check : if (0 <= C) then true else false;
ta5 : if ta5_check then [0, 0] else [-8 * sign_A, -8 * sign_B];
ta5_wrong : [8 * sign_A, 8 * sign_B];

/* Text variables for the worked solution */
origin_statement : if ta5_check then "is <b>true</b>, so the origin \\((0,0)\\) belongs in the solution region" else "is <b>false</b>, so the origin \\((0,0)\\) is NOT in the solution region";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\({@A*x + B*y@} \le {@C@} \implies y_{int} = {@ta1@}, x_{int} = {@ta2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="concept">Graphing Inequalities in Two Unknowns (Inclusive Standard Form)</p>
<p id="description">Grade 11 CBE: Finding axes intercepts from standard form and graphing a solid boundary line with a region marker.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check absolute correctness of y-int</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check absolute correctness of x-int</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Extract coordinates of the two blue points */
x1 : ans3[1];
y1 : ans3[2];
x2 : ans4[1];
y2 : ans4[2];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check if points are distinct</description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is((x1-x2)^2 + (y1-y2)^2 > 0)]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You have dragged both blue points to the exact same location! A line requires two distinct points.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Check if P1 is on the line Ax + By = C (using 0.5 tolerance for JS grid snapping)</description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(abs(A*x1 + B*y1 - C) < 0.5)]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your boundary line is placed incorrectly. Use the intercepts you calculated in Part a to accurately position the blue points.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Check if P2 is on the line Ax + By = C (using 0.5 tolerance for JS grid snapping)</description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(abs(A*x2 + B*y2 - C) < 0.5)]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your boundary line is placed correctly.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your boundary line is placed incorrectly. Use the intercepts you calculated in Part a to accurately position both blue points.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Extract coordinate of the green region marker */
rx : ans5[1];
ry : ans5[2];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description><![CDATA[Check if marker satisfies Ax + By <= C]]></description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(A*rx + B*ry <= C + 0.1)]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your region marker is placed correctly.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your green region marker is in the wrong area. It does not satisfy the inequality \( {@A*x + B*y@} \le {@C@} \).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1866962349</deployedseed>
    <deployedseed>980625921</deployedseed>
    <deployedseed>467993116</deployedseed>
    <deployedseed>742000283</deployedseed>
    <deployedseed>1824327073</deployedseed>
    <deployedseed>452542959</deployedseed>
    <deployedseed>899233517</deployedseed>
    <deployedseed>90841622</deployedseed>
    <deployedseed>1511114865</deployedseed>
    <deployedseed>549617125</deployedseed>
    <deployedseed>619976563</deployedseed>
    <deployedseed>1741431047</deployedseed>
    <deployedseed>1946490121</deployedseed>
    <deployedseed>1506813203</deployedseed>
    <deployedseed>406968633</deployedseed>
    <deployedseed>597982603</deployedseed>
    <deployedseed>795445612</deployedseed>
    <deployedseed>221445806</deployedseed>
    <deployedseed>348399898</deployedseed>
    <deployedseed>187238350</deployedseed>
    <deployedseed>1403941582</deployedseed>
    <deployedseed>459634411</deployedseed>
    <deployedseed>840609740</deployedseed>
    <deployedseed>1826384610</deployedseed>
    <deployedseed>316834790</deployedseed>
    <deployedseed>1777004959</deployedseed>
    <deployedseed>763576718</deployedseed>
    <deployedseed>361901305</deployedseed>
    <deployedseed>2126527256</deployedseed>
    <deployedseed>1921655107</deployedseed>
    <deployedseed>1684565185</deployedseed>
    <deployedseed>729275112</deployedseed>
    <deployedseed>913721598</deployedseed>
    <deployedseed>1287117451</deployedseed>
    <deployedseed>2093912310</deployedseed>
    <deployedseed>1612765062</deployedseed>
    <deployedseed>374973738</deployedseed>
    <deployedseed>1199220676</deployedseed>
    <deployedseed>2061096180</deployedseed>
    <deployedseed>659139071</deployedseed>
    <deployedseed>567810623</deployedseed>
    <deployedseed>620499140</deployedseed>
    <deployedseed>129838073</deployedseed>
    <deployedseed>848091143</deployedseed>
    <deployedseed>10320572</deployedseed>
    <deployedseed>1339997900</deployedseed>
    <deployedseed>1346668222</deployedseed>
    <deployedseed>1651998466</deployedseed>
    <deployedseed>982579376</deployedseed>
    <deployedseed>696194292</deployedseed>
    <deployedseed>1825440825</deployedseed>
    <deployedseed>1312029863</deployedseed>
    <deployedseed>1468333591</deployedseed>
    <deployedseed>684339491</deployedseed>
    <deployedseed>2048813391</deployedseed>
    <deployedseed>1539026913</deployedseed>
    <deployedseed>1562650232</deployedseed>
    <deployedseed>325567466</deployedseed>
    <deployedseed>91241526</deployedseed>
    <deployedseed>464794620</deployedseed>
    <deployedseed>21540774</deployedseed>
    <deployedseed>78309030</deployedseed>
    <deployedseed>196887802</deployedseed>
    <deployedseed>1420413734</deployedseed>
    <deployedseed>1936262081</deployedseed>
    <deployedseed>917272765</deployedseed>
    <deployedseed>1341689496</deployedseed>
    <deployedseed>1121988920</deployedseed>
    <deployedseed>1875847832</deployedseed>
    <deployedseed>1135956046</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Everything correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Correct line, but wrong region</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5_wrong</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Wrong intercepts and wrong line, but marker happens to hit safe zone</description>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>[0, ta1+1]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>[ta2+1, 0]</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
