<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Graphing a linear inequality in two unknowns</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Solve and graph the following linear inequality in two unknowns:</p>
<p style="margin-left: 20px;">
\( \displaystyle y > {@m@}x + {@c@} \)
</p>

<p><b>a) </b>< Find the axes intercepts for the boundary equation \( y = {@m@}x + {@c@} \).</p>
<p><i>(Enter the numerical values for the intercepts below.)</i></p>

<p>\(y\)-intercept (when \(x = 0\)): \( y = \) [[input:ans1]]</p>
<p>\(x\)-intercept (when \(y = 0\)): \( x = \) [[input:ans2]]</p>

<div><b>\(y\)-intercept check:</b> [[validation:ans1]] [[feedback:prt1]]</div>
<div><b>\(x\)-intercept check:</b> [[validation:ans2]] [[feedback:prt2]]</div>

<p><b>b)</b> Graph the inequality on the Cartesian plane below:</p>
<ol>
<li>Drag the <b>two blue points</b> to set the boundary line. <i>(Notice the line is dashed because the inequality is strict: \(>\)).</i></li>
<li>Drag the <b>green Region Marker</b> into the area that satisfies the inequality \( y > {@m@}x + {@c@} \).</li>
</ol>

[[jsxgraph input-ref-ans3="rP1" input-ref-ans4="rP2" input-ref-ans5="rP3" width="600px" height="600px"]]
var brd = JXG.JSXGraph.initBoard(divid, {
    /* 10.5 creates a physical wall, preventing gliders from going past the 10 grid line */
    boundingbox: [-10.5, 10.5, 10.5, -10.5],
    axis: false, grid: false, showNavigation: false, showCopyright: false, pan: {enabled: false}, zoom: {enabled: false}
});

/* Create x-axis. Zero is shifted left slightly to clear the y-axis line */
var xaxis = brd.create('axis', [[0, 0], [1, 0]], {
    name: 'x', withLabel: true, label: {position: 'rt', offset: [-15, 20]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-8, -15]}, drawZero: true, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1, /* -1 forces the ticks to become infinite grid lines */
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

/* Create y-axis. Labels shifted to the left side of the axis */
var yaxis = brd.create('axis', [[0, 0], [0, 1]], {
    name: 'y', withLabel: true, label: {position: 'rt', offset: [15, -15]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-25, -5]}, drawZero: false, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

/* Create two boundary points starting at default positions */
var p1 = brd.create('point', [0, 5], {
    name: '', size: 5, strokeColor: '#1565c0', fillColor: '#ffffff', strokeWidth: 3, snapToGrid: true, snapSizeX: 1, snapSizeY: 1
});
var p2 = brd.create('point', [-5, 0], {
    name: '', size: 5, strokeColor: '#1565c0', fillColor: '#ffffff', strokeWidth: 3, snapToGrid: true, snapSizeX: 1, snapSizeY: 1
});

/* Draw the boundary line (dashed for strict inequality) */
var line = brd.create('line', [p1, p2], {
    strokeWidth: 3, strokeColor: '#1565c0', dash: 2, highlight: false
});

/* Create the Region Marker (the point that determines the shaded side) */
var p3 = brd.create('point', [4, -4], {
    name: 'Region', size: 6, strokeColor: '#2e7d32', fillColor: '#4caf50', strokeWidth: 2, snapToGrid: true, snapSizeX: 1, snapSizeY: 1
});

/* Bind the points to hidden STACK inputs */
stack_jxg.bind_point(rP1, p1);
stack_jxg.bind_point(rP2, p2);
stack_jxg.bind_point(rP3, p3);
stack_jxg.starts_moved(p1);
stack_jxg.starts_moved(p2);
stack_jxg.starts_moved(p3);
[[/jsxgraph]]

<div style="display:none">
[[input:ans3]] [[validation:ans3]]
[[input:ans4]] [[validation:ans4]]
[[input:ans5]] [[validation:ans5]]
</div>
<div><b>Boundary Line Status:</b> [[feedback:prt3]]</div>
<div><b>Region Marker Status:</b> [[feedback:prt4]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked Solution:</b></p>

<p><b>Part a: The Algebraic Boundary</b></p>
<p>To plot the boundary line \( y = {@m@}x + {@c@} \), we find the axes intercepts.</p>
<p>For the \(y\)-intercept, we set \( x = 0 \):</p>
<p style="margin-left: 20px;">
\( \displaystyle \begin{aligned}
y &= {@m@}(0) + {@c@} \\
y &= {@ta1@}
\end{aligned} \)
</p>
<p>For the \(x\)-intercept, we set \( y = 0 \) and solve for \( x \):</p>
<p style="margin-left: 20px;">
\( \displaystyle \begin{aligned}
0 &= {@m@}x + {@c@} \\
-{@c@} &= {@m@}x \\
x &= {@ta2@}
\end{aligned} \)
</p>
<p>The boundary line passes through \((0, {@ta1@})\) and \(({@ta2@}, 0)\).</p>

<p><b>Part b: Graphical Representation</b></p>
<p>We plot the two intercepts on the Cartesian plane and draw a <b>dashed line</b> through them because the inequality is strictly greater than (\(>\)).</p>
<p>Next, we determine which region to shade using a test point. The easiest test point is the origin \((0,0)\). We substitute \(x = 0\) and \(y = 0\) into our inequality:</p>
<p style="margin-left: 20px;">
\( \displaystyle \begin{aligned}
y &> {@m@}x + {@c@} \\
0 &> {@m@}(0) + {@c@} \\
0 &> {@c@}
\end{aligned} \)
</p>
<p>This statement {@origin_statement@}. Therefore, we place our region marker on the corresponding side of the dashed line to indicate the entire shaded area.</p>

[[jsxgraph width="400px" height="400px"]]
var m = {#m#};
var c = {#c#};

var brd = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-10.5, 10.5, 10.5, -10.5],
    axis: false, grid: false, showNavigation: false, showCopyright: false, pan: {enabled: false}, zoom: {enabled: false}
});

var xaxis = brd.create('axis', [[0, 0], [1, 0]], {
    name: 'x', withLabel: true, label: {position: 'rt', offset: [-15, 20]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-8, -15]}, drawZero: true, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

var yaxis = brd.create('axis', [[0, 0], [0, 1]], {
    name: 'y', withLabel: true, label: {position: 'rt', offset: [15, -15]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-25, -5]}, drawZero: false, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

/* Draw correct boundary line */
var p1 = brd.create('point', [0, c], {visible: false});
var p2 = brd.create('point', [1, m + c], {visible: false});
var line = brd.create('line', [p1, p2], {strokeWidth: 3, strokeColor: '#1565c0', dash: 2});

/* Shade the correct region (inverse: true means we shade y > mx + c) */
var ineq = brd.create('inequality', [line], {fillColor: '#4caf50', fillOpacity: 0.2, inverse: true});

/* Show the marker in a generic correct spot */
var mx = 0;
var my = c + 4;
brd.create('point', [mx, my], {name: 'Solution Region', size: 6, strokeColor: '#2e7d32', fillColor: '#4caf50', fixed: true});
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>4</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Randomize slope (m) to be non-zero integers */
m : rand([1, 2, -1, -2]);

/* Randomize x-intercept to be a non-zero integer */
x_int_temp : rand_with_step(-4, 4, 1);
x_int : if is(x_int_temp = 0) then 2 else x_int_temp;

/* Calculate y-intercept (c) so that 0 = m*x_int + c */
c : -m * x_int;

/* Target Answers */
ta1 : c;
ta2 : x_int;

/* These target variables must be mathematically correct to pass the PRT nodes when the 'Fill in correct responses' button is used */
ta3 : [0, c];         /* Correct y-intercept */
ta4 : [x_int, 0];     /* Correct x-intercept */
ta5 : [0, c + 4];     /* A guaranteed point in the y > mx + c region */

/* Text variables for the worked solution */
origin_statement : if is(0 > c) then "is <b>true</b>, so the origin \\((0,0)\\) belongs in the solution region" else "is <b>false</b>, so the origin \\((0,0)\\) is NOT in the solution region";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(y > {@m@}x + {@c@} \implies y_{int} = {@ta1@}, x_{int} = {@ta2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="concept">Graphing Inequalities in Two Unknowns (Strict, y > mx + c)</p>
<p id="description">Grade 11 CBE: Finding intercepts and graphing a dashed boundary line with a region marker on a fully gridded 2D Cartesian plane.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check absolute correctness of y-int</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check absolute correctness of x-int</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Extract coordinates of the two blue points */
x1 : ans3[1];
y1 : ans3[2];
x2 : ans4[1];
y2 : ans4[2];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check if points are distinct</description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is((x1-x2)^2 + (y1-y2)^2 > 0)]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You have dragged both blue points to the exact same location! A line requires two distinct points.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Check if P1 is on the line y = mx + c (using absolute tolerance for JS float issues)</description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(abs(y1 - (m*x1 + c)) < 0.1)]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your boundary line is placed incorrectly. Use the intercepts you calculated in Part a to accurately position the blue points.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Check if P2 is on the line y = mx + c (using absolute tolerance for JS float issues)</description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(abs(y2 - (m*x2 + c)) < 0.1)]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your boundary line is placed correctly.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your boundary line is placed incorrectly. Use the intercepts you calculated in Part a to accurately position both blue points.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Extract coordinate of the green region marker */
rx : ans5[1];
ry : ans5[2];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description><![CDATA[Check if marker satisfies y > mx + c]]></description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(ry > (m*rx + c))]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your region marker is placed correctly.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your green region marker is in the wrong area. It does not satisfy the inequality \( y > {@m@}x + {@c@} \).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>454549787</deployedseed>
    <deployedseed>436782757</deployedseed>
    <deployedseed>652750931</deployedseed>
    <deployedseed>2043264177</deployedseed>
    <deployedseed>1026061986</deployedseed>
    <deployedseed>141514305</deployedseed>
    <deployedseed>71619886</deployedseed>
    <deployedseed>1312782308</deployedseed>
    <deployedseed>1383818811</deployedseed>
    <deployedseed>916400037</deployedseed>
    <deployedseed>1429907158</deployedseed>
    <deployedseed>967727194</deployedseed>
    <deployedseed>924168918</deployedseed>
    <deployedseed>1804702747</deployedseed>
    <deployedseed>1325860372</deployedseed>
    <deployedseed>1661489009</deployedseed>
    <deployedseed>311879886</deployedseed>
    <deployedseed>1923543976</deployedseed>
    <deployedseed>2014827602</deployedseed>
    <deployedseed>1315065274</deployedseed>
    <deployedseed>288594272</deployedseed>
    <deployedseed>288636999</deployedseed>
    <deployedseed>135034588</deployedseed>
    <deployedseed>689616333</deployedseed>
    <deployedseed>910890138</deployedseed>
    <deployedseed>57642655</deployedseed>
    <deployedseed>1900380559</deployedseed>
    <deployedseed>1343390334</deployedseed>
    <deployedseed>1110601052</deployedseed>
    <deployedseed>1418092589</deployedseed>
    <deployedseed>1405330619</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Everything correct (using mathematically accurate targets)</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Correct line, but wrong region</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>[0, c-4]</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Wrong intercepts and wrong line</description>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>[0, c+1]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>[1, m+c+1]</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
