<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Graphing a linear inequality passing through the origin</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Solve and graph the following linear inequality in two unknowns:</p>
<p style="margin-left: 20px;">
\( \displaystyle y &lt; {@m@}x \)
</p>

<p><b>a)</b> Find the \(y\)-intercept of the boundary equation \( y = {@m@}x \), and determine the corresponding \(y\)-value when \(x = {@x_val@}\).</p>
<p><i>(Enter the numerical values below.)</i></p>

<p>\(y\)-intercept (when \(x = 0\)): \( y = \) [[input:ans1]]</p>
<p>\(y\)-value when \(x = {@x_val@}\): \( y = \) [[input:ans2]]</p>

<div><b>\(y\)-intercept check:</b> [[validation:ans1]] [[feedback:prt1]]</div>
<div><b>Second point check:</b> [[validation:ans2]] [[feedback:prt2]]</div>

<p><b>b)</b> Graph the inequality on the Cartesian plane below:</p>
<ul>
<li>Drag the <b>two blue points</b> to set the boundary line. <i>(Notice the line is dashed because the inequality is strict: \(&lt;\)).</i></li>
<li>Drag the <b>green Region Marker</b> into the area that satisfies the inequality \( y &lt; {@m@}x \).</li>
</ul>

[[jsxgraph input-ref-ans3="rP1" input-ref-ans4="rP2" input-ref-ans5="rP3" width="600px" height="600px"]]
var brd = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-10.5, 10.5, 10.5, -10.5],
    axis: false, grid: false, showNavigation: false, showCopyright: false, pan: {enabled: false}, zoom: {enabled: false}
});

var xaxis = brd.create('axis', [[0, 0], [1, 0]], {
    name: 'x', withLabel: true, label: {position: 'rt', offset: [-15, 20]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-8, -15]}, drawZero: true, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

var yaxis = brd.create('axis', [[0, 0], [0, 1]], {
    name: 'y', withLabel: true, label: {position: 'rt', offset: [15, -15]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-25, -5]}, drawZero: false, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

/* Create two boundary points starting at default positions */
var p1 = brd.create('point', [2, 2], {
    name: '', size: 5, strokeColor: '#1565c0', fillColor: '#ffffff', strokeWidth: 3, snapToGrid: true, snapSizeX: 1, snapSizeY: 1
});
var p2 = brd.create('point', [-2, -2], {
    name: '', size: 5, strokeColor: '#1565c0', fillColor: '#ffffff', strokeWidth: 3, snapToGrid: true, snapSizeX: 1, snapSizeY: 1
});

/* Draw the boundary line (dashed for strict inequality) */
var line = brd.create('line', [p1, p2], {
    strokeWidth: 3, strokeColor: '#1565c0', dash: 2, highlight: false
});

/* Create the Region Marker */
var p3 = brd.create('point', [4, -4], {
    name: 'Region', size: 6, strokeColor: '#2e7d32', fillColor: '#4caf50', strokeWidth: 2, snapToGrid: true, snapSizeX: 1, snapSizeY: 1
});

stack_jxg.bind_point(rP1, p1);
stack_jxg.bind_point(rP2, p2);
stack_jxg.bind_point(rP3, p3);
stack_jxg.starts_moved(p1);
stack_jxg.starts_moved(p2);
stack_jxg.starts_moved(p3);
[[/jsxgraph]]

<div style="display:none">
[[input:ans3]] [[validation:ans3]]
[[input:ans4]] [[validation:ans4]]
[[input:ans5]] [[validation:ans5]]
</div>
<div><b>Boundary Line Status:</b> [[feedback:prt3]]</div>
<div><b>Region Marker Status:</b> [[feedback:prt4]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked Solution:</b></p>

<p><b>Part a: The Algebraic Boundary</b></p>
<p>To plot the boundary line \( y = {@m@}x \), we need to find two coordinate points.</p>
<p>First, we find the \(y\)-intercept by substituting \( x = 0 \):</p>
<p style="margin-left: 20px;">
\( \displaystyle \begin{aligned}
y &= {@m@}(0) \\
y &= 0
\end{aligned} \)
</p>
<p>Because the line passes exactly through the origin \((0,0)\), the \(x\)-intercept is also \(0\). We must evaluate a different \(x\)-value to find a second point. Substituting \( x = {@x_val@} \):</p>
<p style="margin-left: 20px;">
\( \displaystyle \begin{aligned}
y &= {@m@}({@x_val@}) \\
y &= {@ta2@}
\end{aligned} \)
</p>
<p>The boundary line passes through \((0, 0)\) and \(({@x_val@}, {@ta2@})\).</p>

<p><b>Part b: Graphical Representation</b></p>
<p>We plot our two points on the Cartesian plane and draw a <b>dashed line</b> through them because the inequality is strictly less than (\(&lt;\)).</p>
<p>Next, we determine which region to shade using a test point. Because the boundary line passes through the origin \((0,0)\), we <b>cannot</b> use it as a test point. Instead, we select a point clearly off the line, such as \((1, 0)\). Substituting \(x = 1\) and \(y = 0\) into our inequality:</p>
<p style="margin-left: 20px;">
\( \displaystyle \begin{aligned}
y &lt; {@m@}x \\
0 &lt; {@m@}(1) \\
0 &lt; {@m@}
\end{aligned} \)
</p>
<p>This statement {@test_statement@}. Therefore, we place our region marker on the {@shade_direction@} side of the dashed line as the point \((1,0)\) to indicate the shaded area.</p>

[[jsxgraph width="400px" height="400px"]]
var m_val = {#m#};

var brd = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-10.5, 10.5, 10.5, -10.5],
    axis: false, grid: false, showNavigation: false, showCopyright: false, pan: {enabled: false}, zoom: {enabled: false}
});

var xaxis = brd.create('axis', [[0, 0], [1, 0]], {
    name: 'x', withLabel: true, label: {position: 'rt', offset: [-15, 20]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-8, -15]}, drawZero: true, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

var yaxis = brd.create('axis', [[0, 0], [0, 1]], {
    name: 'y', withLabel: true, label: {position: 'rt', offset: [15, -15]},
    firstArrow: true, lastArrow: true, strokeWidth: 1.5,
    ticks: {
        drawLabels: true, label: {offset: [-25, -5]}, drawZero: false, insertTicks: false,
        ticksDistance: 2, minorTicks: 1,
        majorHeight: -1, minorHeight: -1,
        strokeColor: '#a0a0a0', strokeOpacity: 0.5, strokeWidth: 1,
        minorStrokeColor: '#e0e0e0', minorStrokeOpacity: 0.5, minorStrokeWidth: 1
    }
});

/* Draw correct boundary line */
var p1 = brd.create('point', [0, 0], {visible: false});
var p2 = brd.create('point', [1, m_val], {visible: false});
var line = brd.create('line', [p1, p2], {strokeWidth: 3, strokeColor: '#1565c0', dash: 2});

/* Shade the correct region. y < mx is always below the line, inverse: false handles this correctly for left-to-right drawn lines */
var ineq = brd.create('inequality', [line], {fillColor: '#4caf50', fillOpacity: 0.2, inverse: false});

/* Show the marker in a generic correct spot guaranteed to be y < mx */
brd.create('point', [0, -5], {name: 'Solution Region', size: 6, strokeColor: '#2e7d32', fillColor: '#4caf50', fixed: true});
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>4</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Randomize slope (m) avoiding 1 or -1 to make the substitution step more obvious */
m : rand([-3, -2, 2, 3]);

/* Ensure the resulting y-value doesn't exceed the grid boundary of 10 */
x_val : if abs(m)=3 then rand([2, 3]) else rand([3, 4]);

/* Target Answers */
ta1 : 0;
ta2 : m * x_val;

/* Graph Line Targets for 'Fill correct responses' */
ta3 : [0, 0];
ta4 : [x_val, ta2];

/* Region Marker Target - [0, -5] is mathematically guaranteed to be in y < mx for any m */
ta5 : [0, -5];
ta5_wrong : [0, 5];

/* Text variables for the worked solution */
test_truth : if (0 < m) then true else false;
test_statement : if test_truth then "is <b>true</b>" else "is <b>false</b>";
shade_direction : if test_truth then "same" else "opposite";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\(y < {@m@}x \implies y_{int} = 0, y({@x_val@}) = {@ta2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="concept">Graphing Inequalities in Two Unknowns (Passing through Origin)</p>
<p id="description">Grade 11 CBE: Finding points for and graphing a strict linear inequality that passes directly through the origin, requiring an alternative test point logic.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check absolute correctness of y-int</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check absolute correctness of secondary point</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Extract coordinates of the two blue points */
x1 : ans3[1];
y1 : ans3[2];
x2 : ans4[1];
y2 : ans4[2];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check if points are distinct</description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is((x1-x2)^2 + (y1-y2)^2 > 0)]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>You have dragged both blue points to the exact same location! A line requires two distinct points.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Check if P1 is on the line y = mx (using 0.5 tolerance for JS grid snapping)</description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(abs(y1 - m*x1) < 0.5)]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your boundary line is placed incorrectly. Use the points you calculated in Part a to accurately position the blue dots.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Check if P2 is on the line y = mx (using 0.5 tolerance for JS grid snapping)</description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(abs(y2 - m*x2) < 0.5)]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your boundary line is placed correctly.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>1</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your boundary line is placed incorrectly. Use the points you calculated in Part a to accurately position both blue dots.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* Extract coordinate of the green region marker */
rx : ans5[1];
ry : ans5[2];</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description><![CDATA[Check if marker satisfies y < mx]]></description>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(ry < (m*rx - 0.1))]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your region marker is placed correctly.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your green region marker is in the wrong area. It does not satisfy the inequality \( y &lt; {@m@}x \).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1941829659</deployedseed>
    <deployedseed>1568728986</deployedseed>
    <deployedseed>148666667</deployedseed>
    <deployedseed>1072729497</deployedseed>
    <deployedseed>997245110</deployedseed>
    <deployedseed>407578526</deployedseed>
    <deployedseed>2068897906</deployedseed>
    <deployedseed>12500026</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Everything correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Correct line, but wrong region</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5_wrong</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-3-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt4-1-F</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Wrong intercepts and wrong line</description>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>[0, ta1+1]</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>[x_val, ta2+1]</value>
      </testinput>
      <testinput>
        <name>ans5</name>
        <value>ta5</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-1-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt3-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-1-T</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
