<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Matrix Types and Orders (INNODEMS)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<style>
.que.stack table.cbgrid, .que.stack table.ordgrid { border-collapse: collapse; margin: 0.4em 0; }
.que.stack table.cbgrid td { padding: 0.12em 0.5em 0.12em 0; vertical-align: top; }
.que.stack table.cbgrid td.cbtype { white-space: nowrap; padding-right: 1em; }
.que.stack table.cbgrid .answer { display: grid; grid-template-columns: repeat(6, 3.5em); align-items: center; margin: 0; }
.que.stack table.cbgrid .answer > .option { display: inline-flex; align-items: center; white-space: nowrap; }
.que.stack table.cbgrid .answer > .option label { margin: 0 0 0 0.25em; }
.que.stack table.ordgrid td { padding: 0.1em 0.5em; white-space: nowrap; vertical-align: middle; }
</style>
<p>Consider the following matrices:</p>
<p>\[ A = {@mats[1]@}, \quad B = {@mats[2]@}, \quad C = {@mats[3]@} \]</p>
<p>\[ D = {@mats[4]@}, \quad E = {@mats[5]@}, \quad F = {@mats[6]@} \]</p>
<p><strong>(a)</strong> For each type below, tick every matrix that is of that type. A matrix may belong to more than one type.</p>
<p> Remember that a <strong> rectangular </strong> matrix must have a different number of rows and columns, it cannot be square.</p>
<table class="cbgrid"><tbody>
<tr><td class="cbtype"><strong>(i) Null matrices:</strong></td><td>[[input:cbnull]] [[validation:cbnull]] [[feedback:prtnull]]</td></tr>
<tr><td class="cbtype"><strong>(ii) Square matrices:</strong></td><td>[[input:cbsquare]] [[validation:cbsquare]] [[feedback:prtsquare]]</td></tr>
<tr><td class="cbtype"><strong>(iii) Rectangular matrices:</strong></td><td>[[input:cbrect]] [[validation:cbrect]] [[feedback:prtrect]]</td></tr>
<tr><td class="cbtype"><strong>(iv) Identity matrices:</strong></td><td>[[input:cbident]] [[validation:cbident]] [[feedback:prtident]]</td></tr>
<tr><td class="cbtype"><strong>(v) Row matrices:</strong></td><td>[[input:cbrow]] [[validation:cbrow]] [[feedback:prtrow]]</td></tr>
<tr><td class="cbtype"><strong>(vi) Column matrices:</strong></td><td>[[input:cbcol]] [[validation:cbcol]] [[feedback:prtcol]]</td></tr>
</tbody></table>
<p><strong>(b)</strong> Write down the order of each matrix: the number of rows in the first box and the number of columns in the second.</p>
<table class="ordgrid"><tbody>
<tr><td>\(A:\)</td><td>[[input:rowA]] [[validation:rowA]]</td><td>\(\times\)</td><td>[[input:colA]] [[validation:colA]]</td><td>[[feedback:prtordA]]</td></tr>
<tr><td>\(B:\)</td><td>[[input:rowB]] [[validation:rowB]]</td><td>\(\times\)</td><td>[[input:colB]] [[validation:colB]]</td><td>[[feedback:prtordB]]</td></tr>
<tr><td>\(C:\)</td><td>[[input:rowC]] [[validation:rowC]]</td><td>\(\times\)</td><td>[[input:colC]] [[validation:colC]]</td><td>[[feedback:prtordC]]</td></tr>
<tr><td>\(D:\)</td><td>[[input:rowD]] [[validation:rowD]]</td><td>\(\times\)</td><td>[[input:colD]] [[validation:colD]]</td><td>[[feedback:prtordD]]</td></tr>
<tr><td>\(E:\)</td><td>[[input:rowE]] [[validation:rowE]]</td><td>\(\times\)</td><td>[[input:colE]] [[validation:colE]]</td><td>[[feedback:prtordE]]</td></tr>
<tr><td>\(F:\)</td><td>[[input:rowF]] [[validation:rowF]]</td><td>\(\times\)</td><td>[[input:colF]] [[validation:colF]]</td><td>[[feedback:prtordF]]</td></tr>
</tbody></table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The <strong>order</strong> of a matrix is its number of rows \(\times\) its number of columns. The types are: <strong>null</strong> (every entry is 0); <strong>square</strong> (equal numbers of rows and columns); <strong>rectangular</strong> (rows and columns differ &ndash; this includes row and column matrices); <strong>identity</strong> (square, with 1s on the leading diagonal and 0s elsewhere); <strong>row</strong> (a single row); <strong>column</strong> (a single column).</p>
<p><strong>For this version:</strong></p>
<ul>
<li>Null: \({@key_null@}\)</li>
<li>Square: \({@key_square@}\)</li>
<li>Rectangular: \({@key_rect@}\)</li>
<li>Identity: \({@key_ident@}\)</li>
<li>Row: \({@key_row@}\)</li>
<li>Column: \({@key_col@}\)</li>
</ul>
<p>Orders: \(A:{@nrows[1]@}\times{@ncols[1]@}\), \(B:{@nrows[2]@}\times{@ncols[2]@}\), \(C:{@nrows[3]@}\times{@ncols[3]@}\), \(D:{@nrows[4]@}\times{@ncols[4]@}\), \(E:{@nrows[5]@}\times{@ncols[5]@}\), \(F:{@nrows[6]@}\times{@ncols[6]@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>12</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>labs : [A,B,C,D,E,F];
Mrole : [[1,0,1,0,0,0],[0,1,0,1,0,0],[0,0,1,0,1,0],[1,0,1,0,0,1],[0,0,1,0,0,0],[0,1,0,0,0,0]];
s_id : rand([2,3]);
s_sq : rand([2,3]);
dimsR : rand([[2,3],[3,2]]);
mR : dimsR[1];
nR : dimsR[2];
dimsN : rand([[2,3],[3,2]]);
mN : dimsN[1];
nN : dimsN[2];
n_row : rand([3,4]);
m_col : rand([3,4]);
mat1 : zeromatrix(mN,nN);
mat2 : ident(s_id);
mat3 : genmatrix(lambda([ii,jj], rand_with_prohib(-5,5,[0])),1,n_row);
mat4 : zeromatrix(m_col,1);
mat5 : genmatrix(lambda([ii,jj], rand_with_prohib(-4,4,[0])),mR,nR);
mat6 : genmatrix(lambda([ii,jj], rand_with_prohib(-4,4,[0])),s_sq,s_sq);
roles : [mat1,mat2,mat3,mat4,mat5,mat6];
perm : rand_selection([1,2,3,4,5,6],6);
mats : makelist(roles[perm[k]],k,1,6);
memb : makelist(Mrole[perm[k]],k,1,6);
nrows : makelist(length(mats[k]),k,1,6);
ncols : makelist(length(first(mats[k])),k,1,6);
idx : makelist(k,k,1,6);
key_null : setify(map(lambda([k],labs[k]),sublist(idx,lambda([k],is(memb[k][1]=1)))));
key_square : setify(map(lambda([k],labs[k]),sublist(idx,lambda([k],is(memb[k][2]=1)))));
key_rect : setify(map(lambda([k],labs[k]),sublist(idx,lambda([k],is(memb[k][3]=1)))));
key_ident : setify(map(lambda([k],labs[k]),sublist(idx,lambda([k],is(memb[k][4]=1)))));
key_row : setify(map(lambda([k],labs[k]),sublist(idx,lambda([k],is(memb[k][5]=1)))));
key_col : setify(map(lambda([k],labs[k]),sublist(idx,lambda([k],is(memb[k][6]=1)))));
opts_null : makelist([labs[k],is(memb[k][1]=1)],k,1,6);
opts_square : makelist([labs[k],is(memb[k][2]=1)],k,1,6);
opts_rect : makelist([labs[k],is(memb[k][3]=1)],k,1,6);
opts_ident : makelist([labs[k],is(memb[k][4]=1)],k,1,6);
opts_row : makelist([labs[k],is(memb[k][5]=1)],k,1,6);
opts_col : makelist([labs[k],is(memb[k][6]=1)],k,1,6);
lab_ident : first(listify(key_ident));</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>Null \({@key_null@}\), Sq \({@key_square@}\), Rect \({@key_rect@}\), Id \({@key_ident@}\), Row \({@key_row@}\), Col \({@key_col@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>cbcol</name>
      <type>checkbox</type>
      <tans>opts_col</tans>
      <boxsize>0</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>cbident</name>
      <type>checkbox</type>
      <tans>opts_ident</tans>
      <boxsize>0</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>cbnull</name>
      <type>checkbox</type>
      <tans>opts_null</tans>
      <boxsize>0</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>cbrect</name>
      <type>checkbox</type>
      <tans>opts_rect</tans>
      <boxsize>0</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>cbrow</name>
      <type>checkbox</type>
      <tans>opts_row</tans>
      <boxsize>0</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>cbsquare</name>
      <type>checkbox</type>
      <tans>opts_square</tans>
      <boxsize>0</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>colA</name>
      <type>numerical</type>
      <tans>ncols[1]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>colB</name>
      <type>numerical</type>
      <tans>ncols[2]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>colC</name>
      <type>numerical</type>
      <tans>ncols[3]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>colD</name>
      <type>numerical</type>
      <tans>ncols[4]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>colE</name>
      <type>numerical</type>
      <tans>ncols[5]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>colF</name>
      <type>numerical</type>
      <tans>ncols[6]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>rowA</name>
      <type>numerical</type>
      <tans>nrows[1]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>rowB</name>
      <type>numerical</type>
      <tans>nrows[2]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>rowC</name>
      <type>numerical</type>
      <tans>nrows[3]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>rowD</name>
      <type>numerical</type>
      <tans>nrows[4]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>rowE</name>
      <type>numerical</type>
      <tans>nrows[5]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>rowF</name>
      <type>numerical</type>
      <tans>nrows[6]</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prtcol</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sa : setify(cbcol);
miss : setdifference(key_col, sa);
extra : setdifference(sa, key_col);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sa</sans>
        <tans>key_col</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>clscol-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>clscol-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. A <strong>column matrix</strong> has a single column (order \(m \times 1\)). You missed \({@miss@}\) and wrongly included \({@extra@}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtident</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sa : setify(cbident);
miss : setdifference(key_ident, sa);
extra : setdifference(sa, key_ident);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sa</sans>
        <tans>key_ident</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>clsid-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>clsid-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. An <strong>identity matrix</strong> is square with 1s on the leading diagonal and 0s elsewhere. You missed \({@miss@}\) and wrongly included \({@extra@}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtnull</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sa : setify(cbnull);
miss : setdifference(key_null, sa);
extra : setdifference(sa, key_null);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sa</sans>
        <tans>key_null</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>clsnull-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>clsnull-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. A <strong>null (zero) matrix</strong> has every entry equal to 0, whatever its shape. You missed \({@miss@}\) and wrongly included \({@extra@}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtordA</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowA</sans>
        <tans>nrows[1]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>ordA-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>ordA-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colA</sans>
        <tans>ncols[1]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordA-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordA-col</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The number of rows of \(A\) is right, but recount the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowA</sans>
        <tans>ncols[1]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>ordA-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordA-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The order is (number of rows) \(\times\) (number of columns). Count the rows of \(A\) first, then the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colA</sans>
        <tans>nrows[1]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordA-swap</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have the right two numbers but the wrong way round: the order is rows \(\times\) columns, not columns \(\times\) rows.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordA-no2</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not correct. Count the number of rows and the number of columns of \(A\) again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtordB</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowB</sans>
        <tans>nrows[2]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>ordB-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>ordB-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colB</sans>
        <tans>ncols[2]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordB-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordB-col</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The number of rows of \(B\) is right, but recount the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowB</sans>
        <tans>ncols[2]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>ordB-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordB-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The order is (number of rows) \(\times\) (number of columns). Count the rows of \(B\) first, then the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colB</sans>
        <tans>nrows[2]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordB-swap</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have the right two numbers but the wrong way round: the order is rows \(\times\) columns, not columns \(\times\) rows.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordB-no2</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not correct. Count the number of rows and the number of columns of \(B\) again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtordC</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowC</sans>
        <tans>nrows[3]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>ordC-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>ordC-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colC</sans>
        <tans>ncols[3]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordC-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordC-col</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The number of rows of \(C\) is right, but recount the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowC</sans>
        <tans>ncols[3]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>ordC-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordC-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The order is (number of rows) \(\times\) (number of columns). Count the rows of \(C\) first, then the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colC</sans>
        <tans>nrows[3]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordC-swap</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have the right two numbers but the wrong way round: the order is rows \(\times\) columns, not columns \(\times\) rows.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordC-no2</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not correct. Count the number of rows and the number of columns of \(C\) again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtordD</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowD</sans>
        <tans>nrows[4]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>ordD-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>ordD-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colD</sans>
        <tans>ncols[4]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordD-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordD-col</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The number of rows of \(D\) is right, but recount the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowD</sans>
        <tans>ncols[4]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>ordD-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordD-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The order is (number of rows) \(\times\) (number of columns). Count the rows of \(D\) first, then the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colD</sans>
        <tans>nrows[4]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordD-swap</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have the right two numbers but the wrong way round: the order is rows \(\times\) columns, not columns \(\times\) rows.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordD-no2</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not correct. Count the number of rows and the number of columns of \(D\) again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtordE</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowE</sans>
        <tans>nrows[5]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>ordE-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>ordE-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colE</sans>
        <tans>ncols[5]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordE-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordE-col</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The number of rows of \(E\) is right, but recount the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowE</sans>
        <tans>ncols[5]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>ordE-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordE-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The order is (number of rows) \(\times\) (number of columns). Count the rows of \(E\) first, then the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colE</sans>
        <tans>nrows[5]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordE-swap</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have the right two numbers but the wrong way round: the order is rows \(\times\) columns, not columns \(\times\) rows.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordE-no2</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not correct. Count the number of rows and the number of columns of \(E\) again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtordF</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowF</sans>
        <tans>nrows[6]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>1</truenextnode>
        <trueanswernote>ordF-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>ordF-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colF</sans>
        <tans>ncols[6]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordF-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordF-col</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The number of rows of \(F\) is right, but recount the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>rowF</sans>
        <tans>ncols[6]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>ordF-2-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordF-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The order is (number of rows) \(\times\) (number of columns). Count the rows of \(F\) first, then the columns.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>colF</sans>
        <tans>nrows[6]</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>ordF-swap</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You have the right two numbers but the wrong way round: the order is rows \(\times\) columns, not columns \(\times\) rows.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>ordF-no2</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not correct. Count the number of rows and the number of columns of \(F\) again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtrect</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text><![CDATA[sa : setify(cbrect);
miss : setdifference(key_rect, sa);
extra : setdifference(sa, key_rect);
inc_square : intersection(sa, key_square);
has_sq_err : is(cardinality(inc_square)>0);]]></text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sa</sans>
        <tans>key_rect</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>clsrect-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>clsrect-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>has_sq_err</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>clsrect-square</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>In this convention a <strong>rectangular</strong> matrix is one that is <strong>not square</strong> (its numbers of rows and columns differ). You included the square matrix \({@inc_square@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>clsrect-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. Here <strong>rectangular</strong> means non-square &ndash; remember that row and column matrices count as rectangular too. You missed \({@miss@}\) and wrongly included \({@extra@}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtrow</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sa : setify(cbrow);
miss : setdifference(key_row, sa);
extra : setdifference(sa, key_row);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sa</sans>
        <tans>key_row</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>clsrow-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>clsrow-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. A <strong>row matrix</strong> has a single row (order \(1 \times n\)). You missed \({@miss@}\) and wrongly included \({@extra@}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtsquare</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>sa : setify(cbsquare);
miss : setdifference(key_square, sa);
extra : setdifference(sa, key_square);
forgot_id : is(not elementp(lab_ident, sa));</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>sa</sans>
        <tans>key_square</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>clssq-ok</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>clssq-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>forgot_id</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>clssq-forgotid</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Remember that an <strong>identity matrix is also square</strong>, so matrix \({@lab_ident@}\) belongs here too.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>clssq-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. A <strong>square matrix</strong> has equal numbers of rows and columns. You missed \({@miss@}\) and wrongly included \({@extra@}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>101013</deployedseed>
    <deployedseed>202024</deployedseed>
    <deployedseed>303035</deployedseed>
    <deployedseed>404046</deployedseed>
    <deployedseed>505057</deployedseed>
    <deployedseed>606068</deployedseed>
    <deployedseed>707079</deployedseed>
    <deployedseed>808081</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All correct</description>
      <testinput>
        <name>cbcol</name>
        <value>mcq_correct(opts_col)</value>
      </testinput>
      <testinput>
        <name>cbident</name>
        <value>mcq_correct(opts_ident)</value>
      </testinput>
      <testinput>
        <name>cbnull</name>
        <value>mcq_correct(opts_null)</value>
      </testinput>
      <testinput>
        <name>cbrect</name>
        <value>mcq_correct(opts_rect)</value>
      </testinput>
      <testinput>
        <name>cbrow</name>
        <value>mcq_correct(opts_row)</value>
      </testinput>
      <testinput>
        <name>cbsquare</name>
        <value>mcq_correct(opts_square)</value>
      </testinput>
      <testinput>
        <name>colA</name>
        <value>ncols[1]</value>
      </testinput>
      <testinput>
        <name>colB</name>
        <value>ncols[2]</value>
      </testinput>
      <testinput>
        <name>colC</name>
        <value>ncols[3]</value>
      </testinput>
      <testinput>
        <name>colD</name>
        <value>ncols[4]</value>
      </testinput>
      <testinput>
        <name>colE</name>
        <value>ncols[5]</value>
      </testinput>
      <testinput>
        <name>colF</name>
        <value>ncols[6]</value>
      </testinput>
      <testinput>
        <name>rowA</name>
        <value>nrows[1]</value>
      </testinput>
      <testinput>
        <name>rowB</name>
        <value>nrows[2]</value>
      </testinput>
      <testinput>
        <name>rowC</name>
        <value>nrows[3]</value>
      </testinput>
      <testinput>
        <name>rowD</name>
        <value>nrows[4]</value>
      </testinput>
      <testinput>
        <name>rowE</name>
        <value>nrows[5]</value>
      </testinput>
      <testinput>
        <name>rowF</name>
        <value>nrows[6]</value>
      </testinput>
      <expected>
        <name>prtcol</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>clscol-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtident</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>clsid-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtnull</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>clsnull-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtordA</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>ordA-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtordB</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>ordB-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtordC</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>ordC-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtordD</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>ordD-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtordE</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>ordE-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtordF</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>ordF-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtrect</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>clsrect-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtrow</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>clsrow-ok</expectedanswernote>
      </expected>
      <expected>
        <name>prtsquare</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>clssq-ok</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Square: forgot the identity is square</description>
      <testinput>
        <name>cbsquare</name>
        <value>ev(listify(disjoin(lab_ident,key_square)),simp)</value>
      </testinput>
      <expected>
        <name>prtsquare</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>clssq-forgotid</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Rectangular: wrongly included a square</description>
      <testinput>
        <name>cbrect</name>
        <value>ev(append(mcq_correct(opts_rect),[lab_ident]),simp)</value>
      </testinput>
      <expected>
        <name>prtrect</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>clsrect-square</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Null: generic wrong set</description>
      <testinput>
        <name>cbnull</name>
        <value>mcq_correct(opts_square)</value>
      </testinput>
      <expected>
        <name>prtnull</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>clsnull-no</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Order A: columns wrong</description>
      <testinput>
        <name>colA</name>
        <value>ev(ncols[1]+1,simp)</value>
      </testinput>
      <testinput>
        <name>rowA</name>
        <value>nrows[1]</value>
      </testinput>
      <expected>
        <name>prtordA</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>ordA-col</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description>Order A: rows wrong (generic)</description>
      <testinput>
        <name>colA</name>
        <value>ncols[1]</value>
      </testinput>
      <testinput>
        <name>rowA</name>
        <value>ev(nrows[1]+5,simp)</value>
      </testinput>
      <expected>
        <name>prtordA</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>ordA-no</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
