<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>q2</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A shopkeeper weighed two bags of sugar using a weighing scale. The mass of the first bag was recorded as {@val1@} kg  and the mass of the second bag was recorded as {@val2@} kg.</p>

<p>Use these measurements to determine:</p>

<p>(a). The absolute error in the sum of the masses of the two bags.</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

<p>(b). The absolute error in the difference between the masses of the two bags.</p>
<p>[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>

<p>(c). The relative error in the sum of the masses of the two bags.</p>
<p>[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>

<p>(d). The relative error in the difference between the masses of the two bags.</p>
<p>[[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</p>

<p>(e). The percentage error in the sum of the masses of the two bags.</p>
<p>[[input:ans5]] [[validation:ans5]] [[feedback:prt5]]</p>

<p>(f). The percentage error in the difference between the masses of the two bags.</p>
<p>[[input:ans6]] [[validation:ans6]] [[feedback:prt6]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>The masses of the two bags of sugar were recorded as {@val1@} kg (to {@dp1@} decimal place(s)) and {@val2@} kg (to {@dp2@} decimal place(s)). Since each measurement is rounded to a given number of decimal places, the true mass lies within half of the smallest recorded unit on either side of the recorded value.</p>

<p>Precision of first bag = {@pre1@} kg</p>
<p>Precision of second bag = {@pre2@} kg</p>

<p>So the mass of the first bag lies between {@min1@} kg and {@max1@} kg, and the mass of the second bag lies between {@min2@} kg and {@max2@} kg.</p>

<p><strong>(a). Absolute error in the sum</strong></p>
<p>Maximum sum = {@max1@} + {@max2@} = {@maxs@} kg</p>
<p>Minimum sum = {@min1@} + {@min2@} = {@mins@} kg</p>
<p>Absolute error = \( \dfrac{{@maxs@} - {@mins@}}{2} = {@ta1@} \) kg</p>

<p><strong>(b). Absolute error in the difference</strong></p>
<p>Maximum difference = {@max1@} − {@min2@} = {@maxd@} kg</p>
<p>Minimum difference = {@min1@} − {@max2@} = {@mind@} kg</p>
<p>Absolute error = \( \dfrac{{@maxd@} - {@mind@}}{2} = {@ta2@} \) kg</p>

<p><strong>(c). Relative error in the sum</strong></p>
<p>Working sum = {@val1@} + {@val2@} = {@acts@} kg</p>
<p>Relative error = \( \dfrac{\text{absolute error}}{\text{working sum}} = \dfrac{{@ta1@}}{{@acts@}} = {@ta3@} \)</p>

<p><strong>(d). Relative error in the difference</strong></p>
<p>Working difference = {@val1@} − {@val2@} = {@actd@} kg</p>
<p>Relative error = \( \dfrac{\text{absolute error}}{\text{working difference}} = \dfrac{{@ta2@}}{{@actd@}} = {@ta4@} \)</p>

<p><strong>(e). Percentage error in the sum</strong></p>
<p>Percentage error = relative error \times 100\% = {@ta3@} \times 100\% = {@ta5@}\%</p>

<p><strong>(f). Percentage error in the difference</strong></p>
<p>Percentage error = relative error \times 100\% = {@ta4@} \times 100\% = {@ta6@}\%</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>val1: float((30 + rand(9)) + (rand(9) + 1) / 10);

val2: float((15 + rand(9)) + ((rand(9) + 1) * 10 + (rand(9) + 1)) / 100);

dp1: 1;
dp2: 2;

pre1: 1 / (2 * 10^dp1);
pre2: 1 / (2 * 10^dp2);

max1: val1 + pre1;
max2: val2 + pre2;

min1: val1 - pre1;
min2: val2 - pre2;

acts: val1 + val2;
actd: val1 - val2;

mins: min1 + min2;
maxs: max1 + max2;

mind: min1 - max2;
maxd: max1 - min2;

ta1: (maxs - mins) / 2;
ta2: (maxd - mind) / 2;
ta3: ta1 / acts;
ta4: ta2 / actd;
ta5: ta3 * 100;
ta6: ta4 * 100;
ta7: ta1 / actd;
ta8: ta2 / acts;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans5</name>
      <type>algebraic</type>
      <tans>ta5</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans6</name>
      <type>algebraic</type>
      <tans>ta6</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta7</tans>
        <testoptions>0.0001</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you may have divided the absolute error by the <strong>working difference</strong> of the length and width instead of their <strong>working sum</strong>. Since this part asks for the relative error <em>in the sum</em>, you need to divide by the working sum, not the working difference.</p>
<p>Relative error in the sum = \( \dfrac{\text{Absolute error in sum}}{\text{working sum}} = \dfrac{{@ta1@}}{{@acts@}} = {@ta3@} \)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans4</sans>
        <tans>ta8</tans>
        <testoptions>0.0001</testoptions>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you may have divided the absolute error by the <strong>working sum</strong> of the length and width instead of their <strong>working difference</strong>. Since this part asks for the relative error <em>in the difference</em>, you need to divide by the working difference, not the working sum.</p>
<p>Relative error in the sum = \( \dfrac{\text{Absolute error in sum}}{\text{working difference}} = \dfrac{{@ta2@}}{{@actd@}} = {@ta4@} \)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt5</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans5</sans>
        <tans>ta5</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt5-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt5-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt6</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans6</sans>
        <tans>ta6</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt6-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt6-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>437531962</deployedseed>
  </question>
</quiz>
