<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Quiz1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Consider the numbers:</p>
<p>\[ {@a1@}, \quad {@a2@}, \quad  {@a3@}, \quad \sqrt{ {@a5@}^2 } \]</p>
<p>Which one of these is irrational? <strong> You can write \(\sqrt 1\) as sqrt(1) and \(7\sqrt 20\) as 7*sqrt(20)</strong></p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>A number is <strong>irrational</strong> if it cannot be written as a ratio of two integers (i.e. its decimal expansion is non-terminating and non-repeating).</p>
<p>In this question:<p>
<p>{@a1@} is rational - it is an integer (or a fraction of integers).</p>
<p>{@a2@} is rational - it is a ratio of two integers, {@a2@}.</p>
<p>\(\sqrt{ {@a5@}^2 }\) <strong>looks</strong> irrational because of the square root symbol, but it simplifies to a whole number, so it is rational. </p>
<p>This is a common trap: not every square root is irrational only square roots of numbers that are NOT perfect squares are irrational.</p>
<p>{@a3@} = {@(a3)@} is irrational, because {@a3@} is not a perfect square, so its square root cannot be simplified to an integer or fraction.</p>
<p><strong> Always simplify a surd fully before deciding whether it
is rational or irrational.</strong></p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>a1: rand(rand(10)+1);        

a2: (rand(8)+1)/(rand(11)+10);


a3:sqrt( rand_with_prohib(1, 20, [1,4,9,16]));         

a5: sqrt((rand(5)+2)^2);                     

perfect_squares: [1,4,9,16,25,36,49,64,81,100,121,144,169,196,225,256,289,324,361,400];
ta1: a3;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}{@a1@}{@a2@}{@a3@}{@a5@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>803247031</deployedseed>
    <deployedseed>605788497</deployedseed>
    <deployedseed>1570270422</deployedseed>
    <deployedseed>495300067</deployedseed>
    <deployedseed>1750567940</deployedseed>
    <deployedseed>1690614933</deployedseed>
    <deployedseed>1738442145</deployedseed>
    <deployedseed>596575584</deployedseed>
    <deployedseed>294121560</deployedseed>
    <deployedseed>990597472</deployedseed>
    <deployedseed>1785977782</deployedseed>
    <deployedseed>919625520</deployedseed>
    <deployedseed>1016682722</deployedseed>
    <deployedseed>953547686</deployedseed>
    <deployedseed>769115340</deployedseed>
    <deployedseed>1141275053</deployedseed>
    <deployedseed>1741479450</deployedseed>
    <deployedseed>1243301736</deployedseed>
    <deployedseed>1618712035</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
