<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>carpet for a rug</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A rectangular rug has length \({@a@}\sqrt{@n@}\) metres and width \({@b@}\sqrt{@m@}\) metres.</p>
<p>Find the exact area of the rug, in m². Give your answer in the form \(k\sqrt{m}\).</p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p><strong>Step 1:</strong> Area = length × width.</p>
\[ \text{Area} = {@a@}\sqrt{@n@} \times {@b@}\sqrt{@m@} \]

<p><strong>Step 2:</strong> Multiply the numbers outside the roots.</p>
\[ {@a@} \times {@b@} = {@ab@} \]

<p><strong>Step 3:</strong> Multiply the numbers inside the roots.</p>
\[ \sqrt{@n@} \times \sqrt{@m@} = \sqrt{{@n@} \times {@m@}} = \sqrt{@nm@} \]

<p><strong>Answer:</strong> \( {@ab@}\sqrt{@nm@} \) m²</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>simp:true;
P: [2,3,5,7,11,13,17,19];
n: rand(P);
m: rand(delete(n,P));
a: rand(makelist(i,i,2,9));
b: rand(makelist(i,i,2,9));
ab: a*b;
nm: n*m;
ta1: ab*sqrt(nm);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}{@a@}\sqrt{@n@} \times {@b@}\sqrt{@m@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1062685336</deployedseed>
    <deployedseed>499156163</deployedseed>
    <deployedseed>1427468834</deployedseed>
    <deployedseed>817349467</deployedseed>
    <deployedseed>1385093124</deployedseed>
    <deployedseed>1041621634</deployedseed>
    <deployedseed>1785695805</deployedseed>
    <deployedseed>522645794</deployedseed>
    <deployedseed>463542579</deployedseed>
    <deployedseed>745355613</deployedseed>
    <deployedseed>690856201</deployedseed>
    <deployedseed>870586025</deployedseed>
    <deployedseed>758278013</deployedseed>
    <deployedseed>33725779</deployedseed>
    <deployedseed>1032472748</deployedseed>
    <deployedseed>1744722991</deployedseed>
    <deployedseed>342891994</deployedseed>
    <deployedseed>1200843156</deployedseed>
    <deployedseed>1277791604</deployedseed>
    <deployedseed>561952223</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
