<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>fencing a square garden</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A square garden has an area of \({@x@}\) m². Fencing is needed around the whole garden.</p>
<p>Find the exact length of fencing needed, in metres. Give your answer in the form \(k\sqrt{m}\), where \(m\) is as small as possible.</p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p><strong>Step 1:</strong> Find the side length. The side of a square is the square root of its area.</p>
\[ \text{side} = \sqrt{@x@} \]

<p><strong>Step 2:</strong> Simplify the root.</p>
\[ \sqrt{@x@} = \sqrt{{@a2@} \times {@n@}} = \sqrt{@a2@} \times \sqrt{@n@} = {@a@}\sqrt{@n@} \]

<p><strong>Step 3:</strong> The fence goes around all 4 sides, so multiply by 4.</p>
\[ 4 \times {@a@}\sqrt{@n@} = {@c@}\sqrt{@n@} \]

<p><strong>Answer:</strong> \( {@c@}\sqrt{@n@} \) metres</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>simp:true;
S: [2,3,5,6,7,10,11,13,14,15,17,19];
n: rand(S);
a: rand(makelist(i,i,2,12));
a2: a^2;
x: a2*n;
c: 4*a;
ta1: c*sqrt(n);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}{@x@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>415455306</deployedseed>
    <deployedseed>1422930189</deployedseed>
    <deployedseed>113148853</deployedseed>
    <deployedseed>814404143</deployedseed>
    <deployedseed>179166329</deployedseed>
    <deployedseed>226397834</deployedseed>
    <deployedseed>475448720</deployedseed>
    <deployedseed>1807396010</deployedseed>
    <deployedseed>1022542102</deployedseed>
    <deployedseed>280113982</deployedseed>
    <deployedseed>485721944</deployedseed>
    <deployedseed>1842442493</deployedseed>
    <deployedseed>1546400797</deployedseed>
    <deployedseed>1057498574</deployedseed>
    <deployedseed>1371223358</deployedseed>
    <deployedseed>1709904574</deployedseed>
    <deployedseed>664491487</deployedseed>
    <deployedseed>909183458</deployedseed>
    <deployedseed>1461305511</deployedseed>
    <deployedseed>1883772590</deployedseed>
  </question>
</quiz>
