<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Derive the quadratic formula — fill in the missing steps</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Work through the derivation of the quadratic formula from \(a x^2 + b x + c = 0\) (with \(a \neq 0\)) by completing the square. Fill in each missing step. Give <strong>exact</strong> expressions in terms of \(a\), \(b\) and \(c\) (no decimals).</p>
<p>Dividing every term by \(a\) gives \[ x^2 + \frac{b}{a}x + \frac{c}{a} = 0. \]</p>

<p><strong>Step 1.</strong> Move the constant term to the right-hand side.</p>
<div>\( x^2 + \frac{b}{a}x = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>

<p><strong>Step 2.</strong> To complete the square, what single term must be added to <em>both</em> sides?</p>
<div>Add \(\) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</div>

<p><strong>Step 3.</strong> The left-hand side is now \(\left(x+\frac{b}{2a}\right)^2\). Write the right-hand side as a single fraction over \(4a^2\).</p>
<div>\( \left(x+\frac{b}{2a}\right)^2 = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</div>

<p><strong>Step 4.</strong> Take square roots and solve for \(x\). Give the two solutions as a set, e.g. \(\{\,\alpha,\ \beta\,\}\).</p>
<div>\( x \in \) [[input:ans4]] [[validation:ans4]] [[feedback:prt4]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Completing the square on the general equation gives the quadratic formula:</p>
\[
\begin{align*}
a x^2 + b x + c &= 0\\
x^2 + \tfrac{b}{a}x + \tfrac{c}{a} &= 0 &&\text{(divide by } a\text{)}\\
x^2 + \tfrac{b}{a}x &= -\tfrac{c}{a} &&\textbf{(Step 1)}\\
x^2 + \tfrac{b}{a}x + \left(\tfrac{b}{2a}\right)^2 &= -\tfrac{c}{a} + \tfrac{b^2}{4a^2} &&\textbf{(Step 2: add } \tfrac{b^2}{4a^2}\text{)}\\
\left(x+\tfrac{b}{2a}\right)^2 &= \tfrac{b^2-4ac}{4a^2} &&\textbf{(Step 3)}\\
x + \tfrac{b}{2a} &= \pm\tfrac{\sqrt{b^2-4ac}}{2a}\\
x &= \dfrac{-b \pm \sqrt{b^2-4ac}}{2a} &&\textbf{(Step 4)}
\end{align*}
\]
<p>The two solutions are \(\left\{\dfrac{-b-\sqrt{b^2-4ac}}{2a},\ \dfrac{-b+\sqrt{b^2-4ac}}{2a}\right\}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>4</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025010100</text>
    </stackversion>
    <questionvariables>
      <text>/* Symbolic derivation: a, b, c stay as free variables (the general result). */
ta1 : -c/a;                       /* Step 1: RHS after moving the constant */
ta2 : b^2/(4*a^2);                /* Step 2: term added to complete the square */
ta3 : (b^2-4*a*c)/(4*a^2);        /* Step 3: simplified RHS as a single fraction */
ta4 : {(-b-sqrt(b^2-4*a*c))/(2*a), (-b+sqrt(b^2-4*a*c))/(2*a)};  /* Step 4: solution set */
/* Common misconceptions, used to give targeted feedback. */
w1_sign   : c/a;                  /* forgot to negate the moved constant */
w2_half   : b/(2*a);              /* took half but forgot to square */
w2_nohalf : (b/a)^2;              /* squared b/a without halving first */
w3_4c     : (b^2-4*c)/(4*a^2);    /* used 4c instead of 4ac in the numerator */
w4_denom  : {(-b+sqrt(b^2-4*a*c)/(2*a)), (-b-sqrt(b^2-4*a*c)/(2*a))};  /* only the root divided by 2a */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>Symbolic quadratic-formula derivation (fixed variant).</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>14</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans4</name>
      <type>algebraic</type>
      <tans>ta4</tans>
      <boxsize>30</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct RHS after moving the constant?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Sign-error diagnostic: kept the constant positive.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>w1_sign</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-sign</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Watch the sign. Moving \(\frac{c}{a}\) across the equals sign changes it to \(-\frac{c}{a}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Move only the constant \(\frac{c}{a}\) to the right-hand side; the \(x\)-terms stay on the left.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct completing-the-square term?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Forgot to square (gave half the coefficient).</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>w2_half</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-half</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is <em>half</em> the coefficient of \(x\). Completing the square needs its <em>square</em>: \(\left(\frac{b}{2a}\right)^2 = \frac{b^2}{4a^2}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Squared b/a without halving first.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>w2_nohalf</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-nohalf</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You squared \(\frac{b}{a}\) directly. First take <em>half</em> of \(\frac{b}{a}\), then square: \(\left(\frac{b}{2a}\right)^2\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Take half of the coefficient of \(x\) (which is \(\frac{b}{a}\)) and square the result.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct combined right-hand side?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>4c instead of 4ac in the numerator.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>w3_4c</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-4c</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Almost. To get the common denominator \(4a^2\), rewrite \(\frac{c}{a}=\frac{4ac}{4a^2}\). The numerator is \(b^2-4ac\), not \(b^2-4c\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Combine \(\frac{b^2}{4a^2}-\frac{c}{a}\) over the common denominator \(4a^2\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt4</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct solution set?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt4-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Is the answer even a set (both roots given)?</description>
        <answertest>SameType</answertest>
        <sans>ans4</sans>
        <tans>ta4</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>2</truenextnode>
        <trueanswernote>prt4-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-single</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>A quadratic has two solutions (the \(\pm\)). Enter <em>both</em> as a set, e.g. \(\{\alpha,\ \beta\}\).</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Bracketing error: only the square root divided by 2a.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans4</sans>
        <tans>w4_denom</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt4-denom</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The <em>whole</em> numerator \(-b\pm\sqrt{b^2-4ac}\) must be divided by \(2a\). Use brackets: \((-b\pm\sqrt{\ldots})/(2a)\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt4-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Solve \(x+\frac{b}{2a}=\pm\frac{\sqrt{b^2-4ac}}{2a}\) for \(x\) and combine each root over \(2a\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description>All four steps correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt4</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Step 1 sign error</description>
      <testinput>
        <name>ans1</name>
        <value>w1_sign</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-sign</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Step 2 forgot to square</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>w2_half</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-half</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Step 2 squared without halving</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>w2_nohalf</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-nohalf</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Step 3 uses 4c instead of 4ac</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>w3_4c</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>ta4</value>
      </testinput>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-4c</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description>Step 4 gives only one root (not a set)</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>(-b+sqrt(b^2-4*a*c))/(2*a)</value>
      </testinput>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-single</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>7</testcase>
      <description>Step 4 bracketing error (only the root divided by 2a)</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <testinput>
        <name>ans4</name>
        <value>w4_denom</value>
      </testinput>
      <expected>
        <name>prt4</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt4-denom</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
