<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>dividing irratinal numbers</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Simplify the following expression, giving your answer in the form \(k\sqrt{m}\).</p>
\[ \frac{{@p@}\sqrt{@nm@}}{{@q@}\sqrt{@m@}} \]
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p><strong>Step 1:</strong> Divide the numbers outside the roots.</p>
\[ {@p@} \div {@q@} = {@r@} \]

<p><strong>Step 2:</strong> Divide the roots.</p>
\[ \frac{\sqrt{@nm@}}{\sqrt{@m@}} = \sqrt{\frac{@nm@}{@m@}} = \sqrt{@n@} \]

<p><strong>Step 3:</strong> Put the two results together.</p>
\[ \frac{{@p@}\sqrt{@nm@}}{{@q@}\sqrt{@m@}} = {@r@}\sqrt{@n@} \]

<p><strong>Answer:</strong> \( {@r@}\sqrt{@n@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>simp:true;
P: [2,3,5,7,11,13];
n: rand(P);
m: rand(delete(n,P));
q: rand(makelist(i,i,2,10));
r: rand(makelist(i,i,2,20));
p: q*r;
nm: n*m;
ta1: r*sqrt(n);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}\frac{{@p@}\sqrt{@nm@}}{{@q@}\sqrt{@m@}}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1833881167</deployedseed>
    <deployedseed>579658153</deployedseed>
    <deployedseed>10742014</deployedseed>
    <deployedseed>1053659964</deployedseed>
    <deployedseed>1668049210</deployedseed>
    <deployedseed>1864251785</deployedseed>
    <deployedseed>2013972711</deployedseed>
    <deployedseed>494105406</deployedseed>
    <deployedseed>2034959698</deployedseed>
    <deployedseed>1621486142</deployedseed>
    <deployedseed>926439713</deployedseed>
    <deployedseed>1597725557</deployedseed>
    <deployedseed>1059474578</deployedseed>
    <deployedseed>1363140893</deployedseed>
    <deployedseed>1227007880</deployedseed>
    <deployedseed>1193002297</deployedseed>
    <deployedseed>1791788052</deployedseed>
    <deployedseed>332442559</deployedseed>
    <deployedseed>632961237</deployedseed>
    <deployedseed>1837700123</deployedseed>
  </question>
</quiz>
