<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>multiplying and dividing irrationals</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Simplify the following expression, giving your answer in the form \(k\sqrt{m}\).</p>
\[ \frac{{@a@}\sqrt{@n@} \times {@b@}\sqrt{@m@}}{{@c@}\sqrt{@m@}} \]
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p><strong>Step 1:</strong> Multiply the top of the fraction. Multiply the numbers outside the roots, then the numbers inside the roots.</p>
\[ {@a@} \times {@b@} = {@ab@}, \qquad \sqrt{@n@} \times \sqrt{@m@} = \sqrt{@nm@} \]
\[ {@a@}\sqrt{@n@} \times {@b@}\sqrt{@m@} = {@ab@}\sqrt{@nm@} \]

<p><strong>Step 2:</strong> The fraction is now:</p>
\[ \frac{{@ab@}\sqrt{@nm@}}{{@c@}\sqrt{@m@}} \]

<p><strong>Step 3:</strong> Divide the numbers outside the roots.</p>
\[ {@ab@} \div {@c@} = {@e@} \]

<p><strong>Step 4:</strong> Divide the roots.</p>
\[ \frac{\sqrt{@nm@}}{\sqrt{@m@}} = \sqrt{\frac{@nm@}{@m@}} = \sqrt{@n@} \]

<p><strong>Step 5:</strong> Put the two results together.</p>
\[ \frac{{@a@}\sqrt{@n@} \times {@b@}\sqrt{@m@}}{{@c@}\sqrt{@m@}} = {@e@}\sqrt{@n@} \]

<p><strong>Answer:</strong> \( {@e@}\sqrt{@n@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>simp:true;
P: [2,3,5,7,11];
n: rand(P);
m: rand(delete(n,P));
c: rand(makelist(i,i,2,8));
k: rand(makelist(i,i,2,10));
b: rand(makelist(i,i,2,20));
a: c*k;
ab: a*b;
nm: n*m;
e: ab/c;
ta1: e*sqrt(n);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}{@a@}, {@b@}, {@c@}, {@n@}, {@m@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1510692815</deployedseed>
    <deployedseed>615194241</deployedseed>
    <deployedseed>861081153</deployedseed>
    <deployedseed>1072491969</deployedseed>
    <deployedseed>1626682607</deployedseed>
    <deployedseed>1401270568</deployedseed>
    <deployedseed>1021674426</deployedseed>
    <deployedseed>1939942513</deployedseed>
    <deployedseed>1951120087</deployedseed>
    <deployedseed>922908880</deployedseed>
    <deployedseed>1655453660</deployedseed>
    <deployedseed>1510915732</deployedseed>
    <deployedseed>1464384152</deployedseed>
    <deployedseed>1005304248</deployedseed>
    <deployedseed>1924601410</deployedseed>
    <deployedseed>924390372</deployedseed>
    <deployedseed>345065506</deployedseed>
    <deployedseed>1061259845</deployedseed>
    <deployedseed>860453759</deployedseed>
    <deployedseed>444357880</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
