<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>multiplying irrational numbers</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Simplify the following expression, giving your answer in the form \(k\sqrt{m}\).</p>
\[ {@a@}\sqrt{@n@} \times {@b@}\sqrt{@m@} \]
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p><strong>Step 1:</strong> Multiply the numbers outside the roots.</p>
\[ {@a@} \times {@b@} = {@ab@} \]

<p><strong>Step 2:</strong> Multiply the numbers inside the roots.</p>
\[ \sqrt{@n@} \times \sqrt{@m@} = \sqrt{{@n@} \times {@m@}} = \sqrt{@nm@} \]

<p><strong>Step 3:</strong> Put the two results together.</p>
\[ {@a@}\sqrt{@n@} \times {@b@}\sqrt{@m@} = {@ab@}\sqrt{@nm@} \]

<p><strong>Answer:</strong> \( {@ab@}\sqrt{@nm@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>simp:true;
P: [2,3,5,7,11,13,17,19];
n: rand(P);
m: rand(delete(n,P));
a: rand(makelist(i,i,2,20));
b: rand(makelist(i,i,2,20));
ab: a*b;
nm: n*m;
ta1: ab*sqrt(nm);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}{@a@}{@b@}{@m@}{@n@}{@p@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>641830320</deployedseed>
    <deployedseed>1575076078</deployedseed>
    <deployedseed>652607987</deployedseed>
    <deployedseed>847420907</deployedseed>
    <deployedseed>650004002</deployedseed>
    <deployedseed>1204834184</deployedseed>
    <deployedseed>176996641</deployedseed>
    <deployedseed>1021089069</deployedseed>
    <deployedseed>2101013916</deployedseed>
    <deployedseed>452884096</deployedseed>
    <deployedseed>1487994643</deployedseed>
    <deployedseed>1015762902</deployedseed>
    <deployedseed>2028007341</deployedseed>
    <deployedseed>283255374</deployedseed>
    <deployedseed>760322867</deployedseed>
    <deployedseed>540612554</deployedseed>
    <deployedseed>365002549</deployedseed>
    <deployedseed>244510382</deployedseed>
    <deployedseed>691295480</deployedseed>
    <deployedseed>156563909</deployedseed>
  </question>
</quiz>
