<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Circumradius of a right-angled triangle (INNODEMS)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The triangle \(ABC\) below is right-angled at \(B\), with legs of \({@lega@}\) cm and \({@legb@}\) cm. The circle through its three vertices is the <strong>circumscribed circle</strong>, with centre \(O\).</p>
<p>Find the radius \(R\) of this circle, in centimetres. (The hypotenuse is not given &mdash; you will need it first.)</p>
[[jsxgraph width="380px" height="380px"]]
JXG.Options.text.useMathJax = true;
var lega = {#lega#};
var legb = {#legb#};
var hyp = {#hyp#};
var s = 6.0/hyp;
var B = [0,0], C = [lega*s,0], A = [0,legb*s];
var O = [lega*s/2, legb*s/2], Rd = hyp*s/2, pad = 0.9;
var board = JXG.JSXGraph.initBoard(divid, {boundingbox:[O[0]-Rd-pad, O[1]+Rd+pad, O[0]+Rd+pad, O[1]-Rd-pad], keepaspectratio:true, axis:false, showCopyright:false, showNavigation:false, pan:{enabled:false}, zoom:{enabled:false}});
var pO = board.create('point', O, {name:'O', size:1.5, fixed:true, fillColor:'#0d8e8e', strokeColor:'#0d8e8e', label:{offset:[8,6], anchorX:'left', anchorY:'bottom'}});
board.create('circle', [pO, Rd], {strokeColor:'#0d8e8e', strokeWidth:2, fillColor:'#0d8e8e', fillOpacity:0.04, highlight:false});
var pA = board.create('point', A, {name:'A', size:1.5, fixed:true, fillColor:'#333', strokeColor:'#333', label:{offset:[-9,6], anchorX:'right', anchorY:'bottom'}});
var pB = board.create('point', B, {name:'B', size:1.5, fixed:true, fillColor:'#333', strokeColor:'#333', label:{offset:[-9,-8], anchorX:'right', anchorY:'top'}});
var pC = board.create('point', C, {name:'C', size:1.5, fixed:true, fillColor:'#333', strokeColor:'#333', label:{offset:[9,-8], anchorX:'left', anchorY:'top'}});
board.create('polygon', [pA,pB,pC], {borders:{strokeColor:'#1565c0', strokeWidth:2, highlight:false}, fillColor:'#1565c0', fillOpacity:0.05, highlight:false, vertices:{visible:false}});
board.create('angle', [pA,pB,pC], {type:'square', radius:0.5, strokeColor:'#333', fillColor:'#333', fillOpacity:0.12, name:'', label:{visible:false}, highlight:false});
board.create('text', [C[0]/2, -0.3, lega+' cm'], {anchorX:'middle', anchorY:'top', fontSize:14, strokeColor:'#333', fixed:true});
board.create('text', [-0.3, A[1]/2, legb+' cm'], {anchorX:'right', anchorY:'middle', fontSize:14, strokeColor:'#333', fixed:true});
[[/jsxgraph]]
<div>\(R = \) [[input:ans1]] cm [[validation:ans1]] [[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>In a right-angled triangle the hypotenuse is a <strong>diameter</strong> of the circumscribed circle (the right angle is an angle in a semicircle). So the centre \(O\) is the midpoint of the hypotenuse \(AC\), and the circumradius is half the hypotenuse.</p>
[[jsxgraph width="380px" height="380px"]]
JXG.Options.text.useMathJax = true;
var lega = {#lega#};
var legb = {#legb#};
var hyp = {#hyp#};
var s = 6.0/hyp;
var B = [0,0], C = [lega*s,0], A = [0,legb*s];
var O = [lega*s/2, legb*s/2], Rd = hyp*s/2, pad = 0.9;
var board = JXG.JSXGraph.initBoard(divid, {boundingbox:[O[0]-Rd-pad, O[1]+Rd+pad, O[0]+Rd+pad, O[1]-Rd-pad], keepaspectratio:true, axis:false, showCopyright:false, showNavigation:false, pan:{enabled:false}, zoom:{enabled:false}});
var pO = board.create('point', O, {name:'O', size:1.5, fixed:true, fillColor:'#0d8e8e', strokeColor:'#0d8e8e', label:{offset:[-10,-10], anchorX:'right', anchorY:'top'}});
board.create('circle', [pO, Rd], {strokeColor:'#0d8e8e', strokeWidth:2, fillColor:'#0d8e8e', fillOpacity:0.04, highlight:false});
var pA = board.create('point', A, {name:'A', size:1.5, fixed:true, fillColor:'#333', strokeColor:'#333', label:{offset:[-9,6], anchorX:'right', anchorY:'bottom'}});
var pB = board.create('point', B, {name:'B', size:1.5, fixed:true, fillColor:'#333', strokeColor:'#333', label:{offset:[-9,-8], anchorX:'right', anchorY:'top'}});
var pC = board.create('point', C, {name:'C', size:1.5, fixed:true, fillColor:'#333', strokeColor:'#333', label:{offset:[9,-8], anchorX:'left', anchorY:'top'}});
board.create('polygon', [pA,pB,pC], {borders:{strokeColor:'#1565c0', strokeWidth:2, highlight:false}, fillColor:'#1565c0', fillOpacity:0.05, highlight:false, vertices:{visible:false}});
board.create('angle', [pA,pB,pC], {type:'square', radius:0.5, strokeColor:'#333', fillColor:'#333', fillOpacity:0.12, name:'', label:{visible:false}, highlight:false});
board.create('text', [C[0]/2, -0.3, lega+' cm'], {anchorX:'middle', anchorY:'top', fontSize:14, strokeColor:'#333', fixed:true});
board.create('text', [-0.3, A[1]/2, legb+' cm'], {anchorX:'right', anchorY:'middle', fontSize:14, strokeColor:'#333', fixed:true});
board.create('segment', [pO,pA], {strokeColor:'#0d8e8e', strokeWidth:1.6, dash:2, highlight:false});
board.create('text', [(O[0]+A[0])/2-0.15, (O[1]+A[1])/2+0.15, 'R'], {anchorX:'right', anchorY:'bottom', fontSize:14, strokeColor:'#0d8e8e', fixed:true});
var Gx=(A[0]+B[0]+C[0])/3, Gy=(A[1]+B[1]+C[1])/3;
var mx=(A[0]+C[0])/2, my=(A[1]+C[1])/2, dx=mx-Gx, dy=my-Gy, L=Math.sqrt(dx*dx+dy*dy)||1;
board.create('text', [mx+0.6*dx/L, my+0.6*dy/L, hyp+' cm'], {anchorX:'middle', anchorY:'middle', fontSize:14, strokeColor:'#0d8e8e', fixed:true});
[[/jsxgraph]]
<p>First find the hypotenuse \(AC\) by Pythagoras:</p>
\[ AC = \sqrt{ {@lega@}^2 + {@legb@}^2 } = \sqrt{ {@lega^2 + legb^2@} } = {@hyp@}\text{ cm}. \]
<p>The circumradius is half of this:</p>
\[ R = \frac{AC}{2} = \frac{ {@hyp@} }{2} = {@ta1@}\text{ cm}. \]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025010100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[mlist : [[2,1],[3,2],[4,1],[4,3]];
cands : [];
for mn in mlist do
  for kk : 1 thru 5 do (
    l1 : kk*(mn[1]^2 - mn[2]^2),
    l2 : kk*(2*mn[1]*mn[2]),
    hh : kk*(mn[1]^2 + mn[2]^2),
    if max(l1, l2, hh) <= 30
      then cands : endcons([max(l1,l2), min(l1,l2), hh], cands)
  );
pick : rand(cands);
lega : pick[1];                 /* longer leg  (horizontal) */
legb : pick[2];                 /* shorter leg (vertical)   */
hyp  : pick[3];                 /* hypotenuse (a whole number, NOT shown) */
ta1  : hyp/2;                   /* model answer: circumradius = half the hypotenuse */
wrong_hyp : hyp;               /* gave the hypotenuse = diameter, forgot to halve */
wrong_addlegs : (lega+legb)/2; /* used hypotenuse = lega+legb instead of Pythagoras */]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( {@lega@},\ {@legb@}\ \text{cm};\ h={@hyp@},\ R={@ta1@}\ \text{cm} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>4</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct circumradius?</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.05</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Gave the hypotenuse (diameter), forgot to halve?</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>wrong_hyp</tans>
        <testoptions>0.05</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-diam</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That looks like the length of the hypotenuse, not the circumradius. What's the relationship between the two?</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Used hypotenuse = lega + legb (Pythagoras slip)?</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>wrong_addlegs</tans>
        <testoptions>0.05</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-addlegs</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you used \(a+b\) for the hypotenuse. The hypotenuse is \(\sqrt{a^2+b^2}\), not \(a+b\) &mdash; recompute it with Pythagoras, then take half.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-incorrect</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not correct. Find the hypotenuse with Pythagoras, \(AC=\sqrt{a^2+b^2}\), then halve it: \(R=\tfrac12 AC\). See the worked solution after this attempt.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>4001</deployedseed>
    <deployedseed>4002</deployedseed>
    <deployedseed>4003</deployedseed>
    <deployedseed>4004</deployedseed>
    <deployedseed>4005</deployedseed>
    <deployedseed>4006</deployedseed>
    <deployedseed>4007</deployedseed>
    <deployedseed>4008</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct circumradius</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Gave the hypotenuse (forgot to halve)</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_hyp</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-diam</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Added the legs for the hypotenuse</description>
      <testinput>
        <name>ans1</name>
        <value>ev(float(wrong_addlegs),simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-addlegs</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Generic wrong value</description>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+0.3,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-incorrect</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
