<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Even and Odd Functions — combining two functions of known parity</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Suppose \(f\) is an <strong>{@f_parity@}</strong> function and \(g\) is an <strong>{@g_parity@}</strong> function (both defined for all real numbers).</p>
<p>What can you say about the function \(h(x) = f(x)\ {@opsymbol@}\ g(x)\)?</p>
<div>[[input:ans1]] [[validation:ans1]]</div>
<div>[[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Rules for combining even (E) and odd (O) functions:</p>
<ul>
<li>Sum: E+E=E, O+O=O, E+O=Neither (in general)</li>
<li>Product: E&times;E=E, O&times;O=E, E&times;O=O</li>
</ul>
<p>Here \(f\) is {@f_parity@} and \(g\) is {@g_parity@}, combined with {@opdesc@}, so \(h\) is {@ta@}.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[f_is_even : is(rand(2)=1);
g_is_even : is(rand(2)=1);
op : rand(2);                     /* 0 = sum, 1 = product */
f_parity : if f_is_even then "even" else "odd";
g_parity : if g_is_even then "even" else "odd";
opsymbol : if op=0 then "+" else "&#215;";
opdesc : if op=0 then "addition" else "multiplication";
/* rules: even+even=even, odd+odd=odd, even+odd=neither (in general);
          even*even=even, odd*odd=even, even*odd=odd */
ta : if op=0 then (if f_is_even and g_is_even then "even" elseif (not f_is_even) and (not g_is_even) then "odd" else "neither")
     else (if f_is_even and g_is_even then "even" elseif (not f_is_even) and (not g_is_even) then "even" else "odd");
opts : [["even", ta="even", "Always even"],["odd", ta="odd", "Always odd"],["neither", ta="neither", "Could be neither (not necessarily even or odd)"]];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.6 (Even and Odd Functions). Applies the rules for the parity of a sum or product of an even/odd function pair, reinforcing the Learner Experience 16.6.2 exploration rather than d</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.6 (Even and Odd Functions). Applies the rules for the parity of a sum or product of an even/odd function pair, reinforcing the Learner Experience 16.6.2 exploration rather than direct substitution.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>opts</tans>
      <boxsize>25</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correctly applied the parity combination rule?</description>
        <answertest>String</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>Try the substitution test directly: compute h(-x) in terms of f(-x) and g(-x), using the parities of f and g, and compare it with h(x) and -h(x).</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
