<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Even and Odd Functions — reading symmetry from a graph (JSXGraph)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The graph of \(y=k(x)\) is shown below.</p>
[[jsxgraph width="420px" height="360px"]]
var scenario = {#scenario#};
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-5, 8, 5, -8],
    axis: true, showCopyright: false, showNavigation: false
});
var f;
if (scenario === 0) { f = board.jc.snippet('0.3*x^2+1', true, 'x', true); }        /* even: symmetric about y-axis */
else if (scenario === 1) { f = board.jc.snippet('0.1*x^3-1.5*x', true, 'x', true); } /* odd: symmetric about the origin */
else { f = board.jc.snippet('0.2*x^2+0.6*x', true, 'x', true); }                    /* neither */
board.create('functiongraph', [f, -4.5, 4.5], {strokeColor: '#1565c0', strokeWidth: 2});
[[/jsxgraph]]
<p>Based on the symmetry of the graph, is \(k(x)\) even, odd, or neither?</p>
<div>[[input:ans1]] [[validation:ans1]]</div>
<div>[[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>An <strong>even</strong> function's graph is a mirror image of itself across the \(y\)-axis. An <strong>odd</strong> function's graph looks the same after a \(180^\circ\) rotation about the origin. If neither of these holds, the function is neither.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[scenario : rand(3);
ta : if scenario=0 then "even" elseif scenario=1 then "odd" else "neither";
opts : [["even", scenario=0, "Even (symmetric about the y-axis)"],["odd", scenario=1, "Odd (symmetric about the origin)"],["neither", scenario=2, "Neither"]];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.6 (Even and Odd Functions). Classifies a function as even, odd, or neither purely from the visual symmetry of its JSXGraph-plotted graph, complementing the algebraic-substitution</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.6 (Even and Odd Functions). Classifies a function as even, odd, or neither purely from the visual symmetry of its JSXGraph-plotted graph, complementing the algebraic-substitution version of this question.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>opts</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correctly read the symmetry of the graph?</description>
        <answertest>String</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>Check carefully: does folding the graph along the y-axis map it onto itself (even)? Does rotating it 180 degrees about the origin map it onto itself (odd)? Or neither?</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
