<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Errors in real life Q1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A site engineer is mixing concrete for a foundation slab. The volume of cement used is measured as \({@cement@} \;\text{m}^3\) and the volume of sand as \({@sand@} \;\text{m}^3\), both correct to 1 decimal place. The sand-to-cement ratio is calculated by dividing the volume of sand by the volume of cement. This figure for ensuring the concrete meets strength specifications.</p>

<p>(a) Find the absolute error in the calculated sand-to-cement ratio.</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>

<p>(b) Find the relative error in the ratio.</p>
<p>[[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</p>

<p>(c) Find the percentage error in the ratio.</p>
<p>[[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3> Worked Solution.</h3>
<p>The volume of cement was recorded as \({@cement@}\;\text{m}^3\) and the volume of sand as \({@sand@}\;\text{m}^3\), both correct to 1 decimal place.</p>

<p><strong>Absolute errors in the individual measurements:</strong></p>
<p>Since each measurement is correct to 1 decimal place, the absolute error in each is half of the smallest unit (\(0.1\;\text{m}^3\)):</p>
<p>\[ \text{Cement error} = \frac{0.1}{2} = {@cement_err@}\; \text{ m}^3, \qquad \text{Sand error} = \frac{0.1}{2} = {@sand_err@}\; \text{ m}^3 \]</p>

<p><strong>Important:</strong> When dividing two measured quantities, the <em>maximum</em> possible value of the result comes from pairing the <em>largest</em> possible numerator with the <em>smallest</em> possible denominator (since dividing by a smaller number gives a larger result). So:</p>
<p>\[ \text{Max sand} = {@sand@} + {@sand_err@} = {@max_sand@}\; \text{ m}^3, \qquad \text{Min cement} = {@cement@} - {@cement_err@} = {@min_cement@}\; \text{ m}^3 \]</p>

<p>(a) <strong>Absolute error in the ratio:</strong></p>
<p>Using the recorded measurements:</p>
<p>\[ \text{Working ratio} = \frac{{@sand@}}{{@cement@}} = {@ratio@} \]</p>
<p>Using the maximum sand and minimum cement:</p>
<p>\[ \text{Maximum ratio} = \frac{{@max_sand@}}{{@min_cement@}} = {@max_ratio@} \]</p>
<p>\[ \text{Absolute error} = \text{Maximum ratio} - \text{Working ratio} = {@max_ratio@} - {@ratio@} = {@ta1@} \]</p>

<p>(b) <strong>Relative error:</strong></p>
<p>\[ \text{Relative error} = \frac{\text{Absolute error}}{\text{Working ratio}} = \frac{{@ta1@}}{{@ratio@}} \approx {@ta2@} \]</p>

<p>(c) <strong>Percentage error:</strong></p>
<p>\[ \text{Percentage error} = \text{Relative error} \times 100\% = {@ta2@} \times 100\% \approx {@ta3@}\% \]</p>

<p><strong>Why this matters on site:</strong> The sand-to-cement ratio directly affects the strength of the concrete. Too much sand relative to cement weakens the mix and can compromise structural integrity, while too little sand makes the mix overly rigid and increases material costs unnecessarily. Even small measurement errors, when compounded across a large pour, can lead to real differences in the final concrete strength, which is why engineers track these error margins carefully.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>cement_scale: 10;
cement_steps: rand(100*cement_scale/10 + 1);
cement: float(20 + cement_steps/cement_scale);
cement_err: float(0.5/cement_scale);

sand_scale: 10;
sand_steps: rand(200*sand_scale/10 + 1);
sand: float(50 + sand_steps/sand_scale);
sand_err: float(0.5/sand_scale);

max_cement: cement + cement_err;
min_cement: cement - cement_err;

max_sand: sand + sand_err;
min_sand: sand - sand_err;

ratio: float(sand / cement);
max_ratio: float(max_sand / min_cement);
min_ratio: float(min_sand / max_cement);

ta1: max_ratio - ratio;
ta2: ta1 / ratio;
ta3: ta2 * 100;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@} {@ta2@} {@ta3@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions>0.1</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>990552160</deployedseed>
    <deployedseed>1718196725</deployedseed>
    <deployedseed>930748714</deployedseed>
    <deployedseed>515045158</deployedseed>
    <deployedseed>823203403</deployedseed>
    <deployedseed>915233461</deployedseed>
    <deployedseed>209789995</deployedseed>
    <deployedseed>1465307131</deployedseed>
    <deployedseed>1914072638</deployedseed>
    <deployedseed>1716845717</deployedseed>
    <deployedseed>1923327154</deployedseed>
    <deployedseed>1263296727</deployedseed>
    <deployedseed>1335182440</deployedseed>
    <deployedseed>324604136</deployedseed>
    <deployedseed>1490567032</deployedseed>
    <deployedseed>1362640969</deployedseed>
    <deployedseed>1874856440</deployedseed>
    <deployedseed>1926206234</deployedseed>
    <deployedseed>880629989</deployedseed>
  </question>
</quiz>
