<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>irrational numbers 1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Simplify \(\sqrt{@number@}\) to the simplest form.</p>
<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>To simplify \( \sqrt{ {@number@} } \), look for the largest perfect-square factor of {@number@}.</p>
<p>Since {@number@} = {@a@} × {@b@}, and {@a@} is a perfect square:</p>
<p>\[ \sqrt{ {@number@} } = \sqrt{ {@a@} } \times \sqrt{ {@b@} } = {@sqrt(a)@}\sqrt{ {@b@} } \]</p>
<p><strong>Key idea:</strong> an irrational number is fully simplified when the number left under the root sign has no perfect-square factors remaining. Always check for the largest such factor first otherwise you may simplify only part-way.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>/* Perfect square factors to choose from */
squares: [4,9,16,25];
a: squares[rand(length(squares))+1];

/* Square-free number (2-7, excluding 4) to multiply with */
b: rand(6)+2;
while member(b, [4]) do b: rand(6)+2;

/* The number the student is asked to simplify */
number: a * b;

/* Correct fully-simplified answer, e.g. sqrt(16)*sqrt(3) = 4*sqrt(3) */
ta1: sqrt(a) * sqrt(b);

ta2: sqrt(b) * a;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}{@ta2@}{@a@}{@b@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You've correctly identified the factor pair {@a@} and {@b@}, but they're the wrong way round.</p>
<p>Remember: the <strong>perfect square</strong> ({@a@}) is the part that goes <em>under</em> the root first, so you can take its square root - it should end up <strong>outside</strong> the root sign as \( \sqrt{ {@a@} } \).</p>
<p>The <strong>square-free</strong> number ({@b@}) is the part that <strong>stays under</strong> the root, since it has no perfect-square factor left to simplify.</p>
<p>Try again: \( \sqrt{ {@number@} } = \sqrt{ {@a@} \times {@b@} } = \sqrt{ {@a@} } \times \sqrt{ {@b@} } \)</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>That's not correct. Check that you've found the <strong>largest</strong> perfect-square factor of {@number@}, and that your final answer is in the form \( a\sqrt{b} \), where \(\sqrt{@a@}\) is outside and {@b@} is inside the root.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1616688395</deployedseed>
    <deployedseed>1180168136</deployedseed>
    <deployedseed>141628609</deployedseed>
    <deployedseed>292926003</deployedseed>
    <deployedseed>1362244051</deployedseed>
    <deployedseed>1072472544</deployedseed>
    <deployedseed>1473479125</deployedseed>
    <deployedseed>1326476696</deployedseed>
    <deployedseed>1312653993</deployedseed>
    <deployedseed>296209971</deployedseed>
    <deployedseed>1903067699</deployedseed>
  </question>
</quiz>
