<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>quiz 2</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A school is preparing a square-shaped display area for a mathematics exhibition. The side length of the display area is \(\sqrt{@n@}\) cm.If the side length is doubled, express the new length in its simplest form.

<p>Your answer should be written in the form \(a\sqrt{b}\), where \(b\) has no perfect-square factors.</p><p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[To find the new length, first double the original side length.

If the original length is \(\sqrt{n}\) cm, then:

\[
\text{New length}=2\sqrt{n}.
\]

Next, look for a perfect-square factor inside \(n\). Take the square root of the perfect square outside the radical and leave the remaining factor inside the square root.

For example, if the original side length is \( \sqrt{@n@} cm \):
<p>For example, the original value is \( \sqrt{ {@n@} }\):</p>
<p>
\[
\begin{array}{rl}
\sqrt{ {@n@} } &= \sqrt{ {@(a^2)@} \times {@b@} } \\
               &= \sqrt{ {@(a^2)@} } \times \sqrt{ {@b@} } \\
               &= {@a@}\sqrt{ {@b@} } 
\end{array}
\]
</p>
Therefore, the final answer should have no perfect-square factors remaining inside the square root.]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>a: rand([2,3,4,5,6,7]);
b: rand([2,3,5,6,7]);
n: (a^2)*b;
ta1: a*sqrt(b);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@} {@a@} {@b@} {@n@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>8471973</deployedseed>
    <deployedseed>913292242</deployedseed>
    <deployedseed>1332708170</deployedseed>
    <deployedseed>979587873</deployedseed>
    <deployedseed>840186958</deployedseed>
    <deployedseed>435628352</deployedseed>
    <deployedseed>2113557173</deployedseed>
    <deployedseed>635360188</deployedseed>
    <deployedseed>600354468</deployedseed>
    <deployedseed>1086009465</deployedseed>
    <deployedseed>1412092851</deployedseed>
    <deployedseed>1973069650</deployedseed>
    <deployedseed>487038779</deployedseed>
    <deployedseed>3817980</deployedseed>
    <deployedseed>1176210072</deployedseed>
    <deployedseed>993207567</deployedseed>
    <deployedseed>1715103298</deployedseed>
  </question>
</quiz>
