<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating exact value of sine theta</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>
The angle \(\theta\) lies in Quadrant {@quadrant@} and satisfies (\cos(\theta)={@cos_value@}.\) Determine the exact value of \(\sin(\theta).\)
</p>

<p>
Give your answer in exact form.
</p>

<p>
[[input:ans1]]
</p>

[[validation:ans1]] [[feedback:prt1]]]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>
To find the value of \(\sin(\theta)\), use the Pythagorean identity
\(\sin^2(\theta)+\cos^2(\theta)=1\).
</p>

<p>
You are given \(\cos(\theta)={@cos_value@}\).
Therefore,
\(\sin^2(\theta)=1-\cos^2(\theta)\).
</p>

<p>
Substituting the given value gives
\(\sin^2(\theta)=1-({@cos_value@})^2\).
</p>

<p>
Taking the square root gives
\(\sin(\theta)=\pm\sqrt{1-({@cos_value@})^2}\).
</p>

<p>
The quadrant determines which sign is required.
Sine is positive in Quadrants I and II and negative in Quadrants III and IV.
</p>

<p>
Since \(\theta\) is in Quadrant {@quadrant@}, the appropriate sign must be selected.
Therefore,
\(\sin(\theta)={@ta@}\).
</p>

<p>
<strong>Remember:</strong> first find the magnitude using
\(\sin^2(\theta)+\cos^2(\theta)=1\), and then use the quadrant to determine the sign.
</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>/* Possible cosine values */
cos_values : [
    [3/5, 4/5],
    [4/5, 3/5],
    [5/13, 12/13],
    [12/13, 5/13]
];

/* Select a random cosine/reference pair */
variant : rand(cos_values);

cos_abs : variant[1];
sin_abs : variant[2];

/* Randomly select a quadrant: II or III */
quadrant : rand([2,3]);

/* Determine the signs */
if quadrant = 2 then (
    cos_value : -cos_abs,
    sin_value : sin_abs
) else (
    cos_value : -cos_abs,
    sin_value : -sin_abs
);

/* Correct answer */
ta : sin_value;

/* Student mistake 1:
   correct magnitude but wrong sign */
ta_mistake1 : -sin_value;

/* Student mistake 2:
   uses the cosine magnitude as the sine value */
ta_mistake2 : cos_abs;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Checks for correct answer</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Using the wrong sign</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta_mistake1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>
The magnitude of your answer is correct, but the sign is incorrect.
</p>

<p>
Remember that sine is positive in Quadrants I and II, and negative in
Quadrants III and IV.
Since \(\theta\) is in Quadrant {@quadrant@}, the sign of
\(\sin(\theta)\) must be chosen accordingly.
</p>

<p>
The correct value is \(\sin(\theta)={@ta@}\).
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Using cosine instead of sine</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans> ta_mistake2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>
You have used the magnitude of the given cosine value as the value of sine.
However, \(\sin(\theta)\) and \(\cos(\theta)\) are generally different.
</p>

<p>
Use the Pythagorean identity
\(\sin^2(\theta)+\cos^2(\theta)=1\)
to find the magnitude of \(\sin(\theta)\).
Then use the quadrant to determine its sign.
</p>

<p>
The correct value is \(\sin(\theta)={@ta@}\).
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
