<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Expanding an equation of the form  (x + n)^2</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Expand the expression \[ {@expr@}. \]</p>
<div>\( = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Multiply out and collect like terms, using \((a+b)^2 = a^2 + 2ab + b^2\):</p>
\[
\begin{align*}
{@expr@} &= (x+{@nn@})(x+{@nn@}) \\
         &= x^2 + 2\times{@nn@}\cdot x + {@nn^2@} \\
         &= {@ta1@}.
\end{align*}
\]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>nn : rand_with_prohib(-9,15,[-1,0,1]);   /* n=4 gives (x+4)^2 */
expr : (x+nn)^2;
ta1 : expand(expr);                     /* the fully expanded model answer */
wrong_foil : x^2 + nn^2;                /* squared each term, dropped the cross term */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( {@expr@} = {@ta1@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>expand</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Is the answer equivalent to the expansion?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>3</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Cross-term misconception: x^2 + n^2.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_foil</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-foil</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You squared each term but dropped the cross term \(2\cdot{@nn@}\cdot x\). Recall \((a+b)^2 = a^2 + 2ab + b^2\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Restated the original expression.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>expr</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-restate</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the original expression &mdash; now multiply it out.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrongvalue</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not equivalent to \({@expr@}\). See the worked solution.</p>]]></text>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Correct value: is it fully expanded?</description>
        <answertest>Expanded</answertest>
        <sans>ans1</sans>
        <tans>ans1</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-notexpanded</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>This is equivalent, but not fully expanded into separate terms. Write it in the form \(x^2 + \ldots\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>3001</deployedseed>
    <deployedseed>3002</deployedseed>
    <deployedseed>3003</deployedseed>
    <deployedseed>3004</deployedseed>
    <deployedseed>3005</deployedseed>
    <deployedseed>3006</deployedseed>
    <deployedseed>1808365208</deployedseed>
    <deployedseed>1145484455</deployedseed>
    <deployedseed>247248595</deployedseed>
    <deployedseed>53307089</deployedseed>
    <deployedseed>1313660966</deployedseed>
    <deployedseed>114907111</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct fully expanded answer</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Cross-term dropped</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_foil</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-foil</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Restated the question</description>
      <testinput>
        <name>ans1</name>
        <value>expr</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-0-T | prt1-notexpanded ]</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Equivalent but not expanded</description>
      <testinput>
        <name>ans1</name>
        <value>(x+nn)*(x+nn)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-notexpanded</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Wrong value (generic)</description>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-wrongvalue</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
