<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Perfect-square trinomial equation — recognise, complete the square, solve</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The height of a structure is modelled by the quadratic equation \[ {@expr@} = 0. \]</p>
<p><strong>a)</strong> Is the left-hand side a perfect-square trinomial?</p>
<div>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>
<p><strong>b)</strong> Rewrite the equation in the form \((x-a)^2 = 0\).</p>
<div>\( \) [[input:ans2]] \( = 0 \) [[validation:ans2]] [[feedback:prt2]]</div>
<p><strong>c)</strong> Hence find the value of \(x\).</p>
<div>\( x = \) [[input:ans3]] [[validation:ans3]] [[feedback:prt3]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The coefficient of \(x\) is \({@-2*aa@}\) and the constant is \({@aa^2@}\). Since \(\left(\tfrac{-2\cdot{@aa@}}{2}\right)^2 = {@aa^2@}\) matches the constant, the left-hand side is a perfect square.</p>
\[
\begin{align*}
{@expr@} &= 0\\
(x-{@aa@})^2 &= 0 &&\text{(complete the square)}\\
x-{@aa@} &= 0 &&\text{(square root of both sides)}\\
x &= {@aa@} &&\text{(a single repeated root)}
\end{align*}
\]]]></text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025010100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[aa : rand(8)+2;              /* a in 2..9; a=7 gives x^2-14x+49 */
expr : x^2 - 2*aa*x + aa^2;   /* always the perfect square (x-a)^2 */
/* Part (a): dropdown model answer [value, correct, display]. Always "Yes". */
ta1 : [[optyes, true, "Yes, it is a perfect-square trinomial."], [optno, false, "No, it is not a perfect-square trinomial."]];
/* Part (b): the completed square (expression to the left of = 0). */
ta2 : (x-aa)^2;
w2_sign  : (x+aa)^2;          /* wrong sign: used (x+a) */
w2_const : (x-aa^2)^2;        /* used the constant term as a */
/* Part (c): the repeated root. */
ta3 : aa;
w3_neg : -aa;                 /* sign slip */
w3_two : {aa, -aa};           /* treated it as two distinct roots */]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( a={@aa@}\); equation \({@expr@}=0\), root \(x={@aa@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>factor</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>6</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description><![CDATA[Did the student choose "Yes"?]]></description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>optyes</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Check again: a trinomial \(x^2+px+q\) is a perfect square when \(q=\left(\frac{p}{2}\right)^2\). Compare the square of half the coefficient of \(x\) with the constant term.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct value AND written as a single square (x-a)^2?</description>
        <answertest>CompSquare</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Right value but not written as a single square (e.g. expanded)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-notform</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the right expression, but write it as a single square \((x-a)^2\), not expanded or multiplied out.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Sign error: used (x+a) instead of (x-a)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>w2_sign</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-sign</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Check the sign. The middle term is negative, so the factor is \((x-a)\), giving \((x-a)^2\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Used the constant term itself as a?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>w2_const</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-const</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Here \(a\) is the <em>square root</em> of the constant term, not the constant itself. Choose \(a\) so that \(2a\) equals the size of the coefficient of \(x\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not equivalent to the left-hand side. Take \(a\) to be half the coefficient of \(x\), then write \((x-a)^2\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct repeated root x = a?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Sign slip: gave x = -a.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>w3_neg</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-sign</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Solving \(x-a=0\) gives \(x=a\). Check the sign — the repeated root is positive here.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Gave two distinct roots instead of one repeated root.</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>w3_two</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-two</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Because the left-hand side is a perfect square, there is a single repeated root \(x=a\), not two different values.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Set the squared factor equal to zero: \((x-a)=0\), so \(x=a\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>2001</deployedseed>
    <deployedseed>2002</deployedseed>
    <deployedseed>2003</deployedseed>
    <deployedseed>2004</deployedseed>
    <deployedseed>2005</deployedseed>
    <deployedseed>2006</deployedseed>
    <deployedseed>2007</deployedseed>
    <deployedseed>2008</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All three parts correct</description>
      <testinput>
        <name>ans1</name>
        <value>optyes</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Part (a) wrong (chose No)</description>
      <testinput>
        <name>ans1</name>
        <value>optno</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-wrong</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Part (b) correct value but expanded</description>
      <testinput>
        <name>ans1</name>
        <value>optyes</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>expr</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-notform</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Part (b) sign error (x+a)^2</description>
      <testinput>
        <name>ans1</name>
        <value>optyes</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>w2_sign</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-sign</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Part (b) used the constant as a</description>
      <testinput>
        <name>ans1</name>
        <value>optyes</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>w2_const</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-const</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description>Part (c) sign slip x = -a</description>
      <testinput>
        <name>ans1</name>
        <value>optyes</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>w3_neg</value>
      </testinput>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-sign</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>7</testcase>
      <description>Part (c) gave two roots</description>
      <testinput>
        <name>ans1</name>
        <value>optyes</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>w3_two</value>
      </testinput>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-two</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
