<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Perfect-square trinomial — garden area (dropdown + complete the square)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The area of a garden is modelled by the expression \[ x^2 + 2{@aa@}x +{@cc@}; \]</p>
<p><strong>a)</strong> Determine whether the expression is a perfect-square trinomial.</p>
<div>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>
[[ if test="is_ps" ]]
<p><strong>b)</strong> Write the expression in the form \((x+a)^2\).</p>
[[ else ]]
<p><strong>b)</strong> Write the expression in the form \((x+a)^2+k\).</p>
[[/ if ]]
<div>\(x^2 + 2{@aa@}x +{@cc@} = \) [[input:ans2]] [[validation:ans2]] [[feedback:prt2]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>A trinomial \(x^2+px+q\) is a <em>perfect square</em> exactly when the constant term equals the square of half the coefficient of \(x\), i.e. \(q=\left(\tfrac{p}{2}\right)^2\).</p>
<p>Here \(p={@2*aa@}\), so \(\left(\tfrac{p}{2}\right)^2={@aa^2@}\), while the constant term is \({@cc@}\).</p>
[[ if test="is_ps" ]]
<p>These are equal, so it <strong>is</strong> a perfect square, and  \[ x^2 + 2 \times {@aa@} \times x +{@cc@} \] =\[ (x+{@aa@})^2. \]</p>
[[ else ]]
<p>These are not equal, so it is <strong>not</strong> a perfect square. Completing the square: \[ x^2 + 2 {@aa@}x +{@cc@}\] \[= (x+{@aa@})^2 + ({@dd@})\] \[ = {@ta2@}. \]</p>
[[/ if ]]]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[aa : rand(6)+2;                       /* a in 2..7; a=6 gives x^2+12x+36 */
ps : rand(2);                         /* 0 or 1: is it a perfect square? */
dd : if ps=1 then 0 else (rand(3)+1)*(2*rand(2)-1);  /* nonzero offset when not a perfect square */
cc : aa^2 + dd;                       /* constant term */
expr : x^2 + 2*aa*x + cc;             /* the trinomial shown to the student */
is_ps : is(dd=0);                     /* true exactly when a perfect square */
/* Part (a): dropdown model answer as [value, correct, display] lists. */
ta1 : [[optyes, is_ps, "Yes, it is a perfect-square trinomial."], [optno, not is_ps, "No, it is not a perfect-square trinomial."]];
ta1v : if is_ps then optyes else optno;   /* the correct dropdown value */
ta1w : if is_ps then optno else optyes;   /* the wrong dropdown value */
/* Part (b): completed-square form; equals (x+a)^2 when it is a perfect square. */
ta2 : (x+aa)^2 + dd;
/* Misconception for (b): using the constant term itself as a. */
w2_const : (x+cc)^2;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( a={@aa@},\ c={@cc@},\ d={@dd@}\); perfect square: {@is_ps@}</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>16</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>2</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>factor</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Did the student choose the correct Yes/No option?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1v</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not quite. A trinomial \(x^2+px+q\) is a perfect square exactly when the constant term equals the square of half the coefficient of \(x\), i.e. \(q=\left(\frac{p}{2}\right)^2\). Compare \(\left(\frac{p}{2}\right)^2\) with the constant term.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct value AND written in completed-square form?</description>
        <answertest>CompSquare</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>x</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Right value but not in completed-square form (e.g. left expanded)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-notform</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is the right expression, but write it in completed-square form \((x+a)^2\) (plus a constant if needed), not expanded.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Used the constant term itself as a (e.g. (x+constant)^2)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>w2_const</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-const</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Here \(a\) is the <em>square root</em> of the constant, not the constant itself. For \(x^2+2ax+a^2=(x+a)^2\), choose \(a\) so that \(2a\) equals the coefficient of \(x\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not equivalent to the expression. Take \(a\) to be half the coefficient of \(x\), then check \((x+a)^2\) against the original.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1001</deployedseed>
    <deployedseed>1002</deployedseed>
    <deployedseed>1003</deployedseed>
    <deployedseed>1004</deployedseed>
    <deployedseed>1005</deployedseed>
    <deployedseed>1006</deployedseed>
    <deployedseed>1007</deployedseed>
    <deployedseed>1008</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Both parts correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta1v</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Wrong Yes/No choice</description>
      <testinput>
        <name>ans1</name>
        <value>ta1w</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Part (b) correct value but expanded (wrong form)</description>
      <testinput>
        <name>ans1</name>
        <value>ta1v</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>expr</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-notform</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Part (b) used the constant as a</description>
      <testinput>
        <name>ans1</name>
        <value>ta1v</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>w2_const</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-const</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Part (b) generic wrong answer</description>
      <testinput>
        <name>ans1</name>
        <value>ta1v</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>(x+aa+1)^2</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-other</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
