<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Inverse Trigonometric Functions — why the domain must be restricted</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The ordinary sine function \(\sin(x)\) is <strong>not</strong> one-to-one over all real numbers, so it cannot have a genuine inverse function unless we first restrict its domain.</p>
<p>Which of the following best explains <em>why</em> this restriction is necessary?</p>
<div>[[input:ans1]] [[validation:ans1]]</div>
<div>[[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Since \(\sin(x)\) repeats every \(2\pi\), infinitely many different \(x\)-values (e.g. \(x=\pi/6\), \(x=5\pi/6+2\pi\), etc.) all give the same \(y\)-value. An inverse function must send each input to exactly one output, so we restrict the domain of \(\sin\) (usually to \([-\pi/2,\pi/2]\)) to make it one-to-one before inverting it.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[opts : [
  ["correct", true, "Because sin(x) is periodic, many different x-values give the same y-value, so reversing the function (swapping inputs and outputs) would not give a single well-defined output for each input."],
  ["wrong1", false, "Because sin(x) is only defined for x between -1 and 1."],
  ["wrong2", false, "Because sin(x) is an odd function, and odd functions never have inverses."],
  ["wrong3", false, "Because the graph of sin(x) does not pass the horizontal line test at exactly one point."]
];
ta : "correct";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.10 (Inverse of Trigonometric Functions). Conceptual reasoning question on why sin(x) needs a restricted domain before it can be inverted, distinguishing this from plausible-sound</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.10 (Inverse of Trigonometric Functions). Conceptual reasoning question on why sin(x) needs a restricted domain before it can be inverted, distinguishing this from plausible-sounding but incorrect explanations.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>opts</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Chose the correct explanation?</description>
        <answertest>String</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>Think about what makes a function invertible: it must be one-to-one. Why does the periodicity of sin(x) break that requirement?</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
