<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Inverse Trigonometric Functions — exact values and the principal range</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Find the exact value of \( {@invname@}({@val@}) \), giving your answer in radians (a multiple of \(\pi\)).</p>
<div>[[input:ans1]] [[validation:ans1]]</div>
<div>[[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The inverse trig functions are defined with a restricted (<strong>principal</strong>) range so that they give exactly one answer:</p>
<ul>
<li>\(\sin^{-1}\) always returns a value in \([-\pi/2,\pi/2]\)</li>
<li>\(\cos^{-1}\) always returns a value in \([0,\pi]\)</li>
<li>\(\tan^{-1}\) always returns a value in \((-\pi/2,\pi/2)\)</li>
</ul>
<p>Even though other angles also satisfy the underlying trig equation, only the one inside the principal range is the correct value of the inverse function. Here \( {@invname@}({@val@}) = {@ta@} \).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[funcidx : rand(3);                    /* 0 arcsin, 1 arccos, 2 arctan */
invname : if funcidx=0 then "sin^{-1}" elseif funcidx=1 then "cos^{-1}" else "tan^{-1}";
signidx : rand(2);                     /* 0 = positive input, 1 = negative input (not available for arccos in this simple set) */
baseidx : rand_with_prohib(1,3,[]);     /* which of a small standard-value table to use */
/* standard values, kept simple: use k = pi/6, pi/4, pi/3 as principal outputs for POSITIVE inputs */
basepi : [1/6,1/4,1/3][baseidx];
val_pos : if funcidx=0 then (if baseidx=1 then 1/2 elseif baseidx=2 then sqrt(2)/2 else sqrt(3)/2)
          elseif funcidx=1 then (if baseidx=1 then sqrt(3)/2 elseif baseidx=2 then sqrt(2)/2 else 1/2)
          else (if baseidx=1 then 1/sqrt(3) elseif baseidx=2 then 1 else sqrt(3));
usesign : if funcidx=1 then 0 else signidx;   /* keep arccos example with a positive input only, for simplicity */
val : if usesign=1 then -val_pos else val_pos;
/* principal value ranges: arcsin in [-pi/2,pi/2], arccos in [0,pi], arctan in (-pi/2,pi/2) */
ta : if funcidx=0 then (if usesign=1 then -basepi*%pi else basepi*%pi)
     elseif funcidx=1 then (if usesign=1 then %pi-basepi*%pi else basepi*%pi)
     else (if usesign=1 then -basepi*%pi else basepi*%pi);
/* common mistake: giving a coterminal / out-of-principal-range angle instead */
wrong_outofrange : if funcidx=0 then (if usesign=1 then %pi+basepi*%pi else %pi-basepi*%pi)
                    elseif funcidx=1 then (if usesign=1 then -( %pi-basepi*%pi) else -basepi*%pi)
                    else (if usesign=1 then %pi-basepi*%pi else %pi+basepi*%pi);]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.10 (Inverse of Trigonometric Functions). Evaluates arcsin/arccos/arctan at a standard value; catches giving a mathematically valid but out-of-principal-range angle instead of the</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.10 (Inverse of Trigonometric Functions). Evaluates arcsin/arccos/arctan at a standard value; catches giving a mathematically valid but out-of-principal-range angle instead of the one true value the inverse function returns.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>12</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct principal value?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Gave a coterminal/out-of-range angle instead of the principal value?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_outofrange</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-outofrange</trueanswernote>
        <truefeedback format="html">
          <text>That angle does satisfy the underlying trig ratio, but it lies outside the principal range of this inverse function, so it is not the value the inverse function returns.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>That is not the correct exact value.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Out-of-principal-range angle</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_outofrange</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-outofrange</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
