<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Inverse of Polynomials — inverse of a restricted quadratic</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Let \( f(x) = (x-{@h@})^2 + {@k@} \) for \( x \geq {@h@} \) (the domain is restricted so that \(f\) is invertible).</p>
<p>Find \( f^{-1}(x) \).</p>
<div>\( f^{-1}(x) = \) [[input:ans1]] [[validation:ans1]]</div>
<div>[[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Swap \(x\) and \(y\) in \(y=(x-{@h@})^2+{@k@}\) to get \(x=(y-{@h@})^2+{@k@}\), then solve for \(y\):</p>
\[
\begin{align*}
x-{@k@} &= (y-{@h@})^2\\
\pm\sqrt{x-{@k@}} &= y-{@h@}\\
y &= {@h@} \pm \sqrt{x-{@k@}}
\end{align*}
\]
<p>Taking the square root gives <strong>two</strong> possible branches, \(\pm\). The domain restriction \(x\geq {@h@}\) on the <em>original</em> function tells us which branch to keep for the inverse: since the original domain only included \(x\geq {@h@}\), the inverse's output must also be \(\geq {@h@}\), which means we need the \(+\) branch:</p>
\[ f^{-1}(x) = {@h@}+\sqrt{x-{@k@}}. \]
<p>Choosing the \(-\) branch would give outputs \(\leq {@h@}\), which contradicts the domain restriction on \(f\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[h : rand_with_prohib(-6,6,[]);
k : rand_with_prohib(-6,6,[]);
ta : h+sqrt(x-k);                        /* correct branch, since x >= h means the "+" root */
wrong_branch : h-sqrt(x-k);              /* wrong sign choice on the square root */
wrong_noshift : sqrt(x-k);               /* forgot to add h back on */]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.8 (Inverse of Polynomials). Finds the inverse of a domain-restricted quadratic; the key trick is choosing the correct +/- branch of the square root based on the original domain r</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.8 (Inverse of Polynomials). Finds the inverse of a domain-restricted quadratic; the key trick is choosing the correct +/- branch of the square root based on the original domain restriction, rather than just forgetting the restriction exists.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>20</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct inverse with the correct sign branch?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Chose the wrong sign branch of the square root?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_branch</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-wrongbranch</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[You've chosen the \(-\) branch, but the domain restriction \(x\geq {@h@}\) on the original function means the inverse's outputs must also be \(\geq {@h@}\) &#8212; that requires the \(+\) branch.]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Forgot to add h back on?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_noshift</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-noshift</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[You found \(\sqrt{x-{@k@}}\) correctly but forgot the horizontal shift &#8212; add {@h@} back on.]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>That is not the correct inverse function.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Wrong sign branch</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_branch</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-wrongbranch</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Forgot the shift</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_noshift</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.2000000</expectedpenalty>
        <expectedanswernote>prt1-noshift</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
