<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Trigonometric Functions — amplitude and period from a formula</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the function \( y = {@amp@}\sin({@b@}x) \).</p>
<p><strong>(a)</strong> State the amplitude.</p>
<div>[[input:ans_amp]] [[validation:ans_amp]]</div>
<p><strong>(b)</strong> State the period (in radians).</p>
<div>[[input:ans_period]] [[validation:ans_period]]</div>
<div>[[feedback:prt_amp]] [[feedback:prt_period]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>For \(y=A\sin(Bx)\), the amplitude is \(|A|\), and the period is \(\dfrac{2\pi}{B}\) (not \(2\pi \times B\), and not \(\dfrac{\pi}{B}\)).</p>
<p>Here \(A={@amp@}\) and \(B={@b@}\), so the amplitude is {@ta_amp@} and the period is \(\dfrac{2\pi}{ {@b@} } = {@ta_period@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>amp : rand_with_prohib(2,6,[]);
b : rand_with_prohib(2,5,[]);
ta_amp : amp;
ta_period : 2*%pi/b;
wrong_period : 2*%pi*b;                    /* multiplied by b instead of dividing */
wrong_period2 : %pi/b;                     /* forgot the factor of 2 */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.9 (Trigonometric Functions). Reads amplitude and period off A*sin(Bx); catches the very common period-formula errors of multiplying by B instead of dividing, and forgetting the f</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.9 (Trigonometric Functions). Reads amplitude and period off A*sin(Bx); catches the very common period-formula errors of multiplying by B instead of dividing, and forgetting the factor of 2.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans_amp</name>
      <type>algebraic</type>
      <tans>ta_amp</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans_period</name>
      <type>algebraic</type>
      <tans>ta_period</tans>
      <boxsize>12</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt_amp</name>
      <value>0.4000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct amplitude?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_amp</sans>
        <tans>ta_amp</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_amp-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_amp-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>The amplitude of A*sin(Bx) is simply |A|.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt_period</name>
      <value>0.6000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct period?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_period</sans>
        <tans>ta_period</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_period-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt_period-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Multiplied by B instead of dividing?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_period</sans>
        <tans>wrong_period</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_period-multiplied</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[The period formula is \(2\pi\) <strong>divided by</strong> \(B\), not multiplied.]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt_period-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Forgot the factor of 2?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans_period</sans>
        <tans>wrong_period2</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt_period-nofactor2</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[You are close, but the period formula is \(\dfrac{2\pi}{B}\), not \(\dfrac{\pi}{B}\) &#8212; don't drop the factor of 2.]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt_period-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>That is not the correct period. The period of A*sin(Bx) is 2*pi/B.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Both correct</description>
      <testinput>
        <name>ans_amp</name>
        <value>ta_amp</value>
      </testinput>
      <testinput>
        <name>ans_period</name>
        <value>ta_period</value>
      </testinput>
      <expected>
        <name>prt_amp</name>
        <expectedscore>0.4000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_amp-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt_period</name>
        <expectedscore>0.6000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_period-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Multiplied by B instead of dividing</description>
      <testinput>
        <name>ans_amp</name>
        <value>ta_amp</value>
      </testinput>
      <testinput>
        <name>ans_period</name>
        <value>wrong_period</value>
      </testinput>
      <expected>
        <name>prt_amp</name>
        <expectedscore>0.4000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_amp-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt_period</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt_period-multiplied</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Forgot the factor of 2</description>
      <testinput>
        <name>ans_amp</name>
        <value>ta_amp</value>
      </testinput>
      <testinput>
        <name>ans_period</name>
        <value>wrong_period2</value>
      </testinput>
      <expected>
        <name>prt_amp</name>
        <expectedscore>0.4000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt_amp-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt_period</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.2000000</expectedpenalty>
        <expectedanswernote>prt_period-nofactor2</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
