<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Inverse of a scaling matrix (INNODEMS)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The scaling matrix \( S = {@Smat@} \) stretches the \(x\)-coordinate by \({@a@}\) and the \(y\)-coordinate by \({@b@}\).</p>
<p>(a) Write the matrix \( S^{-1} \) that undoes this scaling.</p>
<table style="border-collapse: collapse;">
  <tr>
    <td style="white-space: nowrap; vertical-align: middle;">
      \( S^{-1} = \)
    </td>
    <td style="white-space: nowrap; vertical-align: middle;">
      [[input:ans1]]
    </td>
    <td style="white-space: nowrap; vertical-align: middle;">
      [[validation:ans1]]
    </td>
    <td style="white-space: nowrap; vertical-align: middle;">
      [[feedback:prt1]]
    </td>
  </tr>
</table>
<p>(b) Verify that this is an inverse of \(S\) by working out the matrix product \( S\,S^{-1} \).</p>
<table style="border-collapse: collapse;">
  <tr>
    <td style="white-space: nowrap; vertical-align: middle;">
      \( S\,S^{-1} = \)
    </td>
    <td style="white-space: nowrap; vertical-align: middle;">
      [[input:ans2]]
    </td>
    <td style="white-space: nowrap; vertical-align: middle;">
      [[validation:ans2]]
    </td>
    <td style="white-space: nowrap; vertical-align: middle;">
      [[feedback:prt2]]
    </td>
  </tr>
</table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>To undo a horizontal stretching by \({@a@}\), we stretch by \({@1/a@}\).</p>
<p>To undo a vertical stretching by \({@b@}\), we stretch by \({@1/b@}\). Thus the inverse matrix is:</p>
\[ S^{-1} = {@Smat@}^{-1} = {@ta1@}. \]
<p>Multiplying \(S\) by this matrix returns the identity, confirming it is the inverse:</p>
\[ S\,S^{-1} = {@Smat@}\,{@ta1@} = {@ta2@} = I. \]]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>L : random_permutation([2,3,4,5]);
a : L[1];
b : L[2];
Smat : matrix([a,0],[0,b]);
ta1 : invert(Smat);
ta2 : ev(Smat . ta1, simp);
wrong_orig : Smat;
wrong_neg : matrix([-a,0],[0,-b]);
wrong_swap : matrix([1/b,0],[0,1/a]);
wrong_sq : ev(Smat . Smat, simp);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( S=\operatorname{diag}({@a@},{@b@}),\ S^{-1}={@ta1@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>ta1</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>invert,adjoint</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>matrix</type>
      <tans>ta2</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>invert,ident,diag</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_orig</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-orig</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is \(S\) itself. The inverse should <em>undo</em> the scaling, so reciprocate each diagonal entry.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_neg</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-neg</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The inverse is not the negative of \(S\). For a diagonal matrix, invert each diagonal entry: \(a \mapsto 1/a\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_swap</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-swap</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Right idea, but the diagonal entries are in the wrong positions. The \((1,1)\) entry undoes the \(x\)-scaling, so it should be \(1/{@a@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not the inverse of \(S\). To undo a diagonal scaling, take the reciprocal of each diagonal entry. See the worked solution.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The product is the identity matrix \(I\), which confirms the two matrices are inverses.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>wrong_sq</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-square</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That is \(S\,S\), not \(S\,S^{-1}\). Multiply \(S\) by the inverse you found in part (a).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The product of a matrix and its inverse should be the identity matrix \( \left[\begin{array}{cc}1&0\\0&1\end{array}\right] \). Check your multiplication.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1001</deployedseed>
    <deployedseed>2002</deployedseed>
    <deployedseed>3003</deployedseed>
    <deployedseed>4004</deployedseed>
    <deployedseed>5005</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Both parts correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Part (a): gave S itself</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_orig</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.3333333</expectedpenalty>
        <expectedanswernote>prt1-orig</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Part (a): negated instead of inverted</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_neg</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-neg</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Part (a): reciprocals in swapped positions</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_swap</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-swap</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Part (a): generic wrong matrix</description>
      <testinput>
        <name>ans1</name>
        <value>ev(2*ta1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-wrong</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description>Part (b): computed S.S instead of S.S^{-1}</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>wrong_sq</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-square</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>7</testcase>
      <description>Part (b): generic wrong product</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>wrong_orig</value>
      </testinput>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-wrong</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
