<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Undoing a Transformation (INNODEMS)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Let
\[ A = {@A@} \qquad\text{and}\qquad A^{-1} = {@Ai@}, \]
and take the vector
\[ \mathbf{x} = {@xx@}. \]</p>
<p>Compute \( \mathbf{y} = A\mathbf{x} \):</p>
<table style="border-collapse: collapse;">
  <tr>
    <td style="white-space: nowrap; vertical-align: middle;">
      \( \mathbf{y} = \)
    </td>
    <td style="white-space: nowrap; vertical-align: middle;">
      [[input:ans1]]
    </td>
<td style="white-space: nowrap; vertical-align: middle;">
      [[validation:ans1]]
    </td>
    <td style="white-space: nowrap; vertical-align: middle;">
      [[feedback:prt1]]
    </td>
  </tr>
</table>
<p>Now use the inverse to recover \( \mathbf{x} \), i.e. compute \( A^{-1}\mathbf{y} \):</p>
<table style="border-collapse: collapse;">
  <tr>
    <td style="white-space: nowrap; vertical-align: middle;">
      \( A^{-1}\mathbf{y} = \)
    </td>
    <td style="white-space: nowrap; vertical-align: middle;">
      [[input:ans2]]
    </td>
<td style="white-space: nowrap; vertical-align: middle;">
      [[validation:ans2]]
    </td>
    <td style="white-space: nowrap; vertical-align: middle;">
      [[feedback:prt2]]
    </td>
  </tr>
</table>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Multiply \(A\) by \(\mathbf{x}\), taking each row of \(A\) against the column \(\mathbf{x}\):</p>
\[ \mathbf{y} = A\mathbf{x} = {@A@}{@xx@} = {@ta1@}. \]
<p>Applying the inverse undoes the transformation, returning the original vector:</p>
\[ A^{-1}\mathbf{y} = {@Ai@}{@ta1@} = {@ta2@} = \mathbf{x}. \]
<p>This is why \(A^{-1}A = I\): the inverse exactly reverses what \(A\) does.</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>s1 : rand_with_prohib(-3,3,[0]);
s2 : rand_with_prohib(-3,3,[0]);
t  : rand_with_prohib(-2,2,[0]);
r  : rand_with_prohib(-2,2,[t]);
A  : matrix([s1+r*t*s2, r*s2],[t*s2, s2]);
Ai : invert(A);
x1 : rand_with_prohib(-4,4,[0]);
x2 : rand_with_prohib(-4,4,[0]);
xx : matrix([x1],[x2]);
ta1 : ev(A . xx, simp);
ta2 : ev(Ai . ta1, simp);
wrongy : ev(transpose(A) . xx, simp);
wrongx : ev(A . ta1, simp);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( A={@A@},\ \mathbf{x}={@xx@},\ \mathbf{y}={@ta1@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>matrix</type>
      <tans>ta1</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>matrix</type>
      <tans>ta2</tans>
      <boxsize>2</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrongy</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-transpose</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you multiplied using the columns of \(A\) instead of its rows (equivalently, you used \(A^{\mathsf{T}}\)). For \(A\mathbf{x}\), take each <em>row</em> of \(A\) against the column \(\mathbf{x}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrongvalue</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not equal to \(A\mathbf{x}\). Multiply each row of \(A\) by \(\mathbf{x}\) and add the products.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>wrongx</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-applieda</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You applied \(A\) again rather than its inverse. To undo the transformation, multiply \(\mathbf{y}\) by \(A^{-1}\), not by \(A\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-wrongvalue</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not equal to \(A^{-1}\mathbf{y}\). Since \(A^{-1}\) undoes \(A\), the result should be your original \(\mathbf{x}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1001</deployedseed>
    <deployedseed>2002</deployedseed>
    <deployedseed>3003</deployedseed>
    <deployedseed>4004</deployedseed>
    <deployedseed>5005</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Both correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Transpose (column) error on y</description>
      <testinput>
        <name>ans1</name>
        <value>wrongy</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-transpose</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Applied A instead of inverse</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>wrongx</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-applieda</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Generic wrong values both boxes</description>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+matrix([1],[0]),simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta2+matrix([0],[1]),simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-wrongvalue</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-wrongvalue</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
