<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Identifying the centre and radius, then use it to sketch the circle</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>A circle has equation \[ {@lhs@} = {@R@}. \]</p>
<p><strong>1)</strong> State the centre of the circle as a coordinate \([x,y]\):</p>
<div>centre \(=\) [[input:ansCentre]] [[validation:ansCentre]] [[feedback:prtCentre]]</div>
<p><strong>2)</strong> State the radius:</p>
<div>radius \(=\) [[input:ansR]] [[validation:ansR]] [[feedback:prtRadius]]</div>
<p><strong>3) </strong> Sketch the circle by dragging the red point to the <strong>centre</strong>, and dragging the blue point to a point <strong>on the circle</strong>.</p>
[[jsxgraph input-ref-ansC="rC" input-ref-ansP="rP" width="520px" height="520px"]]
var brd = JXG.JSXGraph.initBoard(divid, {boundingbox:[-8, 8, 8, -8], axis:true, grid:true, keepaspectratio:true, showCopyright:false, showNavigation:false});
var C = brd.create('point',[-5,-5],{name:'centre',color:'#cc0000',fillColor:'#cc0000',size:5,snapToGrid:true,snapSizeX:1,snapSizeY:1,label:{offset:[8,8],cssStyle:'font-weight:bold;color:#cc0000'}});
var P = brd.create('point',[-5,-3],{name:'',color:'#1f4fd8',fillColor:'#1f4fd8',size:5,snapToGrid:true,snapSizeX:1,snapSizeY:1});
brd.create('circle',[C,P],{strokeColor:'#1f4fd8',strokeWidth:2.5,fillOpacity:0,highlight:false});
stack_jxg.bind_point(rC, C);
stack_jxg.bind_point(rP, P);
[[/jsxgraph]]
<div style="display:none">[[input:ansC]] [[validation:ansC]] [[input:ansP]] [[validation:ansP]]</div>
<div>[[feedback:prtGraph]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Compare with the standard form \((x-h)^2+(y-k)^2=r^2\).</p>
<ol>
<li><strong>Centre.</strong> \(h={@h@}\) and \(k={@k@}\), so the centre is \(({@h@},\,{@k@})\).</li>
<li><strong>Radius.</strong> The right-hand side is \(r^2={@R@}\), so \(r=\sqrt{@R@}={@r@}\).</li>
<li><strong>Sketch.</strong> Plot the centre \(({@h@},{@k@})\), then mark points \({@r@}\) units to the left, right, above and below it, and draw a smooth circle through them.</li>
</ol>
<p>The circle is shown below.</p>
[[jsxgraph width="480px" height="480px"]]
var H={#h#}, K={#k#}, R={#r#};
var brd = JXG.JSXGraph.initBoard(divid, {boundingbox:[-8, 8, 8, -8], axis:true, grid:true, keepaspectratio:true, showCopyright:false, showNavigation:false});
brd.create('circle',[[H,K], R],{strokeColor:'#1f4fd8',strokeWidth:2.5,fillOpacity:0,highlight:false});
brd.create('point',[H,K],{name:'('+H+','+K+')',fixed:true,color:'#cc0000',fillColor:'#cc0000',size:4,label:{offset:[8,-4],cssStyle:'font-weight:bold;color:#cc0000'}});
brd.create('segment',[[H,K],[H+R,K]],{dash:2,strokeColor:'#cc0000',strokeWidth:1.5});
brd.create('text',[H+R/2, K+0.4, 'r = '+R],{anchorX:'middle',fixed:true,fontSize:12,cssStyle:'font-weight:bold;color:#cc0000'});
[[/jsxgraph]]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>h : rand([-3,-2,-1,0,1,2,3]);
k : rand([-3,-2,-1,0,1,2,3]);
r : rand([2,3,4]);
R : r^2;
lhs : (x-h)^2 + (y-k)^2;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>centre \(({@h@},{@k@})\), \(r={@r@}\), \(R={@R@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ansC</name>
      <type>algebraic</type>
      <tans>[h,k]</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ansCentre</name>
      <type>algebraic</type>
      <tans>[h,k]</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>1</insertstars>
      <syntaxhint>[a,b]</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ansP</name>
      <type>algebraic</type>
      <tans>[h+r,k]</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options>hideanswer</options>
    </input>
    <input>
      <name>ansR</name>
      <type>numerical</type>
      <tans>r</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prtCentre</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ansCentre</sans>
        <tans>[h,k]</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtCentre-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtCentre-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The centre is \((h,k)\): read \(h\) and \(k\) from \((x-h)^2+(y-k)^2\). Watch the signs.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtGraph</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>cx : round(ansC[1]); cy : round(ansC[2]);
px : round(ansP[1]); py : round(ansP[2]);
rsq : (px-cx)^2 + (py-cy)^2;
centreok : if (is(cx=h) and is(cy=k)) then 1 else 0;
radok : if is(rsq=R) then 1 else 0;
ng : centreok + radok;</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>centre and radius</description>
        <answertest>AlgEquiv</answertest>
        <sans>ng</sans>
        <tans>2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtGraph-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prtGraph-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>one right</description>
        <answertest>AlgEquiv</answertest>
        <sans>ng</sans>
        <tans>1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtGraph-partial</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Part of the circle is right. Put the centre at \(({@h@},{@k@})\) and drag the blue point exactly \({@r@}\) units away (e.g. to \(({@h@}+{@r@},{@k@})\)).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtGraph-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Drag the red point to the centre \(({@h@},{@k@})\) and the blue point to a point \({@r@}\) units away.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prtRadius</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ansR</sans>
        <tans>r</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prtRadius-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prtRadius-wrong</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>The right-hand side equals \(r^2={@R@}\), so \(r=\sqrt{@R@}\).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>30001</deployedseed>
    <deployedseed>30002</deployedseed>
    <deployedseed>30003</deployedseed>
    <deployedseed>30004</deployedseed>
    <deployedseed>30005</deployedseed>
    <deployedseed>30006</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All correct</description>
      <testinput>
        <name>ansC</name>
        <value>[h,k]</value>
      </testinput>
      <testinput>
        <name>ansCentre</name>
        <value>[h,k]</value>
      </testinput>
      <testinput>
        <name>ansP</name>
        <value>[h+r,k]</value>
      </testinput>
      <testinput>
        <name>ansR</name>
        <value>r</value>
      </testinput>
      <expected>
        <name>prtCentre</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtCentre-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtGraph</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtGraph-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prtRadius</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtRadius-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Centre wrong; graph centre wrong but radius right</description>
      <testinput>
        <name>ansC</name>
        <value>[h+1,k]</value>
      </testinput>
      <testinput>
        <name>ansCentre</name>
        <value>[h+1,k]</value>
      </testinput>
      <testinput>
        <name>ansP</name>
        <value>[h+1+r,k]</value>
      </testinput>
      <testinput>
        <name>ansR</name>
        <value>r</value>
      </testinput>
      <expected>
        <name>prtCentre</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtCentre-wrong</expectedanswernote>
      </expected>
      <expected>
        <name>prtGraph</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtGraph-partial</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Radius = R (forgot square root)</description>
      <testinput>
        <name>ansC</name>
        <value>[h,k]</value>
      </testinput>
      <testinput>
        <name>ansCentre</name>
        <value>[h,k]</value>
      </testinput>
      <testinput>
        <name>ansP</name>
        <value>[h+r,k]</value>
      </testinput>
      <testinput>
        <name>ansR</name>
        <value>R</value>
      </testinput>
      <expected>
        <name>prtRadius</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prtRadius-wrong</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
