<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>GeoGebra Incentre Construction (INNODEMS)</text>
    </name>
    <questiontext format="plain_text">
      <text><![CDATA[<p>The triangle \(ABC\) below has vertices
\(A={@ntupleify(A)@}\), \(B={@ntupleify(B)@}\) and \(C={@ntupleify(C)@}\).</p>

<p>Using only the construction tools provided (which include <strong>angle bisector</strong> and <strong>perpendicular </strong> tools), construct the centre of the <strong>inscribed circle</strong> (the incentre) of triangle \(ABC\), and its inscribed circle.</p>

<p>GeoGebra displays the coordinates of the points you construct. Read off the position of the incentre and the radius of the inscribed circle (the inradius), and enter them below, each rounded to <strong>3 decimal places</strong>.</p>

[[geogebra set="A__fixed,B__fixed,C__fixed" width="1200px" height="600px"]]
params["material_id"]="dcdnc4ee";
/* Restricted compass-and-straightedge toolbar. customToolBar OVERRIDES the file's
   toolbar; numbers are GeoGebra tool IDs, | starts a new menu.
   0 Move, 6 Delete | 1 Point, 5 Intersect, 501 Point-on-object |
   2 Line, 15 Segment, 18 Ray | 10 Circle(centre,point), 53 Compass. */
/* params["appName"]="geometry"; */
params["showToolBar"]=true;
params["customToolBar"]="0 | 6 | 5 | 2 | 10 | 9 | 4 | 8 | 38";
params["showToolBarHelp"]=false;
/* Hide the algebra input bar and menu bar. NOTE: the algebra VIEW panel is
   governed by the saved file (perspective is ignored with material_id), so ALSO
   close the Algebra View in the GeoGebra material and re-save. */
params["showAlgebraInput"]=true;
params["algebraInputPosition"]="bottom";
params["showMenuBar"]=false;
params["showResetIcon"]=true;
params["showFullscreenButton"]=true;
[[/geogebra]]

<div>Incentre \(I = (\,\) [[input:xc]] \(,\) [[input:yc]] \(\,)\) \( (\,\)[[validation:xc]], [[validation:yc]] \(\,)\)</div>

<div>[[feedback:prt1]]</div>

<div>Inradius \(r = \) [[input:r]] [[validation:r]]</div>

<div>[[feedback:prt2]]</div>]]></text>
    </questiontext>
    <generalfeedback format="plain_text">
      <text><![CDATA[<p>The centre of the inscribed circle is the <em>incentre</em>: the point where the three interior angle bisectors meet. It is the unique point equidistant from all three sides, so the perpendicular distance from it to each side is the same &mdash; that common distance is the inradius \(r\).</p>

<p><strong>Construction.</strong> Bisect any two interior angles (say at \(A\) and \(B\)) using the angle-bisector tool; their intersection \(I\) is the incentre. Dropping a perpendicular from \(I\) to any side gives the inradius.</p>


[[geogebra set="A__fixed,B__fixed,C__fixed" width="1200px" height="600px"]]
params["material_id"]="pgd3hywq";
/* Restricted compass-and-straightedge toolbar. customToolBar OVERRIDES the file's
   toolbar; numbers are GeoGebra tool IDs, | starts a new menu.
   0 Move, 6 Delete | 1 Point, 5 Intersect, 501 Point-on-object |
   2 Line, 15 Segment, 18 Ray | 10 Circle(centre,point), 53 Compass. */
/* params["appName"]="geometry"; */
params["showToolBar"]=true;
params["customToolBar"]="0 | 6 | 5 | 2 | 10 | 9 | 4 | 8 | 38";
params["showToolBarHelp"]=false;
/* Hide the algebra input bar and menu bar. NOTE: the algebra VIEW panel is
   governed by the saved file (perspective is ignored with material_id), so ALSO
   close the Algebra View in the GeoGebra material and re-save. */
params["showAlgebraInput"]=true;
params["algebraInputPosition"]="bottom";
params["showMenuBar"]=false;
params["showResetIcon"]=true;
params["showFullscreenButton"]=true;
[[/geogebra]]

<p>For this triangle the incentre is
\(I \approx ({@dispdp(ta_X[1],3)@},\ {@dispdp(ta_X[2],3)@})\) and the inradius is
\(r \approx {@dispdp(ta_r,3)@}\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text>tris : [ [[0,0],[11,0],[-1,4]], [[0,0],[11,0],[0,5]], [[0,0],[10,0],[-1,5]], [[0,0],[9,0],[-1,4]], [[0,0],[11,0],[2,4]], [[0,0],[11,0],[-1,6]], [[0,0],[10,0],[0,4]], [[0,0],[11,0],[1,6]], [[0,0],[10,0],[2,4]], [[0,0],[11,0],[2,5]] ];
tri : tris[rand(length(tris))+1];
A : tri[1] + [-1,-1];
B : tri[2] + [-1,-1];
C : tri[3] + [-1,-1];
/* side lengths, each opposite its like-named vertex */
la : Distance(B,C);  /* a = |BC|, opposite A */
lb : Distance(A,C);  /* b = |CA|, opposite B */
lc : Distance(A,B);  /* c = |AB|, opposite C */
perim : la+lb+lc;
ssemi : perim/2;
/* incentre = (a*A + b*B + c*C)/(a+b+c) */
ta_X : [(la*A[1]+lb*B[1]+lc*C[1])/perim, (la*A[2]+lb*B[2]+lc*C[2])/perim];
/* area by the shoelace formula; inradius r = area/s */
area : abs(A[1]*(B[2]-C[2])+B[1]*(C[2]-A[2])+C[1]*(A[2]-B[2]))/2;
ta_r : area/ssemi;
/* distractor centres used only for misconception feedback */
ta_G : [(A[1]+B[1]+C[1])/3, (A[2]+B[2]+C[2])/3];   /* centroid: medians */
ddet : 2*(A[1]*(B[2]-C[2])+B[1]*(C[2]-A[2])+C[1]*(A[2]-B[2]));
ta_O : [((A[1]^2+A[2]^2)*(B[2]-C[2])+(B[1]^2+B[2]^2)*(C[2]-A[2])+(C[1]^2+C[2]^2)*(A[2]-B[2]))/ddet, ((A[1]^2+A[2]^2)*(C[1]-B[1])+(B[1]^2+B[2]^2)*(A[1]-C[1])+(C[1]^2+C[2]^2)*(B[1]-A[1]))/ddet];   /* circumcentre: perpendicular bisectors */
/* rounded model answers (to 3 d.p.), used for the inputs and the question tests */
rnd3(z) := float(round(float(z)*1000)/1000);
ans_xc : rnd3(ta_X[1]);
ans_yc : rnd3(ta_X[2]);
ans_r  : rnd3(ta_r);
ans_Ox : rnd3(ta_O[1]);
ans_Oy : rnd3(ta_O[2]);
ans_Gx : rnd3(ta_G[1]);
ans_Gy : rnd3(ta_G[2]);
ans_Rcirc : rnd3(sqrt((ta_O[1]-A[1])^2+(ta_O[2]-A[2])^2));   /* circumradius */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\(A={@ntupleify(A)@},\ B={@ntupleify(B)@},\ C={@ntupleify(C)@},\ I \approx ({@dispdp(ta_X[1],3)@},\ {@dispdp(ta_X[2],3)@}),\ r \approx {@dispdp(ta_r,3)@}\)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>r</name>
      <type>numerical</type>
      <tans>ans_r</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>xc</name>
      <type>numerical</type>
      <tans>ans_xc</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>yc</name>
      <type>numerical</type>
      <tans>ans_yc</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>/* largest single-coordinate error, so x AND y must each be correct to 3 d.p. */
errInc  : max(abs(xc-ta_X[1]), abs(yc-ta_X[2]));
errCirc : max(abs(xc-ta_O[1]), abs(yc-ta_O[2]));
errCent : max(abs(xc-ta_G[1]), abs(yc-ta_G[2]));</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Both incentre coordinates correct to 3 d.p.?</description>
        <answertest>NumAbsolute</answertest>
        <sans>errInc</sans>
        <tans>0</tans>
        <testoptions>0.0005</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-incentre</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The centre is correctly placed at the incentre.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Did they find the circumcentre instead?</description>
        <answertest>NumAbsolute</answertest>
        <sans>errCirc</sans>
        <tans>0</tans>
        <testoptions>0.05</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-circumcentre</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That point is the <em>circumcentre</em> &mdash; where the perpendicular bisectors of the <strong>sides</strong> meet. It is equidistant from the three vertices, not the three sides. For the incircle you must bisect the <strong>angles</strong>, not the sides.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Did they find the centroid instead?</description>
        <answertest>NumAbsolute</answertest>
        <sans>errCent</sans>
        <tans>0</tans>
        <testoptions>0.05</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-centroid</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>That point is the <em>centroid</em> &mdash; where the three medians meet (the triangle's balance point). The incircle's centre is equidistant from the sides, which you get from the angle bisectors.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-other</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your point \(I\) is not the incentre. The incentre is where the angle bisectors cross &mdash; check that each bisector makes equal angles with the two sides meeting at that vertex.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Inradius correct to 3 d.p.?</description>
        <answertest>NumAbsolute</answertest>
        <sans>r</sans>
        <tans>ta_r</tans>
        <testoptions>0.0005</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-radius-ok</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>The inradius is correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-radius-no</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Your inradius is not correct (to 3 decimal places). The inradius is the perpendicular distance from the incentre to any side, equal to \(r = \dfrac{\text{area}}{s}\) where \(s\) is the semi-perimeter.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1001</deployedseed>
    <deployedseed>1002</deployedseed>
    <deployedseed>1003</deployedseed>
    <deployedseed>1004</deployedseed>
    <deployedseed>1005</deployedseed>
    <deployedseed>1006</deployedseed>
    <deployedseed>1007</deployedseed>
    <deployedseed>1008</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct: incentre and inradius both correct</description>
      <testinput>
        <name>r</name>
        <value>ans_r</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_xc</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_yc</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-incentre</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-radius-ok</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Centre correct but inradius wrong</description>
      <testinput>
        <name>r</name>
        <value>ans_r+5</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_xc</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_yc</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-incentre</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-radius-no</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Used the circumcentre and circumradius instead</description>
      <testinput>
        <name>r</name>
        <value>ans_Rcirc</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_Ox</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_Oy</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-0-F | prt1-circumcentre</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-radius-no</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Used the centroid instead (inradius correct)</description>
      <testinput>
        <name>r</name>
        <value>ans_r</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>ans_Gx</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>ans_Gy</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-0-F | prt1-1-F | prt1-centroid</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-radius-ok</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Neither a recognised centre nor inradius (point at vertex A)</description>
      <testinput>
        <name>r</name>
        <value>ans_r+5</value>
      </testinput>
      <testinput>
        <name>xc</name>
        <value>rnd3(A[1])</value>
      </testinput>
      <testinput>
        <name>yc</name>
        <value>rnd3(A[2])</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-0-F | prt1-1-F | prt1-other</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-radius-no</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
