<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>angles at the centre an on the circumference Q1</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>\(A\), \(B\) and \(C\) are points on a circle with centre \(O\). The angle \(ACB = {@ang@}^\circ\).</p>
<p>{@img@}</p>

<p><strong>(a)</strong> Find the angle \(AOB\).</p>
<p>[[input:ans1]] degrees [[validation:ans1]] [[feedback:prt1]]</p>

<p><strong>(b)</strong> Find the reflex angle \(AOB\) (the angle on the other side of \(O\)).</p>
<p>[[input:ans2]] degrees [[validation:ans2]] [[feedback:prt2]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p><strong>(a)</strong> The angle at the centre is twice the angle at the circumference when both stand on the same arc \(AB\).</p>
\[ AOB = 2 \times {@ang@}^\circ = {@cen@}^\circ \]

<p><strong>(b)</strong> Angles around a point add up to \(360^\circ\).</p>
\[ \text{reflex } AOB = 360^\circ - {@cen@}^\circ = {@refl@}^\circ \]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[simp:true;
ang: rand(makelist(i,i,20,80));
cen: 2*ang;
refl: 360-cen;

pa: float((90+ang)*%pi/180);
pb: float((90-ang)*%pi/180);
xA: float(cos(pa));
yA: float(sin(pa));
xB: float(cos(pb));
yB: float(sin(pb));

img: plot(
  [[parametric, cos(s), sin(s), [s,0,2*%pi]],
   [discrete, [[xA,yA],[0,0],[xB,yB]]],
   [discrete, [[xA,yA],[0,-1],[xB,yB]]]],
  [x,-1.4,1.4], [y,-1.4,1.4],
  [box,false], [axes,false], [legend,false], [size,400,400],
  [label, ["A",1.15*xA,1.15*yA], ["B",1.15*xB,1.15*yB], ["C",0,-1.15], ["O",0.1,-0.12]]
);

ta1: cen;
ta2: refl;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}{@t@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>903107706</deployedseed>
    <deployedseed>58320113</deployedseed>
    <deployedseed>1622202458</deployedseed>
    <deployedseed>1111176232</deployedseed>
    <deployedseed>1734788791</deployedseed>
    <deployedseed>1524069297</deployedseed>
    <deployedseed>878461024</deployedseed>
    <deployedseed>205911477</deployedseed>
    <deployedseed>1100429436</deployedseed>
    <deployedseed>1838300744</deployedseed>
    <deployedseed>1476193021</deployedseed>
    <deployedseed>1885597890</deployedseed>
    <deployedseed>206070520</deployedseed>
    <deployedseed>2062717024</deployedseed>
    <deployedseed>2093986038</deployedseed>
    <deployedseed>150081348</deployedseed>
    <deployedseed>1053400946</deployedseed>
    <deployedseed>449193896</deployedseed>
    <deployedseed>432368436</deployedseed>
    <deployedseed>1894173929</deployedseed>
  </question>
</quiz>
