<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Limits of Polynomials — combining limits with the sum and constant-multiple laws</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Given that \( \displaystyle\lim_{x\to {@aa@}} f(x) = {@Lf@} \) and \( \displaystyle\lim_{x\to {@aa@}} g(x) = {@Lg@} \), evaluate</p>
<p>\[ \lim_{x\to {@aa@}} \big( {@ca@}f(x) + {@cb@}g(x) \big). \]</p>
<div>[[input:ans1]] [[validation:ans1]]</div>
<div>[[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The limit laws let you split this apart term by term:</p>
\[ \lim_{x\to {@aa@}}\big({@ca@}f(x)+{@cb@}g(x)\big) = {@ca@}\lim_{x\to {@aa@}}f(x) + {@cb@}\lim_{x\to {@aa@}}g(x) = {@ca@}({@Lf@})+{@cb@}({@Lg@}) = {@ta@}. \]
<p>Don't drop the constant multipliers {@ca@} and {@cb@}, and don't multiply the two limits together &#8212; here they are added, not multiplied.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>aa : rand_with_prohib(1,5,[]);
Lf : rand_with_prohib(-9,9,[0]);
Lg : rand_with_prohib(-9,9,[0]);
ca : rand_with_prohib(2,5,[]);
cb : rand_with_prohib(2,5,[]);
ta : ca*Lf+cb*Lg;
wrong_ignorecoef : Lf+Lg;                 /* ignored the constant multipliers */
wrong_multiplyall : ca*Lf*cb*Lg;          /* wrongly multiplied the two limits together */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.7 (Limits of Polynomials). Applies the sum rule and constant-multiple rule for limits; catches dropping the constant multipliers and wrongly multiplying the two limits together i</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.7 (Limits of Polynomials). Applies the sum rule and constant-multiple rule for limits; catches dropping the constant multipliers and wrongly multiplying the two limits together instead of adding.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correctly applied the sum/constant-multiple limit laws?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Dropped the constant multipliers?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_ignorecoef</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-ignorecoef</trueanswernote>
        <truefeedback format="html">
          <text>You've added \(L_f+L_g\) directly, but forgotten the constant multipliers {@ca@} and {@cb@} in front of \(f(x)\) and \(g(x)\).</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Multiplied the two limits together instead of adding?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>wrong_multiplyall</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-multiplied</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[The two scaled limits should be <strong>added</strong>, not multiplied together.]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>That is not the correct value. Apply the sum rule and the constant multiple rule for limits separately, then add.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Dropped coefficients</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_ignorecoef</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-ignorecoef</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Multiplied instead of added</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_multiplyall</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.2000000</expectedpenalty>
        <expectedanswernote>prt1-multiplied</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
