<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Limits of Polynomials — end behaviour at infinity</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the polynomial \( f(x) = {@p@}x^{@n@} + {@q@}x + {@r@} \).</p>
<p><strong>(a)</strong> What is \( \displaystyle\lim_{x \to \infty} f(x) \)?</p>
<div>[[input:ans1]] [[validation:ans1]]</div>
<p><strong>(b)</strong> What is \( \displaystyle\lim_{x \to -\infty} f(x) \)?</p>
<div>[[input:ans2]] [[validation:ans2]]</div>
<div>[[feedback:prt1]] [[feedback:prt2]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>For large \(|x|\), the leading term \({@p@}x^{@n@}\) dominates everything else. As \(x\to\infty\), the sign of the limit matches the sign of the leading coefficient \({@p@}\).</p>
<p>As \(x\to-\infty\): if the degree \(n={@n@}\) is <strong>even</strong>, \(x^{@n@}\) is positive for large negative \(x\) too, so the behaviour is the <em>same</em> as \(x\to\infty\). If the degree is <strong>odd</strong>, \(x^{@n@}\) is negative for large negative \(x\), so the behaviour <em>flips</em>.</p>
<p>Here \(n={@n@}\), which is {@parityword@}.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[p : rand_with_prohib(-5,5,[0]);
q : rand_with_prohib(1,9,[]);
r : rand_with_prohib(1,9,[]);
n : rand_with_prohib(2,4,[]);           /* even or odd degree, matters! */
n_even : evenp(n);
p_pos : is(p>0);
/* as x -> +infinity: sign follows sign of p, always */
lim_pos_inf : if p_pos then "posinf" else "neginf";
/* as x -> -infinity: if n even, same sign as leading behaviour at +inf; if n odd, opposite */
lim_neg_inf : if n_even then lim_pos_inf else (if lim_pos_inf="posinf" then "neginf" else "posinf");
opts : [["posinf", true, "+∞"],["neginf", true, "-∞"]];
parityword : if n_even then "even" else "odd";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.7 (Limits of Polynomials). End behaviour as x tends to +/- infinity; catches assuming the behaviour at -infinity always matches +infinity, ignoring that this only holds for even-</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.7 (Limits of Polynomials). End behaviour as x tends to +/- infinity; catches assuming the behaviour at -infinity always matches +infinity, ignoring that this only holds for even-degree polynomials.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>opts</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>opts</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description><![CDATA[Correct limit as x -> +infinity (matches sign of leading coefficient)?]]></description>
        <answertest>String</answertest>
        <sans>ans1</sans>
        <tans>lim_pos_inf</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>As x becomes very large and positive, the sign of the whole expression matches the sign of the leading coefficient.</text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>0.5000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description><![CDATA[Correct limit as x -> -infinity (accounting for degree parity)?]]></description>
        <answertest>String</answertest>
        <sans>ans2</sans>
        <tans>lim_neg_inf</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description><![CDATA[Used the same value as x->+infinity without checking degree parity?]]></description>
        <answertest>String</answertest>
        <sans>ans2</sans>
        <tans>lim_pos_inf</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-forgotparity</trueanswernote>
        <truefeedback format="html">
          <text>Did you assume the behaviour at \(-\infty\) is always the same as at \(+\infty\)? That is only true when the degree is even. Here the degree is {@n@}, which is {@parityword@}, so you need to check whether the sign flips.</text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>That is not the correct limit as x approaches negative infinity.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Both correct</description>
      <testinput>
        <name>ans1</name>
        <value>lim_pos_inf</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>lim_neg_inf</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Forgot that odd/even degree matters for -infinity</description>
      <testinput>
        <name>ans1</name>
        <value>lim_pos_inf</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>lim_pos_inf</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-forgotparity</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
