<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Types of Functions — identifying a function family from its graph (JSXGraph)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The graph of a function is shown below.</p>
[[jsxgraph width="420px" height="340px"]]
var scenario = {#scenario#};
var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [-5, 8, 5, -8],
    axis: true, showCopyright: false, showNavigation: false
});
var f;
if (scenario === 0) { f = board.jc.snippet('0.5*x+1', true, 'x', true); board.create('functiongraph',[f,-5,5],{strokeColor:'#1565c0',strokeWidth:2}); }
else if (scenario === 1) { f = board.jc.snippet('0.4*x^2-2', true, 'x', true); board.create('functiongraph',[f,-4.2,4.2],{strokeColor:'#1565c0',strokeWidth:2}); }
else if (scenario === 2) { f = board.jc.snippet('0.15*x^3', true, 'x', true); board.create('functiongraph',[f,-4,4],{strokeColor:'#1565c0',strokeWidth:2}); }
else { f = board.jc.snippet('2^x-3', true, 'x', true); board.create('functiongraph',[f,-5,3.5],{strokeColor:'#1565c0',strokeWidth:2}); }
[[/jsxgraph]]
<p>What type of function does this graph most likely represent?</p>
<div>[[input:ans1]] [[validation:ans1]]</div>
<div>[[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Key visual clues: a <strong>straight line</strong> is linear; a single <strong>symmetric U-shape</strong> (one turning point) is quadratic; an <strong>S-shaped curve</strong> with a turning point and an inflection (up to two turning points) is cubic; a curve that keeps rising (or falling) ever more steeply, flattening out toward a horizontal asymptote on one side, is exponential.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[scenario : rand(4);              /* 0 linear, 1 quadratic, 2 cubic, 3 exponential */
ta : if scenario=0 then "linear" elseif scenario=1 then "quadratic" elseif scenario=2 then "cubic" else "exponential";
opts : [["linear", scenario=0, "Linear"],["quadratic", scenario=1, "Quadratic"],["cubic", scenario=2, "Cubic"],["exponential", scenario=3, "Exponential"]];]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.5 (Types of Functions). Identifies linear, quadratic, cubic, or exponential function families from the shape of a JSXGraph-plotted curve.</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.5 (Types of Functions). Identifies linear, quadratic, cubic, or exponential function families from the shape of a JSXGraph-plotted curve.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>dropdown</type>
      <tans>opts</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correctly identified the function family from its graph?</description>
        <answertest>String</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrong</falseanswernote>
        <falsefeedback format="html">
          <text>Look again at the overall shape: how many turning points does it have, and does it flatten out toward an asymptote anywhere?</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Wrong family</description>
      <testinput>
        <name>ans1</name>
        <value><![CDATA[if ta="linear" then "quadratic" else "linear"]]></value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-wrong</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
