<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>AROC in context - value and sign interpretation (INNODEMS)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>{@contextintro@}</p>
\[ {@fname@}({@ivar@}) = {@fdisp@} \]
<p>where {@varmeaning@}.</p>
<p>Consider the interval from \( {@ivar@} = {@pp@} \) to \( {@ivar@} = {@qq@} \).</p>

<div><p>(a) Find the average rate of change of {@quantity@} over this interval, giving your answer exactly.</p>
<div>\( \text{Average rate of change} = \) [[input:ans1]] {@rateunit@} [[validation:ans1]]</div>
<div>[[feedback:prt1]]</div></div>

<div>(b) Over this interval, is {@quantity@} increasing or decreasing?
[[input:ans2]] [[validation:ans2]]
<div>[[feedback:prt2]]</div></div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The <strong>average rate of change</strong> of a function \(f\) over an interval \([p,q]\) is the slope of the secant line joining the two points on its graph:</p>
\[ \text{AROC} = \frac{f(q) - f(p)}{q - p}. \]
<p>Here \(f({@pp@}) = {@fp@}\) and \(f({@qq@}) = {@fq@}\), so</p>
\[ \text{AROC} = \frac{ {@fq@} - ({@fp@}) }{ {@qq@} - {@pp@} } = \frac{ {@fq-fp@} }{ {@qq-pp@} } = {@ta1@}\ \text{@rateunit@} \]
<p>Because the average rate of change is {@signword@} \(0\), {@quantity@} is <strong>{@dirword@}</strong> on average over this interval. {@interpretsentence@}</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025010100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Design backwards. Each context is
   [tag, A, B, C, p, q, fname, ivar, contextintro, varmeaning, quantity, rateunit].
   AROC over [p,q] for f = A x^2 + B x + C is A*(p+q)+B, chosen to be a tidy
   integer. The four families echo the textbook worked examples. The profit family
   uses A=-1 sampled past the vertex, so it supplies the guaranteed negative-rate
   variants. */

/* ---- Car: s(t)=A t^2 km, t hours; average speed (always positive). ---- */
Acar : rand([2,3,4,5]);
pcar : rand([1,2,3]);
qcar : pcar + rand([2,3,4]);
car : [ "car", Acar, 0, 0, pcar, qcar, "s", "t",
        "A car leaves a town. Its distance from the town (in kilometres) after \\(t\\) hours is modelled by",
        "\\(t\\) is the time in hours and \\(s\\) is the distance in kilometres",
        "the distance", "km/h" ];

/* ---- Profit: P(x)=-x^2+bp*x-cp (thousands KSh); sampled past vertex x=bp/2
   so AROC = -(p+q)+bp is negative. ---- */
bp  : rand([30,40,50,60]);
pp0 : bp/2 + rand([2,4,6]);       /* start past the vertex */
qp0 : pp0 + rand([4,6,8]);
cp  : rand([100,200,300,400]);
profit : [ "profit", -1, bp, -cp, pp0, qp0, "P", "x",
        "A small business models its monthly profit (in thousands of KSh) from selling \\(x\\) units as",
        "\\(x\\) is the number of units sold and \\(P\\) is the profit in thousands of KSh",
        "the profit", "thousand KSh per unit" ];

/* ---- Cost: C(x)=x^2-bc*x+dc KSh; sampled past vertex x=bc/2 so AROC>0. ---- */
bc  : rand([4,6,8]);
pc0 : bc/2 + rand([2,4,6]);
qc0 : pc0 + rand([3,5,7]);
dc  : rand([800,1000,1200]);
cost : [ "cost", 1, -bc, dc, pc0, qc0, "C", "x",
        "The cost (in KSh) of producing \\(x\\) bags of maize flour is modelled by",
        "\\(x\\) is the number of bags and \\(C\\) is the cost in KSh",
        "the cost", "KSh per bag" ];

/* ---- Revenue: R(x)=Ar x^2 + br x KSh; AROC = Ar*(p+q)+br > 0. ---- */
Ar  : rand([2,3]);
br  : rand([5,10,15]);
pr0 : rand([5,10]);
qr0 : pr0 + rand([5,10]);
revenue : [ "revenue", Ar, br, 0, pr0, qr0, "R", "x",
        "The revenue (in KSh) generated by selling \\(x\\) items is modelled by",
        "\\(x\\) is the number of items and \\(R\\) is the revenue in KSh",
        "the revenue", "KSh per item" ];

ctx : rand([car, profit, cost, revenue]);

/* unpack */
tag  : ctx[1];
AA   : ctx[2];
BB   : ctx[3];
CC   : ctx[4];
pp   : ctx[5];
qq   : ctx[6];
fname    : ctx[7];
ivar     : ctx[8];
contextintro : ctx[9];
varmeaning   : ctx[10];
quantity     : ctx[11];
rateunit     : ctx[12];

/* the model as a displayed expression in the context variable */
xv  : if is(ivar="t") then t else x;
fexpr : AA*xv^2 + BB*xv + CC;
fdisp : fexpr;

/* function values and the exact AROC (guaranteed integer, nonzero) */
fp  : AA*pp^2 + BB*pp + CC;
fq  : AA*qq^2 + BB*qq + CC;
ta1 : (fq - fp)/(qq - pp);          /* = AA*(pp+qq)+BB */

/* sign interpretation */
increasing_opt : "increasing";
decreasing_opt : "decreasing";
dirword  : if is(ta1>0) then increasing_opt else decreasing_opt;
signword : if is(ta1>0) then "greater than" else "less than";

/* Dropdown model answer: [value, correct?, display]. The stored value is the
   word itself, so the correct value equals dirword and prt2/qtests can compare
   against dirword / increasing_opt / decreasing_opt directly. */
ta2 : [ ["increasing", is(ta1>0), "increasing"],
        ["decreasing", is(ta1<0), "decreasing"] ];

/* a plain-language interpretation sentence for the worked solution */
interpretsentence : if is(tag="car")
   then "The average speed of the car over this interval is positive because it keeps moving away from the town."
   else if is(tag="profit")
   then "Beyond the most profitable output, producing more units lowers profit, so the average rate is negative."
   else if is(tag="cost")
   then "Producing more bags raises total cost, so the average cost per extra bag is positive."
   else "Selling more items raises revenue, so the average revenue per extra item is positive.";

/* wrong-answer variables for misconception diagnostics (must live here, not in
   feedback variables, so the qtests can reference them). */
wrong_divq  : (fq - fp)/qq;         /* divided by q instead of (q-p) */
wrong_sign  : ev(-ta1, simp);       /* sign slip in the numerator */]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( {@fname@}({@ivar@})={@fdisp@};\ [{@pp@},{@qq@}];\ \text{AROC}={@ta1@}\ ({@dirword@}) \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>7</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>dropdown</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct exact AROC.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.001</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Divided by q instead of (q - p).</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>wrong_divq</tans>
        <testoptions>0.001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-divq</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Check your denominator. The average rate of change divides the change in output by the <em>change</em> in input, \( q - p \), not by \( q \) alone.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Sign slip in the numerator (f(p) - f(q)).</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>wrong_sign</tans>
        <testoptions>0.001</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-sign</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Your value has the right size but the wrong sign. Keep the order consistent: the numerator is \( f(q) - f(p) \) over the denominator \( q - p \).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-incorrect</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Evaluate \(f\) at each endpoint, then compute \( \dfrac{f(q)-f(p)}{q-p} \). Substitute carefully and try again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct increasing/decreasing choice (ties to sign of AROC).</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>dirword</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-wrongsign</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Compare the sign of your average rate of change with \(0\). A <em>positive</em> average rate of change means the quantity is increasing on average; a <em>negative</em> one means it is decreasing. Match your choice to the sign of your answer in part (a).</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1001</deployedseed>
    <deployedseed>2002</deployedseed>
    <deployedseed>3003</deployedseed>
    <deployedseed>4004</deployedseed>
    <deployedseed>5005</deployedseed>
    <deployedseed>6006</deployedseed>
    <deployedseed>7007</deployedseed>
    <deployedseed>8008</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Both correct: exact AROC and the matching direction.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dirword</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Divided by q instead of (q - p); direction still correct.</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_divq</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dirword</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-divq</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Sign slip in the numerator; also picks the wrong direction.</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_sign</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value><![CDATA[if is(ta1>0) then decreasing_opt else increasing_opt]]></value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-sign</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-wrongsign</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Generic wrong value; correct direction.</description>
      <testinput>
        <name>ans1</name>
        <value>ev(ta1+7,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>dirword</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-incorrect</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description><![CDATA[Correct AROC but wrong direction chosen (interpretation slip: picks the opposite of the sign of the AROC). On a negative-rate variant this is "increasing"; on a positive-rate one it is "decreasing".]]></description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value><![CDATA[if is(ta1>0) then decreasing_opt else increasing_opt]]></value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-wrongsign</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
