<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Calculating coordinates of a point on a unit circle.</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Let \(P\) be the point where the terminal side of the angle \(\theta={@theta@}\) intersects the unit circle.</p>

<p>What are the coordinates \((x,y)\) of point \(P\)?</p>

<p>Enter your answer as an ordered pair.</p>

<p>[[input:ans1]] [[validation:ans1]][[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p>For a point \(P=(x,y)\) on the unit circle, the coordinates are determined by the angle \(\theta\):</p>

\(x=\cos(\theta), \qquad y=\sin(\theta).\)

<p>Therefore, the coordinates of \(P\) are</p>

\(P=(\cos(\theta),\sin(\theta)).\)

<p>To find the exact values, identify the reference angle and the quadrant in which \(\theta\) lies. The reference angle determines the magnitudes of the coordinates, while the quadrant determines their signs.</p>

<p>For example, if \(\theta=\frac{5\pi}{6}\), the reference angle is \(\frac{\pi}{6}\) and the angle lies in Quadrant II. Thus \(x\) is negative and \(y\) is positive:</p>

\(x=-\frac{\sqrt{3}}{2}, \qquad y=\frac12.\)

<p>Hence</p>

\(P=\left(-\frac{\sqrt{3}}{2},\frac12\right).\)]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>angles : [
    [%pi/6,   -sqrt(3)/2,  1/2],
    [%pi/4,   -sqrt(2)/2,  sqrt(2)/2],
    [%pi/3,   -1/2,        sqrt(3)/2],
    [2*%pi/3, -1/2,        sqrt(3)/2],
    [3*%pi/4, -sqrt(2)/2,  sqrt(2)/2],
    [5*%pi/6, -sqrt(3)/2,  1/2],
    [7*%pi/6, -sqrt(3)/2, -1/2],
    [5*%pi/4, -sqrt(2)/2, -sqrt(2)/2],
    [4*%pi/3, -1/2,       -sqrt(3)/2],
    [5*%pi/3,  1/2,       -sqrt(3)/2],
    [7*%pi/4,  sqrt(2)/2, -sqrt(2)/2],
    [11*%pi/6, sqrt(3)/2, -1/2]
];

variant : rand(angles);

theta : variant[1];
x0 : variant[2];
y0 : variant[3];

ta : [x0,y0];
/*possible student mistakes*/

ta_mistake2 : [-x0,y0];
ta_mistake3 : [x0,-y0];</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>1</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Using a an incorrect sign on the x coordinate</description>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta_mistake2 </tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><strong>There is a sign error in your \(x\)-coordinate.</strong></p>

<p>
Your answer was
\(
{@ans1@}.
\)
</p>

<p>
The angle selected for this question is
\(
\theta={@theta@}.
\)
For this angle, the corresponding point on the unit circle is
\(
(x_0,y_0)=({@x0@},{@y0@}).
\)
</p>

<p>
Therefore, the correct coordinate pair is
\(
\boxed{[{@x0@},{@y0@}]}.
\)
</p>

<p>
Your answer has the opposite sign for the \(x\)-coordinate:
\(
[{@-x0@},{@y0@}].
\)
This means that you have selected the point on the opposite side of the
\(y\)-axis.
</p>

<p>
Remember:
the \(x\)-coordinate is negative on the left-hand side of the
coordinate plane and positive on the right-hand side.
</p>

<p>
For the angle
\(
\theta={@theta@},
\)
the point lies in the quadrant where \(x={@x0@}\) and \(y={@y0@}\).
Hence,
\(
\boxed{[{@x0@},{@y0@}]}
\)
is the required answer.
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>EqualComAss</answertest>
        <sans>ans1</sans>
        <tans>ta_mistake3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-3-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p><strong>There is a sign error in your \(y\)-coordinate.</strong></p>

<p>
Your answer was
\(
{@ans1@}.
\)
</p>

<p>
The angle selected for this question is
\(
\theta={@theta@}.
\)
For this angle, the corresponding point on the unit circle is
\(
(x_0,y_0)=({@x0@},{@y0@}).
\)
</p>

<p>
Therefore, the correct coordinate pair is
\(
\boxed{[{@x0@},{@y0@}]}.
\)
</p>

<p>
Your answer has the opposite sign for the \(y\)-coordinate:
\(
[{@x0@},{@-y0@}].
\)
This means that you have selected the point on the opposite side of
the \(x\)-axis.
</p>

<p>
Remember:
the \(y\)-coordinate is positive above the \(x\)-axis and negative
below the \(x\)-axis.
</p>

<p>
For the angle
\(
\theta={@theta@},
\)
the point lies in the quadrant where
\(x={@x0@}\) and \(y={@y0@}\).
Hence the correct answer is
\(
\boxed{[{@x0@},{@y0@}]}.
\)
</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <qtest>
      <testcase>1</testcase>
      <description/>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
