<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Adding Irrational Numbers</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Simplify the following expression, giving your answer in the form \(k\sqrt{m}\), where \(m\) is as small as possible.</p>
\[ \sqrt{@x@} + \sqrt{@y@} \]
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>
<p><strong>Step 1: Simplify the first root.</strong> Find a square number that divides \(\sqrt{@x@}\):</p>
\[ \sqrt{@x@} = \sqrt{{@a2@} \times {@n@}} = \sqrt{@a2@} \times \sqrt{@n@} = {@a@}\sqrt{@n@} \]

<p><strong>Step 2: Simplify the second root.</strong> Find a square number that divides \(\sqrt{@y@}\):</p>
\[ \sqrt{@y@} = \sqrt{{@b2@} \times {@n@}} = \sqrt{@b2@} \times \sqrt{@n@} = {@b@}\sqrt{@n@} \]

<p><strong>Step 3: Add the like terms.</strong> Both roots now have the same irrational part \(\sqrt{@n@}\), so we add the coefficients:</p>
\[ {@a@}\sqrt{@n@} + {@b@}\sqrt{@n@} = ({@a@} + {@b@})\sqrt{@n@} = {@c@}\sqrt{@n@} \]

<p><strong>Answer:</strong> \( {@c@}\sqrt{@n@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>simp:true;
n: rand([2,3,5,6,7,10]);
a: rand([2,3,4]);
b: rand(delete(a,[2,3,4,5]));
a2: a^2;
b2: b^2;
x: a2*n;
y: b2*n;
c: a+b;
ta1: c*sqrt(n);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>299449144</deployedseed>
    <deployedseed>104025983</deployedseed>
    <deployedseed>428634585</deployedseed>
    <deployedseed>516604977</deployedseed>
    <deployedseed>998082319</deployedseed>
    <deployedseed>1444542724</deployedseed>
    <deployedseed>138157955</deployedseed>
    <deployedseed>371058</deployedseed>
    <deployedseed>1640713448</deployedseed>
    <deployedseed>1157552018</deployedseed>
    <deployedseed>1435147743</deployedseed>
    <deployedseed>35998818</deployedseed>
    <deployedseed>1791318300</deployedseed>
    <deployedseed>1595302255</deployedseed>
    <deployedseed>1334507987</deployedseed>
    <deployedseed>1135066240</deployedseed>
  </question>
</quiz>
