<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Adding and subtracting surds</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Simplify the following, giving your answer in the form \( a\sqrt{b} \), where \(b\) has no square factors:</p>
<p>
\[ {@A@}\sqrt{ {@m1@} } \;{@signchar@}\; {@Bmag@}\sqrt{ {@m2@} } \]
</p>
<p>
\( a = \) [[input:ans1]] ; \( \sqrt{b} \), \( b = \) [[input:ans2]]
</p>
<p>[[validation:ans1]] [[validation:ans2]]</p>
<p></p><p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<![CDATA[ 
<p>First simplify each surd separately, by finding the largest perfect-square factor of each number under the root:</p>
<p>
\[ \sqrt{ {@m1@} } = \sqrt{ {@(a1^2)@} \times {@p@} } = {@a1@}\sqrt{ {@p@} } \]
</p>
<p>
\[ \sqrt{ {@m2@} } = \sqrt{ {@(a2^2)@} \times {@p@} } = {@a2@}\sqrt{ {@p@} } \]
</p>
<p>Substitute these back into the original expression:</p>
<p>
\[
{@A@}\sqrt{ {@m1@} } \;{@signchar@}\; {@Bmag@}\sqrt{ {@m2@} }
= {@A@}\times {@a1@}\sqrt{ {@p@} } \;{@signchar@}\; {@Bmag@}\times {@a2@}\sqrt{ {@p@} }
= {@(A*a1)@}\sqrt{ {@p@} } \;{@signchar@}\; {@(Bmag*a2)@}\sqrt{ {@p@} }
\]
</p>
<p>Both terms now have the same surd part, \( \sqrt{ {@p@} } \), so they are &ldquo;like surds&rdquo; and their coefficients can be combined:</p>
<p>
\[ = {@coeff@}\sqrt{ {@p@} } \]
</p>
 ]]]]><![CDATA[>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[plist: [2,3,5,6,7,10,11,13,14,15,17,19,21,22,23];
p: plist[1+rand(length(plist))];
a1: 1+rand(4);
a2: 1+rand(4);
while a2=a1 do a2: 1+rand(4);
A: 1+rand(6);
Bmag: 1+rand(6);
sgn: if rand(2)=0 then 1 else -1;
m1: a1^2*p;
m2: a2^2*p;
coeff: A*a1 + sgn*Bmag*a2;
while coeff=0 do (
  sgn: if rand(2)=0 then 1 else -1,
  Bmag: 1+rand(6),
  coeff: A*a1 + sgn*Bmag*a2
);
signchar: if sgn=1 then "+" else "-";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@coeff@}sqrt({@p@}){@p@}({@a1@}){@a2@}({@A@}){@Bmag@}({@sgn@})</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>coeff</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>p</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>coeff</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>787506738</deployedseed>
    <deployedseed>238958275</deployedseed>
    <deployedseed>119473212</deployedseed>
    <deployedseed>1734062193</deployedseed>
    <deployedseed>247217039</deployedseed>
    <deployedseed>246782825</deployedseed>
    <deployedseed>1673680581</deployedseed>
    <deployedseed>65635654</deployedseed>
    <deployedseed>1746781382</deployedseed>
    <deployedseed>350494984</deployedseed>
    <deployedseed>1737849700</deployedseed>
    <deployedseed>1864361363</deployedseed>
    <deployedseed>162837131</deployedseed>
    <deployedseed>1211333129</deployedseed>
    <deployedseed>230818370</deployedseed>
    <deployedseed>86325304</deployedseed>
    <deployedseed>1801398813</deployedseed>
    <deployedseed>2076187930</deployedseed>
    <deployedseed>915374806</deployedseed>
    <deployedseed>1615159368</deployedseed>
  </question>
</quiz>
