<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>angles subtended by the same diameter</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>The diagram shows a circle with centre \(O\). \(AB\) is a diameter, and \(C\) and \(D\) are points on the circumference.
Angle \(CAB = {@a@}^\circ\).</p>

[[jsxgraph width="380px" height="380px"]]
var board = JXG.JSXGraph.initBoard(divid, {boundingbox:[-1.4,1.4,1.4,-1.4], axis:false, showCopyright:false, showNavigation:false, keepAspectRatio:true});
var t = {#a#}*Math.PI/180;
var d = -110*Math.PI/180;
var O = board.create('point',[0,0],{name:'O',fixed:true,size:2,label:{offset:[6,-14]}});
var A = board.create('point',[-1,0],{name:'A',fixed:true,label:{offset:[-16,-4]}});
var B = board.create('point',[1,0],{name:'B',fixed:true,label:{offset:[8,-4]}});
var C = board.create('point',[Math.cos(2*t),Math.sin(2*t)],{name:'C',fixed:true,label:{offset:[6,6]}});
var D = board.create('point',[Math.cos(d),Math.sin(d)],{name:'D',fixed:true,label:{offset:[-4,-16]}});
board.create('circle',[O,A],{fixed:true,strokeColor:'#333'});
board.create('segment',[A,B],{fixed:true,strokeColor:'#333'});
board.create('segment',[A,C],{fixed:true,strokeColor:'#333'});
board.create('segment',[C,B],{fixed:true,strokeColor:'#333'});
board.create('segment',[A,D],{fixed:true,strokeColor:'#333'});
board.create('segment',[D,B],{fixed:true,strokeColor:'#333'});
[[/jsxgraph]]

<p>(a) Find angle \(ACB\). [[input:ans1]] [[validation:ans1]]</p>
<div>[[feedback:prt1]]</div>

<p>(b) Find angle \(ABC\). [[input:ans2]] [[validation:ans2]]</p>
<div>[[feedback:prt2]]</div>

<p>(c) Find angle \(ADB\). [[input:ans3]] [[validation:ans3]]</p>
<div>[[feedback:prt3]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h4>Worked solution</h4>
<p><strong>(a)</strong> \(AB\) is a diameter, so the angle in a semicircle is a right angle:
\(ACB = 90^\circ\).</p>
<p><strong>(b)</strong> The angles in triangle \(ABC\) add up to \(180^\circ\):
\(ABC = 180 - 90 - {@a@} = {@ta2@}^\circ\).</p>
<p><strong>(c)</strong> \(D\) is also on the circumference, so \(AB\) subtends a right angle at \(D\) too:
\(ADB = 90^\circ\), whichever side of the diameter the point lies.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>a: rand([20,25,30,35,40,50,55,65]);
ta1: 90;
ta2: 90-a;
ta3: 90;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@},a = {@a@}, d = {@d@},{@ta2@},{@ta3@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>826958706</deployedseed>
    <deployedseed>2057194529</deployedseed>
    <deployedseed>1973949024</deployedseed>
    <deployedseed>594417850</deployedseed>
    <deployedseed>1189041869</deployedseed>
    <deployedseed>199956645</deployedseed>
    <deployedseed>47443320</deployedseed>
    <deployedseed>2078936213</deployedseed>
  </question>
</quiz>
