<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Domain of a rational function — finding the excluded values</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the function</p>
<p>\[ g(x) = \dfrac{ {@expr_num@} }{ {@expr_den@} }. \]</p>
<p>State the value(s) of \(x\) that must be <strong>excluded</strong> from the domain of \(g\).</p>
<p>Give your answer as a set, e.g. <code>{2,-3}</code>, or <code>{}</code> if there are none.</p>
<div>Excluded values \( = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>A rational function is undefined only where its denominator is zero, so we set the denominator equal to zero and solve:</p>
\[
\begin{align*}
{@expr_den@} &= 0\\
(x-{@b@})(x+{@b@}) &= 0
\end{align*}
\]
<p>This gives \(x = {@b@}\) or \(x = -{@b@}\). The numerator does not affect the domain here, since \(x+{@a@} = 0\) only tells us where \(g(x)=0\), not where it is undefined.</p>
<p>The excluded values are \( \{ {@b@}, -{@b@} \} \).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>b : rand_with_prohib(2,9,[]);            /* denominator root, e.g. b=5 gives x^2-25 */
a : rand_with_prohib(1,9,[b]);            /* numerator root; a != b so nothing cancels */
expr_num : x+a;
expr_den : x^2-b^2;
expr : expr_num/expr_den;
ta : {-b,b};                              /* correct excluded-value set */
wrong_pos : {b};                          /* forgot the negative root */
wrong_numerator : {-a};                   /* set the numerator to zero instead of the denominator */
wrong_nosqrt : {b^2};                     /* forgot to take the square root of b^2 */
wrong_signflip : {b,b};                   /* sign error: used +b twice */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>g(x) = ({@expr_num@})/({@expr_den@}); excluded values {@ta@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions, Section 16.3 (Domain). Tests the common misconception of setting the numerator equal to zero, or forgetting the negative root of a difference of squares.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>{}</syntaxhint>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords>diff,solve</forbidwords>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Is the answer the correct set of excluded values?</description>
        <answertest>Sets</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-0-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Forgot the negative root.</description>
        <answertest>Sets</answertest>
        <sans>ans1</sans>
        <tans>wrong_pos</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0.5</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-onlypositive</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You found one of the two values, but a difference of squares \(x^2-{@b@}^2=0\) has <em>two</em> solutions: don't forget \(x=-{@b@}\) as well as \(x={@b@}\).</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description>Used the numerator instead of the denominator.</description>
        <answertest>Sets</answertest>
        <sans>ans1</sans>
        <tans>wrong_numerator</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-numerator</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you set the <strong>numerator</strong> equal to zero. That tells you where \(g(x)=0\), not where \(g\) is undefined. The domain is restricted by the <strong>denominator</strong>: set \({@expr_den@}=0\) instead.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>3</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>3</name>
        <description>Forgot to take the square root of b^2.</description>
        <answertest>Sets</answertest>
        <sans>ans1</sans>
        <tans>wrong_nosqrt</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-nosqrt</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>You correctly identified \(x^2 = {@b@}^2\), but stopped one step early. You still need to take the square root of both sides to find the two values of \(x\) itself.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-wrongvalue</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>That is not the correct set of excluded values. See the worked solution above.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Fully correct answer</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-0-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Forgot the negative root</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_pos</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.5000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-onlypositive</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Used the numerator instead of the denominator</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_numerator</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.2000000</expectedpenalty>
        <expectedanswernote>prt1-numerator</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>Forgot to take the square root (gave b^2)</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_nosqrt</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.3000000</expectedpenalty>
        <expectedanswernote>prt1-nosqrt</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>5</testcase>
      <description>Generic wrong value</description>
      <testinput>
        <name>ans1</name>
        <value>{0}</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.4000000</expectedpenalty>
        <expectedanswernote>prt1-wrongvalue</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>6</testcase>
      <description>Empty set (student thinks there is no restriction)</description>
      <testinput>
        <name>ans1</name>
        <value>{}</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.4000000</expectedpenalty>
        <expectedanswernote>prt1-wrongvalue</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
