<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Domain — a square root inside a denominator (strict inequality)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the function</p>
<p>\[ h(x) = \dfrac{1}{\sqrt{x-{@a@}}}. \]</p>
<p>State the domain of \(h\) as an inequality in \(x\).</p>
<div>\(x\) [[input:ans_sym]] [[validation:ans_sym]] {@a@}</div>
<div>[[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Here there are <strong>two</strong> restrictions stacked on top of each other: the expression under the square root must be \(\geq 0\), <em>and</em> the whole square root sits in a denominator so it cannot be \(0\) either. Combining \(x-{@a@}\geq 0\) with \(x-{@a@}\neq 0\) gives the strict inequality \(x-{@a@}\gt 0\), i.e. \(x \gt {@a@}\).</p>
<p>A very common mistake is to only apply the square-root rule (\(\geq\)) and forget that dividing by zero is also forbidden.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[a : rand_with_prohib(1,9,[]);
gt_opt : "gt"; ge_opt : "ge";
opts : [[gt_opt, true, "\\gt"],[ge_opt, false, "\\geq"]];
ta : gt_opt;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.3 (Domain). Tests the combined restriction from a square root sitting inside a denominator, which forces a strict inequality rather than the usual >= from the square root rule al</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.3 (Domain). Tests the combined restriction from a square root sitting inside a denominator, which forces a strict inequality rather than the usual >= from the square root rule alone.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans_sym</name>
      <type>dropdown</type>
      <tans>opts</tans>
      <boxsize>8</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description><![CDATA[Correctly used strict > (not >=)?]]></description>
        <answertest>String</answertest>
        <sans>ans_sym</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-onlysqrt</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[You've applied the square-root rule (\(\geq 0\)) but forgotten that \(h(x)\) also has this square root in the <strong>denominator</strong>, so it cannot equal zero either. Combining both restrictions gives a strict inequality.]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description><![CDATA[Correct: strict >]]></description>
      <testinput>
        <name>ans_sym</name>
        <value>gt_opt</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description><![CDATA[Forgot the denominator restriction, used >=]]></description>
      <testinput>
        <name>ans_sym</name>
        <value>ge_opt</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-onlysqrt</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
