<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>The Ferris wheel (angle at the centre is twice the angle at the circumference)</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A Ferris wheel has its centre at \(O\). Two cabins, \(A\) and \(B\), are on the rim of the wheel.
A rider at point \(C\) on the rim sees the two cabins, and angle \(ACB = {@c@}^\circ\).
The arms \(OA\) and \(OB\) are equal in length.</p>

[[jsxgraph width="380px" height="400px"]]
var board = JXG.JSXGraph.initBoard(divid, {boundingbox:[-1.6,1.6,1.6,-1.8], axis:false, showCopyright:false, showNavigation:false, keepAspectRatio:true});
var r = Math.PI/180;
var c = {#c#}, cc = {#cc#};
function pt(nm, deg, col) {
  var t = deg*r;
  return board.create('point',[Math.cos(t),Math.sin(t)],{name:nm,fixed:true,size:5,strokeColor:col,fillColor:col,
    label:{offset:[16*Math.cos(t)-4,16*Math.sin(t)-4],strokeColor:col}});
}
// ground and stand
board.create('segment',[[-1.4,-1.6],[1.4,-1.6]],{fixed:true,strokeColor:'#16a34a',strokeWidth:8,highlight:false});
board.create('segment',[[-0.7,-1.6],[0,0]],{fixed:true,strokeColor:'#92400e',strokeWidth:5,highlight:false});
board.create('segment',[[0.7,-1.6],[0,0]],{fixed:true,strokeColor:'#92400e',strokeWidth:5,highlight:false});
// wheel
board.create('circle',[[0,0],1],{fixed:true,strokeColor:'#1d4ed8',strokeWidth:4,fillColor:'#bfdbfe',fillOpacity:0.5,highlight:false});
var O = board.create('point',[0,0],{name:'O',fixed:true,size:4,strokeColor:'#1e293b',fillColor:'#1e293b',label:{offset:[8,-16]}});
var A = pt('A', 90+c, '#dc2626');
var B = pt('B', 90-c, '#dc2626');
var C = pt('C', cc, '#16a34a');
board.create('segment',[O,A],{fixed:true,strokeColor:'#dc2626',strokeWidth:3});
board.create('segment',[O,B],{fixed:true,strokeColor:'#dc2626',strokeWidth:3});
board.create('segment',[C,A],{fixed:true,strokeColor:'#ea580c',strokeWidth:3});
board.create('segment',[C,B],{fixed:true,strokeColor:'#ea580c',strokeWidth:3});
board.create('segment',[A,B],{fixed:true,strokeColor:'#7c3aed',strokeWidth:2,dash:2});
board.create('angle',[B,C,A],{radius:0.3,fillColor:'#facc15',fillOpacity:0.8,strokeColor:'#ca8a04',withLabel:false});
board.create('angle',[B,O,A],{radius:0.3,fillColor:'#fca5a5',fillOpacity:0.8,strokeColor:'#dc2626',withLabel:false});
[[/jsxgraph]]

<p>(a) Find angle \(AOB\). [[input:ans1]] [[validation:ans1]]</p>
<div>[[feedback:prt1]]</div>

<p>(b) Find angle \(OAB\). [[input:ans2]] [[validation:ans2]]</p>
<div>[[feedback:prt2]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h4>Worked solution</h4>
<p><strong>(a)</strong> The angle at the centre is twice the angle at the circumference when both stand on the same arc:
\(AOB = 2 \times {@c@} = {@ta1@}^\circ\).</p>
<p><strong>(b)</strong> \(OA = OB\) because they are radii, so triangle \(OAB\) is isosceles with equal base angles:
\(OAB = (180 - {@ta1@}) \div 2 = {@ta2@}^\circ\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>c: rand_with_step(20,60,5);
cc: rand_with_step(250,290,10);
ta1: 2*c;
ta2: 90-c;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@},{@ta2@},c = {@c@}, cc = {@cc@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1069207007</deployedseed>
    <deployedseed>278116389</deployedseed>
    <deployedseed>6993729</deployedseed>
    <deployedseed>345945981</deployedseed>
    <deployedseed>1413910182</deployedseed>
    <deployedseed>739381681</deployedseed>
    <deployedseed>211640661</deployedseed>
    <deployedseed>1338217075</deployedseed>
    <deployedseed>2084052915</deployedseed>
    <deployedseed>487401740</deployedseed>
    <deployedseed>1353076892</deployedseed>
    <deployedseed>83143676</deployedseed>
    <deployedseed>789552058</deployedseed>
    <deployedseed>901319377</deployedseed>
    <deployedseed>147289704</deployedseed>
    <deployedseed>1772650085</deployedseed>
    <deployedseed>913054455</deployedseed>
    <deployedseed>1466959640</deployedseed>
    <deployedseed>1820282408</deployedseed>
    <deployedseed>1903828976</deployedseed>
  </question>
</quiz>
