<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Forming and solving a compound inequality for a greenhouse temperature range</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>{@fname@}, a farmer in {@place@}, needs to keep the indoor temperature (\(T\)) of a greenhouse such that it is strictly greater than <b>\({@L@}^\circ\text{C}\)</b> and less than or equal to <b>\({@U@}^\circ\text{C}\)</b>.</p>

<p>An automatic heater in the greenhouse always adds exactly <b>\({@H@}^\circ\text{C}\)</b> to the current outdoor temperature (\(t\)). Therefore, we can say \(T = t + {@H@}\).</p>

<p>By substituting \(T\) into the compound inequality, determine the range of outdoor temperatures (\(t\)) that will keep the greenhouse in the ideal condition.</p>
<p><i>(Enter the lower and upper bounds of \(t\) in the boxes below.)</i></p>

<p>[[input:ans1]] \( < t \le \) [[input:ans2]]</p>

<div><b>Lower bound:</b> [[validation:ans1]] [[feedback:prt1]]</div>
<div><b>Upper bound:</b> [[validation:ans2]] [[feedback:prt2]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<p><b>Worked Solution:</b></p>

<p>First, we translate the initial temperature constraints into a single compound inequality. The indoor temperature \(T\) must be greater than \({@L@}\) and less than or equal to \({@U@}\):</p>
<p style="margin-left: 20px;">
\( \displaystyle {@L@} < T \le {@U@} \)
</p>

<p>Next, we substitute the expression for the indoor temperature (\(T = t + {@H@}\)) into the middle of the inequality:</p>
<p style="margin-left: 20px;">
\( \displaystyle {@L@} < t + {@H@} \le {@U@} \)
</p>

<p>To isolate the outdoor temperature \(t\), we subtract the heater's contribution (\({@H@}\)) from all three parts of the inequality:</p>
<p style="margin-left: 20px;">
\( \displaystyle \begin{aligned}
{@L@} - {@H@} &< t \le {@U@} - {@H@} \\
{@ta1@} &< t \le {@ta2@}
\end{aligned} \)
</p>

<p>Hence, the required range for the outdoor temperature is:</p>
<p style="margin-left: 20px;">
\( \displaystyle \boxed{{@ta1@} < t \le {@ta2@}} \)
</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026010500</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Local Context */
names : ["Mwakidize", "Nafula", "Kiplagat", "Mutua", "Akinyi"];
fname : rand(names);
places : ["Molo", "Kinangop", "Timau", "Kieni"];
place : rand(places);

/* Temperature Bounds */
L : rand_with_step(12, 16, 1);
U : rand_with_step(24, 29, 1);
H : rand_with_step(8, 12, 1);

/* Target Answers */
ta1 : L - H;
ta2 : U - H;

/* Mistake Catchers */
/* Mistake: Added the heater temp instead of subtracting */
m1_L : L + H;
m1_U : U + H;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>\({@L@} < t + {@H@} \le {@U@} \implies {@ta1@} < t \le {@ta2@}\)</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="concept">Compound Inequalities (Real Life Constraints)</p>
<p id="description">Grade 11 CBE: Forming a continuous compound inequality from a localized word problem, applying algebraic substitution, and solving for bounds.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta1</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta2</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check absolute correctness</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Check if they added the heater temp instead of subtracting</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>m1_L</tans>
        <testoptions>0.01</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>To isolate \(t\) in the expression \(t + {@H@}\), you must <b>subtract</b> \({@H@}\) from all parts of the inequality, not add it.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Check absolute correctness</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Check if they added the heater temp instead of subtracting</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>m1_U</tans>
        <testoptions>0.01</testoptions>
        <quiet>1</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>To isolate \(t\) in the expression \(t + {@H@}\), you must <b>subtract</b> \({@H@}\) from all parts of the inequality, not add it.</p>]]></text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1454926149</deployedseed>
    <deployedseed>1999632484</deployedseed>
    <deployedseed>1375373389</deployedseed>
    <deployedseed>994523405</deployedseed>
    <deployedseed>1369784828</deployedseed>
    <deployedseed>1002172452</deployedseed>
    <deployedseed>222143592</deployedseed>
    <deployedseed>1828944267</deployedseed>
    <deployedseed>2108853083</deployedseed>
    <deployedseed>1694082361</deployedseed>
    <deployedseed>1733787698</deployedseed>
    <deployedseed>1658506047</deployedseed>
    <deployedseed>68839506</deployedseed>
    <deployedseed>571894848</deployedseed>
    <deployedseed>116135926</deployedseed>
    <deployedseed>108408297</deployedseed>
    <deployedseed>268454280</deployedseed>
    <deployedseed>1194791708</deployedseed>
    <deployedseed>739746763</deployedseed>
    <deployedseed>1302944130</deployedseed>
    <deployedseed>192937883</deployedseed>
    <deployedseed>441761376</deployedseed>
    <deployedseed>764854639</deployedseed>
    <deployedseed>1317120740</deployedseed>
    <deployedseed>2014565290</deployedseed>
    <deployedseed>1583089034</deployedseed>
    <deployedseed>1767654796</deployedseed>
    <deployedseed>917584765</deployedseed>
    <deployedseed>919873200</deployedseed>
    <deployedseed>2040763991</deployedseed>
    <deployedseed>1692829419</deployedseed>
    <deployedseed>1472871576</deployedseed>
    <deployedseed>1610781339</deployedseed>
    <deployedseed>1832781561</deployedseed>
    <deployedseed>1284431028</deployedseed>
    <deployedseed>1731704269</deployedseed>
    <deployedseed>671938933</deployedseed>
    <deployedseed>517385863</deployedseed>
    <deployedseed>1712356028</deployedseed>
    <deployedseed>1093513250</deployedseed>
    <deployedseed>569869669</deployedseed>
    <deployedseed>887813112</deployedseed>
    <deployedseed>1277219262</deployedseed>
    <deployedseed>1549617572</deployedseed>
    <deployedseed>43623716</deployedseed>
    <deployedseed>271287039</deployedseed>
    <deployedseed>608841263</deployedseed>
    <deployedseed>76278248</deployedseed>
    <deployedseed>659590464</deployedseed>
    <deployedseed>206120969</deployedseed>
    <deployedseed>1463513842</deployedseed>
    <deployedseed>1293182562</deployedseed>
    <deployedseed>118916200</deployedseed>
    <deployedseed>605603446</deployedseed>
    <deployedseed>1659639760</deployedseed>
    <deployedseed>2016594371</deployedseed>
    <deployedseed>470101878</deployedseed>
    <deployedseed>843014866</deployedseed>
    <deployedseed>2075106770</deployedseed>
    <deployedseed>1199840090</deployedseed>
    <deployedseed>1860449808</deployedseed>
    <deployedseed>1795070956</deployedseed>
    <deployedseed>2094876490</deployedseed>
    <deployedseed>90649697</deployedseed>
    <deployedseed>1153318542</deployedseed>
    <deployedseed>1208246688</deployedseed>
    <deployedseed>1787366843</deployedseed>
    <deployedseed>1475396986</deployedseed>
    <deployedseed>850085245</deployedseed>
    <deployedseed>1473487282</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Both limits correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-1-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-1-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Added instead of subtracted</description>
      <testinput>
        <name>ans1</name>
        <value>m1_L</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>m1_U</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-T</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-2-T</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Generic wrong answers</description>
      <testinput>
        <name>ans1</name>
        <value>ta1+1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2+1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-2-F</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt2-2-F</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
