<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Intersecting chords - solve for x, then find AP and PD (INNODEMS)</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>In the figure, the chords \(AB\) and \(CD\) of the circle intersect at \(P\). The lengths are</p>
\[ AP = ({@eAP@})\text{ cm}, \quad PB = {@PB@}\text{ cm}, \quad CP = {@CP@}\text{ cm}, \quad PD = ({@ePD@})\text{ cm}. \]

[[jsxgraph width="470px" height="470px"]]
var AP = {#ta_AP#}, PB = {#PB#}, CP = {#CP#}, PD = {#ta_PD#};
var b1 = {#b1#}, b2 = {#b2#};

var rad = Math.PI / 180;
var aAB = 20 * rad, aCD = 75 * rad;
var ux = Math.cos(aAB), uy = Math.sin(aAB);
var vx = Math.cos(aCD), vy = Math.sin(aCD);

var Ap = [-AP * ux, -AP * uy], Bp = [PB * ux, PB * uy];
var Cp = [-CP * vx, -CP * vy], Dp = [PD * vx, PD * vy];

function circum(p1, p2, p3) {
    var x1 = p1[0], y1 = p1[1], x2 = p2[0], y2 = p2[1], x3 = p3[0], y3 = p3[1];
    var d = 2 * (x1 * (y2 - y3) + x2 * (y3 - y1) + x3 * (y1 - y2));
    var cx = ((x1*x1 + y1*y1) * (y2 - y3) + (x2*x2 + y2*y2) * (y3 - y1) + (x3*x3 + y3*y3) * (y1 - y2)) / d;
    var cy = ((x1*x1 + y1*y1) * (x3 - x2) + (x2*x2 + y2*y2) * (x1 - x3) + (x3*x3 + y3*y3) * (x2 - x1)) / d;
    return [cx, cy];
}
var ctr = circum(Ap, Bp, Cp);
var rr  = Math.hypot(ctr[0] - Ap[0], ctr[1] - Ap[1]);
var pad = 1.6;

var board = JXG.JSXGraph.initBoard(divid, {
    boundingbox: [ctr[0] - rr - pad, ctr[1] + rr + pad, ctr[0] + rr + pad, ctr[1] - rr - pad],
    keepaspectratio: true, axis: false, showCopyright: false, showNavigation: false
});

board.create('circle', [ctr, rr], {strokeColor: '#888', strokeWidth: 1.5, fixed: true});

var A = board.create('point', Ap, {name: 'A', size: 2, fixed: true, fillColor: '#333', strokeColor: '#333', label: {offset: [-12, -2], anchorX: 'right', anchorY: 'middle'}});
var B = board.create('point', Bp, {name: 'B', size: 2, fixed: true, fillColor: '#333', strokeColor: '#333', label: {offset: [10, 8], anchorX: 'left', anchorY: 'bottom'}});
var C = board.create('point', Cp, {name: 'C', size: 2, fixed: true, fillColor: '#333', strokeColor: '#333', label: {offset: [-10, -8], anchorX: 'right', anchorY: 'top'}});
var D = board.create('point', Dp, {name: 'D', size: 2, fixed: true, fillColor: '#333', strokeColor: '#333', label: {offset: [0, 12], anchorX: 'middle', anchorY: 'bottom'}});
var P = board.create('point', [0, 0], {name: 'P', size: 2, fixed: true, fillColor: '#333', strokeColor: '#333', label: {offset: [9, -4], anchorX: 'left', anchorY: 'top'}});

board.create('segment', [A, B], {strokeColor: '#1565c0', strokeWidth: 2.5, fixed: true});
board.create('segment', [C, D], {strokeColor: '#d32f2f', strokeWidth: 2.5, fixed: true});

var nAB = [-uy, ux], nCD = [-vy, vx], off = 0.6;
function lbl(p1, p2, n, sgn, txt) {
    board.create('text',
        [(p1[0] + p2[0]) / 2 + sgn * off * n[0], (p1[1] + p2[1]) / 2 + sgn * off * n[1], txt],
        {anchorX: 'middle', anchorY: 'middle', fontSize: 13, fixed: true, cssStyle: 'color:#333'});
}
function xexp(b) { return '(x' + (b > 0 ? ' + ' + b : ' - ' + (-b)) + ') cm'; }

lbl(Ap, [0, 0], nAB,  1, xexp(b1));
lbl([0, 0], Bp, nAB, -1, PB + ' cm');
lbl(Cp, [0, 0], nCD,  1, CP + ' cm');
lbl([0, 0], Dp, nCD,  1, xexp(b2));
[[/jsxgraph]]

<div>(a) Find \( x = \) [[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</div>
<div>(b) Hence find \( AP = \) [[input:ans2]]\(\text{ cm}\) [[validation:ans2]] [[feedback:prt2]]</div>
<div>(c) and \( PD = \) [[input:ans3]]\(\text{ cm}\) [[validation:ans3]] [[feedback:prt3]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>When two chords intersect inside a circle, \( AP\cdot PB = CP\cdot PD \). Substituting the expressions and the given numbers,</p>
\[ ({@eAP@})\,({@PB@}) = ({@CP@})\,({@ePD@}). \]
<p>Expanding and collecting the \(x\)-terms gives the linear equation</p>
\[ {@PB-CP@}\,x = {@CP*b2 - PB*b1@}, \qquad\text{so}\qquad x = {@ta_x@}. \]
<p>Substituting back, \( AP = {@ta_AP@}\) cm and \( PD = {@ta_PD@}\) cm.</p>]]></text>
    </generalfeedback>
    <defaultgrade>3</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025010100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Design backwards: AP0=p*r, PB=q*s, CP=p*s, PD0=q*r  => AP0*PB = CP*PD0.
   A solution value x is also drawn; AP = x + b1 and PD = x + b2 with
   b1 = AP0 - x, b2 = PD0 - x. p#q forces PB#CP, so (x+b1)PB = CP(x+b2) is a
   genuine linear equation with unique root x. Generate the admissible
   (p,q,r,s,x) on the fly: lengths in [2,12], chord totals <= 15, constants
   nonzero and |b| <= 6, and x distinct from AP and PD. */
params : sublist(
  create_list([p,q,r,s,xx], p,1,4, q,1,4, r,2,5, s,2,5, xx,2,9),
  lambda([z], is(
        z[1]#z[2] and z[3]#z[4]
    and z[1]*z[3]>=2 and z[1]*z[3]<=12 and z[2]*z[4]>=2 and z[2]*z[4]<=12
    and z[1]*z[4]>=2 and z[1]*z[4]<=12 and z[2]*z[3]>=2 and z[2]*z[3]<=12
    and z[1]*z[3]+z[2]*z[4]<=15 and z[1]*z[4]+z[2]*z[3]<=15
    and z[5]#z[1]*z[3] and z[5]#z[2]*z[3]
    and abs(z[1]*z[3]-z[5])<=6 and abs(z[2]*z[3]-z[5])<=6 )));
sel : rand(params);
AP : sel[1]*sel[3];
PB : sel[2]*sel[4];
CP : sel[1]*sel[4];
PD : sel[2]*sel[3];
xsol : sel[5];
b1 : AP - xsol;
b2 : PD - xsol;
eAP : x + b1;
ePD : x + b2;
ta_x : xsol;
ta_AP : AP;
ta_PD : PD;]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>\( AP={@eAP@},\ PB={@PB@},\ CP={@CP@},\ PD={@ePD@};\ x={@ta_x@},\ AP={@ta_AP@},\ PD={@ta_PD@} \)</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>Correct answer, well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>Your answer is partially correct.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>Incorrect answer.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>numerical</type>
      <tans>ta_x</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>numerical</type>
      <tans>ta_AP</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>numerical</type>
      <tans>ta_PD</tans>
      <boxsize>5</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>3</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct value of x.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans1</sans>
        <tans>ta_x</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-incorrect</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Use the intersecting-chords relationship \( AP\cdot PB = CP\cdot PD \). Substitute the two expressions and the two given numbers, expand the brackets, and solve the resulting linear equation for \(x\). Have another go.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct AP.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ta_AP</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Follow-through: AP consistent with the student's own x.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans2</sans>
        <tans>ans1+b1</tans>
        <testoptions>0.01</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-followthrough</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is consistent with your value of \(x\), so your substitution is right &mdash; but \(x\) itself is not correct yet. Re-solve part (a), then substitute again.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-incorrect</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Once you have \(x\), substitute it into the expression shown for \(AP\). Try again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct PD.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ta_PD</tans>
        <testoptions>0.01</testoptions>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description>Follow-through: PD consistent with the student's own x.</description>
        <answertest>NumAbsolute</answertest>
        <sans>ans3</sans>
        <tans>ans1+b2</tans>
        <testoptions>0.01</testoptions>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-followthrough</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This is consistent with your value of \(x\), so your substitution is right &mdash; but \(x\) itself is not correct yet. Re-solve part (a), then substitute again.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-incorrect</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Once you have \(x\), substitute it into the expression shown for \(PD\). Try again.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1001</deployedseed>
    <deployedseed>2002</deployedseed>
    <deployedseed>3003</deployedseed>
    <deployedseed>4004</deployedseed>
    <deployedseed>5005</deployedseed>
    <deployedseed>6006</deployedseed>
    <deployedseed>7007</deployedseed>
    <deployedseed>8008</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>All correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta_x</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta_AP</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta_PD</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Wrong x; AP and PD follow through from that x</description>
      <testinput>
        <name>ans1</name>
        <value>ev(ta_x+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta_AP+1,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(ta_PD+1,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-incorrect</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-followthrough</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-followthrough</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>3</testcase>
      <description>Correct x; AP wrong; PD correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta_x</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta_AP+2,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta_PD</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-incorrect</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>4</testcase>
      <description>All wrong, not consistent (generic)</description>
      <testinput>
        <name>ans1</name>
        <value>ev(ta_x+3,simp)</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ev(ta_AP+50,simp)</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ev(ta_PD+50,simp)</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-incorrect</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-incorrect</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt3-incorrect</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
