<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>16.1 Q3 Horizontal line test on an Eldoret sacco fare curve</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>A matatu sacco in Eldoret charges a fare that depends on distance. The manager wants to know whether two different distances ever give exactly the same fare. The board shows the fare curve and a horizontal line you can slide up and down.</p>
<p>Slide the horizontal line through the whole height of the graph. If some horizontal line meets the curve at two or more points, then two different distances give the same fare. Is this fare function one to one or many to one?</p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>

[[jsxgraph input-ref="ref1" width="500px" height="420px"]]
var board = JXG.JSXGraph.initBoard(divid, {boundingbox:[-1,8,10,-2], axis:true, showCopyright:false, showNavigation:false});
var o={#one2one#};
if (o==1){ board.create('functiongraph',[function(x){return 0.6*x+0.5;},0,9],{strokeColor:'#1565c0',strokeWidth:2}); }
else { board.create('functiongraph',[function(x){return 0.15*(x-4)*(x-4)+1.5;},0,9],{strokeColor:'#1565c0',strokeWidth:2}); }
var hy=board.create('slider',[[6.5,7],[9,7],[0,3,7]],{name:'line',snapWidth:0.1});
board.create('line',[[-1,function(){return hy.Value();}],[10,function(){return hy.Value();}]],{strokeColor:'#c62828',strokeWidth:2,dash:2});
board.create('text',[-0.5,-1.3,'Slide the red line through the whole height'],{fontSize:12,anchorX:'left'});
[[/jsxgraph]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>A function is one to one when no horizontal line meets its graph more than once. If some horizontal line crosses the curve twice, then two different inputs share the same output, and the function is many to one.</p>
<p>A single crossing is expected and is not a problem. The concern is two or more crossings at the same height. For this fare curve the answer is {@correct@}.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025010100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* one2one=1 -> strictly increasing (one to one);
   one2one=0 -> has a flat/turn so a horizontal line meets it twice. */
one2one : rand([0,1]);
correct : if one2one=1 then "one to one" else "many to one";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>Correct answer is {@correct@}.</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>That is correct. Well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>You are part of the way there.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>That is not correct yet.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>string</type>
      <tans>correct</tans>
      <boxsize>16</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>one to one</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>a1 : supcase(sconcat(ans1));
cor : supcase(correct);</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>is(a1=cor)</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct. Sliding the line through the whole graph gives the right conclusion.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(a1="ONE TO ONE") and is(cor="MANY TO ONE")]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-mistake1</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>It looks like you judged from the general upward shape rather than testing with the line. A curve that rises can still flatten and repeat a fare. Drag the horizontal line through the whole graph and you will find a height where it crosses twice, so the function is many to one.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>true</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-mistake2</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Remember that a single crossing is expected and is fine. Reporting many to one whenever the line touches the curve at all is too strict. The test looks for two or more crossings at the same height. Check again for a repeated height before deciding.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>301</deployedseed>
    <deployedseed>302</deployedseed>
    <deployedseed>303</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct</description>
      <testinput>
        <name>ans1</name>
        <value>correct</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
