<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>16.1 Q1 Sorting Kenyan relations into function or not a function</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>A form teacher in Kericho keeps two sets of records.</p>
<p>The first set links {@farmers@} different tea farmers to the single cooperative society {@coop@} that buys all their tea. Several farmers are linked to the same cooperative.</p>
<p>The second set links one national identity card number to two different farmer profiles, because of a mistake in the records.</p>
<p>Look at the diagrams on the interactive board. For the first set of records, is the relation a function?</p>
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>For the second set of records, where one identity number points to two different profiles, is the relation a function?</p>
<p>[[input:ans2]] [[validation:ans2]]</p>
<p>[[feedback:prt1]]</p>
<p>[[feedback:prt2]]</p>

[[jsxgraph input-ref-ans1="ref1" input-ref-ans2="ref2" width="560px" height="320px"]]
var board = JXG.JSXGraph.initBoard(divid, {boundingbox:[-6,5,6,-5], axis:false, grid:false, showCopyright:false, showNavigation:false});
board.create('text',[-5.5,4.3,'Records 1 (many farmers to one cooperative)'],{fontSize:12,anchorX:'left'});
var lyF=[3,2,1,0];
for (var i=0;i<{#farmers#};i++){ board.create('point',[-4.5,lyF[i]],{name:'',size:2,fixed:true,color:'#1565c0'});
  board.create('arrow',[[-4.3,lyF[i]],[-1.7,1.5]],{strokeColor:'#888',lastArrow:{size:5}}); }
board.create('point',[-1.5,1.5],{name:'"'+({#coop#})+'"',size:3,fixed:true,color:'#2e7d32',label:{offset:[8,0],anchorX:'left'}});
board.create('text',[0.5,4.3,'Records 2 (one ID to two profiles)'],{fontSize:12,anchorX:'left'});
board.create('point',[2,1.5],{name:'ID',size:3,fixed:true,color:'#c62828'});
board.create('point',[5,2.5],{name:'Profile X',size:2,fixed:true,color:'#1565c0',label:{offset:[8,0]}});
board.create('point',[5,0.5],{name:'Profile Y',size:2,fixed:true,color:'#1565c0',label:{offset:[8,0]}});
board.create('arrow',[[2.3,1.6],[4.7,2.5]],{strokeColor:'#c62828',lastArrow:{size:5}});
board.create('arrow',[[2.3,1.4],[4.7,0.5]],{strokeColor:'#c62828',lastArrow:{size:5}});
// visual only; no bind needed for classification inputs
[[/jsxgraph]]]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>A relation is a function when every input is linked to exactly one output.</p>
<p>In the first set of records, many farmers are linked to the one cooperative {@coop@}. Each farmer still has exactly one cooperative, so every input has exactly one output. This means the first relation is a function. Many inputs are allowed to share the same output.</p>
<p>In the second set of records, one identity number is linked to two different profiles. That input has more than one output, so the second relation is not a function.</p>]]></text>
    </generalfeedback>
    <defaultgrade>2</defaultgrade>
    <penalty>0.3333333</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2025010100</text>
    </stackversion>
    <questionvariables>
      <text><![CDATA[/* Two mapping diagrams. d1 is many-to-one (a function),
   d2 is one-to-many (not a function). The learner classifies each. */
farmers : rand([3,4,5]);
coop : rand(["Kericho Highlands","Nandi Hills","Sotik"]);
/* Correct classifications. */
ans1_correct : "function";
ans2_correct : "not a function";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="moodle_auto_format">
      <text>Records 1 is a function, Records 2 is not.</text>
    </questionnote>
    <questiondescription format="moodle_auto_format">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text>That is correct. Well done.</text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text>You are part of the way there.</text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text>That is not correct yet.</text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>string</type>
      <tans>ans1_correct</tans>
      <boxsize>16</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>function</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>string</type>
      <tans>ans2_correct</tans>
      <boxsize>16</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint>not a function</syntaxhint>
      <syntaxattribute>1</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>0</mustverify>
      <showvalidation>0</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>a1 : supcase(sconcat(ans1));</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(a1="FUNCTION")]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct. Many farmers sharing one cooperative still gives each farmer exactly one output, so this is a function.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>true</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-mistake1</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>This relation is actually a function. It is true that several arrows leave the left side, but that does not break the rule. A function only fails when a single input has more than one output. Here each farmer has exactly one cooperative, so many different inputs are allowed to share one output and it still counts as a function.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text>a2 : supcase(sconcat(ans2));</text>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans><![CDATA[is(a2="NOT A FUNCTION")]]></sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-correct</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct. One identity number points to two profiles, so that input has two outputs and the relation is not a function.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt-0-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>true</sans>
        <tans>true</tans>
        <testoptions/>
        <quiet>1</quiet>
        <truescoremode>=</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-mistake2</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Look again at the second diagram. Every arrow does point somewhere, but a function needs each input to have exactly one output, not just some output. The single identity number points to two different profiles, so this relation is not a function.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>101</deployedseed>
    <deployedseed>102</deployedseed>
    <deployedseed>103</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Both correct</description>
      <testinput>
        <name>ans1</name>
        <value><![CDATA["function"]]></value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value><![CDATA["not a function"]]></value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Both mistakes</description>
      <testinput>
        <name>ans1</name>
        <value><![CDATA["not a function"]]></value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value><![CDATA["function"]]></value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-mistake1</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt2-mistake2</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
