<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q1 - Meaning of a one-to-one relation</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Which of the following best describes a one-to-one relation?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>A one-to-one relation is a special type of function. In a function, each input in the domain must map to exactly one output in the codomain. For the relation to be one-to-one, we add one more condition. No two different inputs are allowed to share the same output. This means each output in the codomain can be traced back to at most one input.</p>
<p>Option A describes a relation where one input has many outputs, so it is not even a function.</p>
<p>Option B describes many inputs sharing one output, which is a many-to-one relation, not one-to-one.</p>
<p>Option D also describes one input with many outputs.</p>
<p>Only option C captures both requirements, a well-defined mapping and no repeated outputs. The correct answer is C.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options1: ["A. Every element of the domain maps to more than one element of the codomain.","B. Multiple elements of the domain map to the same element of the codomain.","C. Every element of the domain maps to a unique element of the codomain, and every element of the codomain has at most one pre-image.","D. An input value maps to multiple output values."];
tans1: "C. Every element of the domain maps to a unique element of the codomain, and every element of the codomain has at most one pre-image.";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. Read the worked solution below to confirm your reasoning.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. Read the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. Read the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
  </question>
</quiz>
