<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q7 - Which equation is one-to-many</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Which of the following equations represents a one-to-many relation?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Check each equation by asking whether a single input value of x can give more than one value of y.</p>
<p>For y = 3x minus 5, each x gives exactly one y, so this is one-to-one.</p>
<p>For y = x cubed, each x also gives exactly one y, and different x values give different y values, so this is one-to-one too.</p>
<p>For y equals the absolute value of x, each x gives one y, though different x values, such as 2 and negative 2, give the same y. This makes it many-to-one, not one-to-many.</p>
<p>Now check x squared plus y squared equals 9. This is the equation of a circle. Rearranging for y, we get two possible values for almost every x between negative 3 and 3, the positive square root of 9 minus x squared, and the negative square root of 9 minus x squared. So one input value of x gives two output values of y. This matches a one-to-many relation. The correct answer is B.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options7: ["A. y = 3x - 5","B. x^2 + y^2 = 9","C. y = x^3","D. y = |x|"];
tans7: "B. x^2 + y^2 = 9";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. Read the worked solution below to confirm your reasoning.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. Read the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. Read the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
  </question>
</quiz>
