<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q2 - Mapping type for f(x) = x squared</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Consider the function f(x) = x<sup>2</sup> for all real numbers. What type of mapping does this function represent?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>To classify the mapping, check whether different inputs can give the same output. Try x = 2 and x = negative 2. Both give f(x) = 4. So two different inputs, 2 and negative 2, map to the same output, 4. This pattern repeats for every nonzero input and its negative.</p>
<p>Since many inputs share the same output, the mapping is many-to-one. It cannot be one-to-many or many-to-many, because each input still gives only one output. It is also not one-to-one, since we just found two inputs sharing an output. The correct answer is B.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options2: ["A. One-to-many","B. Many-to-one","C. One-to-one","D. Many-to-many"];
tans2: "B. Many-to-one";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. Read the worked solution below to confirm your reasoning.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. Read the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. Read the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
  </question>
</quiz>
