<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q9 - National ID numbers as a one-to-one relation</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>In Kenya, every citizen who registers for a national identification card is issued exactly one ID number, and that ID number is never given to any other person.</p>
<p>Treating citizens as the domain and ID numbers as the codomain, which type of relation best describes this pairing?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Each citizen receives exactly one ID number, so no citizen is linked to more than one number. In addition, no ID number is ever reused for a different citizen, so each number traces back to only one person.</p>
<p>Since every element of the domain connects to a unique element of the codomain, and no output is shared, this is a one-to-one relation.</p>
<p>This kind of relation matters in real life, because it lets the government uniquely identify each citizen using their ID number, and it also allows the reverse process, finding the citizen who owns a given ID number. The correct answer is D.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options9: ["A. One-to-many","B. Many-to-one","C. Many-to-many","D. One-to-one"];
tans9: "D. One-to-one";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. Read the worked solution below to confirm your reasoning.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. Read the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. Read the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
  </question>
</quiz>
