<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q4 - Admission numbers at Green Valley Secondary School</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>At Green Valley Secondary School in Nakuru, every new form one student is given exactly one admission number, and no two students ever share the same admission number.</p>
<p>Part of the pairing is shown below.</p>
<p>Amina pairs with GV101.<br/>
Brian pairs with GV102.<br/>
Cynthia pairs with GV103.<br/>
David pairs with GV104.</p>
<p>Treating the students as the domain and the admission numbers as the codomain, what type of relation does this pairing represent?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Look at the pairing carefully. Each student has exactly one admission number, so no student is connected to more than one number. Also, each admission number belongs to exactly one student, so no admission number is shared between two students.</p>
<p>Since every element of the domain, the students, connects to a unique element of the codomain, the admission numbers, and every admission number traces back to only one student, the relation is one-to-one. The correct answer is C.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options4: ["A. One-to-many","B. Many-to-one","C. One-to-one","D. Many-to-many"];
tans4: "C. One-to-one";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. Read the worked solution below to confirm your reasoning.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. Read the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. Read the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
  </question>
</quiz>
