<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q8 - Restricting the domain of x squared</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The function f(x) = x<sup>2</sup> for all real numbers does not have an inverse function because it is many-to-one. A teacher suggests restricting the domain to x greater than or equal to zero. Why does this restriction allow the function to have an inverse?</p>
<p>[[input:ans8]]</p>
<p>[[validation:ans8]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The original function f(x) = x<sup>2</sup> is many-to-one because a positive number and its negative give the same square. For example, 3 and negative 3 both give 9.</p>
<p>An inverse function can only exist if the original relation is one-to-one, because the inverse must be able to send each output back to exactly one input.</p>
<p>By restricting the domain to x greater than or equal to zero, we remove all the negative inputs. Now every remaining input, being nonnegative, gives a different output from every other remaining input. The function has become one-to-one on this restricted domain, so an inverse function can now be defined. The correct answer is A.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options8: ["A. Restricting the domain changes the mapping from many-to-one to one-to-one, so each output now comes from only one input.","B. Restricting the domain removes the need for the function to pass the vertical line test.","C. Restricting the domain makes the function many-to-many, which always has an inverse.","D. Restricting the domain has no effect on whether an inverse exists."];
tans8: "A. Restricting the domain changes the mapping from many-to-one to one-to-one, so each output now comes from only one input.";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt8]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. See the worked solution below.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. See the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. See the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans8</name>
      <type>radio</type>
      <tans>tans8</tans>
      <boxsize>1</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options>options8</options>
    </input>
    <prt>
      <name>prt8</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>String</answertest>
        <sans>ans8</sans>
        <tans>tans8</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt8-0-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt8-0-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not correct.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
