<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q6 - Tea farmers near Kericho delivering to one center</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Farmers from four villages near Kericho, namely Kapsoit, Chilchila, Kabianga, and Ainamoi, all deliver their tea leaves to the very same collection center in Kericho town.</p>
<p>Treating the villages as the domain and the collection center as the codomain, what type of relation is this?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Here the domain is the four villages and the codomain is a single collection center. Every village connects to that same center, so many elements of the domain map to only one element of the codomain. This matches the definition of a many-to-one relation.</p>
<p>It is not one-to-many because no single village is being sent to more than one center. It is not one-to-one because more than one village shares the same center. It is not many-to-many because only one element of the codomain is being used. The correct answer is B.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options6: ["A. One-to-many","B. Many-to-one","C. One-to-one","D. Many-to-many"];
tans6: "B. Many-to-one";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. Read the worked solution below to confirm your reasoning.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. Read the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. Read the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
  </question>
</quiz>
