<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q3 - How relations and functions connect</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Which statement accurately describes the relationship between relations and functions?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>A relation is any set of ordered pairs connecting elements of one set to another. A function is a special kind of relation with an extra rule, every input must be connected to exactly one output.</p>
<p>Because of this extra rule, every function automatically qualifies as a relation, but not every relation qualifies as a function.</p>
<p>Option A is false because some relations allow one input to have several outputs, which breaks the function rule.</p>
<p>Option C is false because functions can be one-to-one or many-to-one, both are valid function types.</p>
<p>Option D is false because one-to-one relations can also be functions. The correct answer is B.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options3: ["A. All relations are functions.","B. All functions are relations.","C. Functions only include one-to-one relations.","D. A relation can only be classified as a function if it is many-to-one."];
tans3: "B. All functions are relations.";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. Read the worked solution below to confirm your reasoning.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. Read the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. Read the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
  </question>
</quiz>
