<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q8 - Restricting the domain of x squared</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>The function f(x) = x<sup>2</sup> for all real numbers does not have an inverse function because it is many-to-one. A teacher suggests restricting the domain to x greater than or equal to zero. Why does this restriction allow the function to have an inverse?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The original function f(x) = x<sup>2</sup> is many-to-one because a positive number and its negative give the same square. For example, 3 and negative 3 both give 9.</p>
<p>An inverse function can only exist if the original relation is one-to-one, because the inverse must be able to send each output back to exactly one input.</p>
<p>By restricting the domain to x greater than or equal to zero, we remove all the negative inputs. Now every remaining input, being nonnegative, gives a different output from every other remaining input. The function has become one-to-one on this restricted domain, so an inverse function can now be defined. The correct answer is A.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options8: ["A. Restricting the domain changes the mapping from many-to-one to one-to-one, so each output now comes from only one input.","B. Restricting the domain removes the need for the function to pass the vertical line test.","C. Restricting the domain makes the function many-to-many, which always has an inverse.","D. Restricting the domain has no effect on whether an inverse exists."];
tans8: "A. Restricting the domain changes the mapping from many-to-one to one-to-one, so each output now comes from only one input.";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. Read the worked solution below to confirm your reasoning.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. Read the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. Read the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
  </question>
</quiz>
