<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q5 - What the vertical line test tells us</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>How does the vertical line test help in determining if a graph represents a function?</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The vertical line test checks whether a graph represents a function. Draw an imaginary vertical line anywhere across the graph. If that line ever crosses the graph at more than one point, it means one input value on the horizontal axis is connected to more than one output value on the vertical axis. That breaks the basic rule of a function, which says each input must give only one output.</p>
<p>So if any vertical line meets the graph more than once, the graph does not represent a function.</p>
<p>Option A wrongly links this test to one-to-one functions, that job actually belongs to the horizontal line test. Options B and D describe the situation the wrong way round. The correct answer is C.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options5: ["A. If any vertical line intersects the graph at more than one point, the graph is a one-to-one function.","B. If any vertical line intersects the graph at exactly one point, the relation is not a function.","C. If any vertical line intersects the graph at more than one point, the graph is not a function.","D. If any vertical line intersects the graph at exactly one point, the graph is a many-to-one function."];
tans5: "C. If any vertical line intersects the graph at more than one point, the graph is not a function.";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. Read the worked solution below to confirm your reasoning.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. Read the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. Read the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
  </question>
</quiz>
