<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q7 - Which equation is one-to-many</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>Which of the following equations represents a one-to-many relation?</p>
<p>[[input:ans7]]</p>
<p>[[validation:ans7]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>Check each equation by asking whether a single input value of x can give more than one value of y.</p>
<p>For y = 3x minus 5, each x gives exactly one y, so this is one-to-one.</p>
<p>For y = x cubed, each x also gives exactly one y, and different x values give different y values, so this is one-to-one too.</p>
<p>For y equals the absolute value of x, each x gives one y, though different x values, such as 2 and negative 2, give the same y. This makes it many-to-one, not one-to-many.</p>
<p>Now check x squared plus y squared equals 9. This is the equation of a circle. Rearranging for y, we get two possible values for almost every x between negative 3 and 3, the positive square root of 9 minus x squared, and the negative square root of 9 minus x squared. So one input value of x gives two output values of y. This matches a one-to-many relation. The correct answer is B.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options7: ["A. y = 3x - 5","B. x^2 + y^2 = 9","C. y = x^3","D. y = |x|"];
tans7: "B. x^2 + y^2 = 9";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt7]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. See the worked solution below.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. See the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. See the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans7</name>
      <type>radio</type>
      <tans>tans7</tans>
      <boxsize>1</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options>options7</options>
    </input>
    <prt>
      <name>prt7</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>String</answertest>
        <sans>ans7</sans>
        <tans>tans7</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt7-0-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt7-0-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not correct.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
