<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Relations Q10 - What the horizontal line test determines</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>What does the horizontal line test determine for a given function?</p>
<p>[[input:ans10]]</p>
<p>[[validation:ans10]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>The horizontal line test is used after we already know a graph represents a function. Draw an imaginary horizontal line anywhere across the graph. If that line only ever touches the graph at one point, no matter where you draw it, it means every output value on the vertical axis came from only one input value. This tells us the function is one-to-one.</p>
<p>Because it is one-to-one, we can reverse the pairing without any conflict, which means the function has an inverse function.</p>
<p>If a horizontal line touches the graph more than once at any point, the function is many-to-one instead, and it will not have an inverse unless we restrict its domain. The correct answer is D.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text/>
    </stackversion>
    <questionvariables>
      <text><![CDATA[options10: ["A. If the function is even or odd.","B. If the function is a many-to-one relation.","C. If the function passes through the origin.","D. If the function is a one-to-one relation (and thus has an inverse function)."];
tans10: "D. If the function is a one-to-one relation (and thus has an inverse function).";]]></text>
    </questionvariables>
    <specificfeedback format="html">
      <text>[[feedback:prt10]]</text>
    </specificfeedback>
    <questionnote format="html">
      <text/>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<p>Correct. See the worked solution below.</p>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<p>Not quite. See the worked solution below.</p>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<p>Not correct. See the worked solution below.</p>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>1</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans10</name>
      <type>radio</type>
      <tans>tans10</tans>
      <boxsize>1</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options>options10</options>
    </input>
    <prt>
      <name>prt10</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>String</answertest>
        <sans>ans10</sans>
        <tans>tans10</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt10-0-T</trueanswernote>
        <truefeedback format="html">
          <text><![CDATA[<p>Correct.</p>]]></text>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt10-0-F</falseanswernote>
        <falsefeedback format="html">
          <text><![CDATA[<p>Not correct.</p>]]></text>
        </falsefeedback>
      </node>
    </prt>
  </question>
</quiz>
