<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>rationalising the denominator 3</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Rationalise the denominator and simplify the following expression, giving your answer in the form \(k\sqrt{m}\).</p>
\[ \frac{@p@}{{@q@}\sqrt{@n@}} \]
<p>[[input:ans1]] [[validation:ans1]]</p>
<p>[[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked solution</h3>

<p><strong>Step 1:</strong> Multiply the top and bottom by \(\sqrt{@n@}\). This removes the root from the denominator.</p>
\[ \frac{@p@}{{@q@}\sqrt{@n@}} \times \frac{\sqrt{@n@}}{\sqrt{@n@}} = \frac{{@p@}\sqrt{@n@}}{{@q@}\sqrt{@n@}\sqrt{@n@}} \]

<p><strong>Step 2:</strong> Simplify the bottom. A root multiplied by itself gives the number inside.</p>
\[ \sqrt{@n@} \times \sqrt{@n@} = {@n@}, \qquad \text{so the bottom is } {@q@} \times {@n@} = {@qn@} \]
\[ \frac{{@p@}\sqrt{@n@}}{@qn@} \]

<p><strong>Step 3:</strong> Divide the number on top by the number on the bottom.</p>
\[ {@p@} \div {@qn@} = {@k@} \]

<p><strong>Answer:</strong> \( {@k@}\sqrt{@n@} \)</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>simp:true;
S: [2,3,5,6,7,10,11,13,14,15,17,19];
n: rand(S);
q: rand(makelist(i,i,2,6));
k: rand(makelist(i,i,2,9));
p: k*q*n;
qn: q*n;
ta1: k*sqrt(n);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}{@p@}, {@q@}, {@n@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>676925765</deployedseed>
    <deployedseed>1374249863</deployedseed>
    <deployedseed>725072021</deployedseed>
    <deployedseed>1903225532</deployedseed>
    <deployedseed>460261343</deployedseed>
    <deployedseed>78010333</deployedseed>
    <deployedseed>277208370</deployedseed>
    <deployedseed>902834998</deployedseed>
    <deployedseed>1082645859</deployedseed>
    <deployedseed>2052690034</deployedseed>
    <deployedseed>111210439</deployedseed>
    <deployedseed>1695859632</deployedseed>
    <deployedseed>2003577017</deployedseed>
    <deployedseed>432908877</deployedseed>
    <deployedseed>829379093</deployedseed>
    <deployedseed>734853377</deployedseed>
    <deployedseed>378352608</deployedseed>
    <deployedseed>1263426154</deployedseed>
    <deployedseed>1906438856</deployedseed>
    <deployedseed>1335739337</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
