<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Rationalizing the denominator Q2</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>Rationalize the denominator of the following expression by multiplying by the conjugate, giving your answer in its simplest form:</p>
<p>\[ \frac{ {@p@} }{ \sqrt{ {@a@} } + \sqrt{ {@b@} } } \]</p>

<p><strong>Hint</strong>: You can type the answer of: \(\frac{2\left(\sqrt{4}-\sqrt{2}\right)}{2}\) as: <code>(2*(sqrt(4)-sqrt(2)))/2</code></p>

<p>[[input:ans1]] [[validation:ans1]] [[feedback:prt1]]</p>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h3>Worked Solution.</h3>
<p>To rationalize \( \dfrac{ {@p@} }{ \sqrt{ {@a@} } + \sqrt{ {@b@} } } \), multiply top and bottom by the <strong>conjugate</strong> of the denominator, \( \sqrt{ {@a@} } - \sqrt{ {@b@} } \):</p>
<p>\[ \frac{ {@p@} }{ \sqrt{ {@a@} } + \sqrt{ {@b@} } } \times \frac{ \sqrt{ {@a@} } - \sqrt{ {@b@} } }{ \sqrt{ {@a@} } - \sqrt{ {@b@} } } = \frac{ {@p@} (\sqrt{ {@a@} } - \sqrt{ {@b@} }) }{ {@a@} - {@b@} } \]</p>
<p><strong>Key idea:</strong> \( (\sqrt{@a@}+\sqrt{@b@})(\sqrt{@a@}-\sqrt{@b@}) = {@a@} - {@b@} \) a difference of squares which removes both square roots from the denominator in one step.</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>p: rand(8)+2;

sf_options: [2,3,5,6,7,10,11,13,14,15];

a: sf_options[rand(length(sf_options))+1];

b: sf_options[rand(length(sf_options))+1];
while b = a do b: sf_options[rand(length(sf_options))+1];

denom_original: sqrt(a) + sqrt(b);
conjugate: sqrt(a) - sqrt(b);

new_denom: a - b;
ta1: (p*(sqrt(a)-sqrt(b))) / (a-b);

ta2: (p*(sqrt(a)+sqrt(b))) / (a+b);

ta3: (p*(sqrt(a)-sqrt(b))) / (b-a);</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1686962858</deployedseed>
    <deployedseed>285438376</deployedseed>
    <deployedseed>705506926</deployedseed>
    <deployedseed>168021422</deployedseed>
    <deployedseed>973536801</deployedseed>
    <deployedseed>1482274541</deployedseed>
    <deployedseed>243379236</deployedseed>
    <deployedseed>1243594016</deployedseed>
    <deployedseed>1709692850</deployedseed>
    <deployedseed>1601243320</deployedseed>
    <deployedseed>796205468</deployedseed>
    <deployedseed>2129258895</deployedseed>
    <deployedseed>1236393058</deployedseed>
    <deployedseed>462532131</deployedseed>
    <deployedseed>1102599601</deployedseed>
    <deployedseed>823846035</deployedseed>
    <deployedseed>2064076209</deployedseed>
    <deployedseed>1101065311</deployedseed>
    <deployedseed>810025949</deployedseed>
  </question>
</quiz>
