<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Functions in Real Life — maximizing the area of a fenced plot</text>
    </name>
    <questiontext format="html">
      <text><![CDATA[<p>A farmer has {@perimeter@} metres of fencing to enclose a rectangular plot against an existing straight wall (so only three sides need fencing: two equal widths \(w\) and one length \(l\), opposite the wall).</p>
<p>The area enclosed, as a function of the width \(w\), is</p>
<p>\[ A(w) = w({@perimeter@}-2w). \]</p>
<p>Find the <strong>maximum possible area</strong>, in square metres.</p>
<div>[[input:ans1]] [[validation:ans1]]</div>
<div>[[feedback:prt1]]</div>]]></text>
    </questiontext>
    <generalfeedback format="html">
      <text><![CDATA[<p>\(A(w)=w({@perimeter@}-2w)={@Aexpand@}\) is a downward-opening parabola in \(w\) (the coefficient of \(w^2\) is negative), so it has a maximum at its vertex.</p>
<p>The vertex occurs at \(w={@wvertex@}\), but that is the <strong>width</strong> that gives the maximum &#8212; you still need to substitute it back into \(A(w)\) to get the maximum <strong>area</strong>:</p>
\[ A({@wvertex@}) = {@wvertex@}({@perimeter@}-2({@wvertex@})) = {@ta@}\ \text{m}^2. \]]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>perimeter : 2*rand_with_prohib(10,40,[]);      /* keep it even so the vertex width is a whole number */
Afun : w*(perimeter-2*w);
Aexpand : expand(Afun);
wvertex : -coeff(Aexpand,w,1)/(2*coeff(Aexpand,w,2));
ta : subst(w=wvertex,Afun);
wrong_vertexw : wvertex;                 /* gave the optimal WIDTH instead of the maximum AREA */</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text><![CDATA[<p>Grade 11 Functions 16.11 (Functions in Real Life). A real-life optimization problem using the vertex of a quadratic; catches the common mistake of reporting the vertex's w-coordinate (the optimal widt</p>]]></text>
    </questionnote>
    <questiondescription format="html">
      <text><![CDATA[<p id="description">Grade 11 Functions 16.11 (Functions in Real Life). A real-life optimization problem using the vertex of a quadratic; catches the common mistake of reporting the vertex's w-coordinate (the optimal width) instead of substituting back to find the maximum area itself.</p>]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta</tans>
      <boxsize>10</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>0</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description>Correct maximum area (not just the optimal width)?</description>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-correct</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-vertexonly</falseanswernote>
        <falsefeedback format="html">
          <text>That is the width that gives the maximum area, not the maximum area itself. Substitute this width back into A(w) to find the actual maximum area.</text>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1</deployedseed>
    <deployedseed>2</deployedseed>
    <deployedseed>3</deployedseed>
    <deployedseed>4</deployedseed>
    <deployedseed>5</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Correct</description>
      <testinput>
        <name>ans1</name>
        <value>ta</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>prt1-correct</expectedanswernote>
      </expected>
    </qtest>
    <qtest>
      <testcase>2</testcase>
      <description>Gave the optimal width instead of the max area</description>
      <testinput>
        <name>ans1</name>
        <value>wrong_vertexw</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>0.0000000</expectedscore>
        <expectedpenalty>0.1000000</expectedpenalty>
        <expectedanswernote>prt1-vertexonly</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
