<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>angles on the same arc with algebra</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>An amphitheatre is a circle with centre \(O\). The two ends of the stage, \(A\) and \(B\), are on its edge.
Two seats, \(C\) and \(D\), are also on the edge, on the same side of the stage.
Angle \(ACB = (3x + {@m@})^\circ\) and angle \(ADB = (5x - {@nn@})^\circ\).</p>

[[jsxgraph width="380px" height="400px"]]
var board = JXG.JSXGraph.initBoard(divid, {boundingbox:[-1.6,1.6,1.6,-1.6], axis:false, showCopyright:false, showNavigation:false, keepAspectRatio:true});
var r = Math.PI/180;
var v = {#v#}, pa = {#pa#}, qa = {#qa#};
function pt(nm, deg, col) {
  var t = deg*r;
  return board.create('point',[Math.cos(t),Math.sin(t)],{name:nm,fixed:true,size:5,strokeColor:col,fillColor:col,
    label:{offset:[16*Math.cos(t)-4,16*Math.sin(t)-4],strokeColor:col}});
}
board.create('circle',[[0,0],1],{fixed:true,strokeColor:'#6d28d9',strokeWidth:4,fillColor:'#ddd6fe',fillOpacity:0.55,highlight:false});
var O = board.create('point',[0,0],{name:'O',fixed:true,size:4,strokeColor:'#1e293b',fillColor:'#1e293b',label:{offset:[8,-18]}});
var A = pt('A', 90+v, '#ea580c');
var B = pt('B', 90-v, '#ea580c');
var C = pt('C', pa, '#2563eb');
var D = pt('D', qa, '#16a34a');
board.create('segment',[A,B],{fixed:true,strokeColor:'#92400e',strokeWidth:6,highlight:false});
board.create('segment',[C,A],{fixed:true,strokeColor:'#2563eb',strokeWidth:2});
board.create('segment',[C,B],{fixed:true,strokeColor:'#2563eb',strokeWidth:2});
board.create('segment',[D,A],{fixed:true,strokeColor:'#16a34a',strokeWidth:2});
board.create('segment',[D,B],{fixed:true,strokeColor:'#16a34a',strokeWidth:2});
board.create('segment',[O,A],{fixed:true,strokeColor:'#db2777',strokeWidth:2,dash:2});
board.create('segment',[O,B],{fixed:true,strokeColor:'#db2777',strokeWidth:2,dash:2});
board.create('angle',[B,C,A],{radius:0.3,fillColor:'#93c5fd',fillOpacity:0.9,strokeColor:'#2563eb',withLabel:false});
board.create('angle',[B,D,A],{radius:0.3,fillColor:'#86efac',fillOpacity:0.9,strokeColor:'#16a34a',withLabel:false});
board.create('angle',[B,O,A],{radius:0.25,fillColor:'#f9a8d4',fillOpacity:0.9,strokeColor:'#db2777',withLabel:false});
[[/jsxgraph]]

<p>(a) Find the value of \(x\). [[input:ans1]] [[validation:ans1]]</p>
<div>[[feedback:prt1]]</div>

<p>(b) Find angle \(ACB\). [[input:ans2]] [[validation:ans2]]</p>
<div>[[feedback:prt2]]</div>

<p>(c) Find angle \(AOB\). [[input:ans3]] [[validation:ans3]]</p>
<div>[[feedback:prt3]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h4>Worked solution</h4>
<p><strong>(a)</strong> \(C\) and \(D\) are on the same side of \(AB\), so angles \(ACB\) and \(ADB\) stand on the same arc and are equal:</p>
<p>\(3x + {@m@} = 5x - {@nn@}\)<br>
\({@m+nn@} = 2x\)<br>
\(x = {@x0@}\)</p>
<p><strong>(b)</strong> \(ACB = 3({@x0@}) + {@m@} = {@ta2@}^\circ\).<br>
Check: \(ADB = 5({@x0@}) - {@nn@} = {@ta2@}^\circ\).</p>
<p><strong>(c)</strong> The angle at the centre is twice the angle at the circumference:
\(AOB = 2 \times {@ta2@} = {@ta3@}^\circ\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>x0: rand_with_step(10,20,1);
m: rand_with_step(5,15,5);
nn: 2*x0-m;
v: 3*x0+m;
pa: rand_with_step(200,230,10);
qa: rand_with_step(280,330,10);
ta1: x0;
ta2: v;
ta3: 2*v;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@},{@ta2@},{@ta3@}, x = {@x0@}, m = {@m@}, pa = {@pa@}, qa = {@qa@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-2-T</trueanswernote>
        <truefeedback format="html">
          <text>Nice try but you answer gives   the angle at the centre. Angle ACB is at the circumference.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>1</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>v</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-2-T</trueanswernote>
        <truefeedback format="html">
          <text>{@v@} is the angle at the circumference. The angle at the centre is twice as big.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>2</falsenextnode>
        <falseanswernote>prt3-2-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
      <node>
        <name>2</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>v/2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>+</truescoremode>
        <truescore>0</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-3-T</trueanswernote>
        <truefeedback format="html">
          <text>You have halved the angle. The centre angle is double the circumference angle.</text>
        </truefeedback>
        <falsescoremode>-</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-3-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1324164089</deployedseed>
    <deployedseed>1744087038</deployedseed>
    <deployedseed>1731576993</deployedseed>
    <deployedseed>1923653760</deployedseed>
    <deployedseed>535835</deployedseed>
    <deployedseed>199743781</deployedseed>
    <deployedseed>1190134363</deployedseed>
    <deployedseed>1296828693</deployedseed>
    <deployedseed>97729882</deployedseed>
    <deployedseed>706422057</deployedseed>
    <deployedseed>32124431</deployedseed>
    <deployedseed>1114555872</deployedseed>
    <deployedseed>1841704938</deployedseed>
    <deployedseed>642917188</deployedseed>
    <deployedseed>873084923</deployedseed>
    <deployedseed>1595348916</deployedseed>
    <deployedseed>352896297</deployedseed>
    <deployedseed>1594782499</deployedseed>
    <deployedseed>1690764030</deployedseed>
    <qtest>
      <testcase>1</testcase>
      <description>Test case assuming the teacher's input gets full marks.</description>
      <testinput>
        <name>ans1</name>
        <value>ta1</value>
      </testinput>
      <testinput>
        <name>ans2</name>
        <value>ta2</value>
      </testinput>
      <testinput>
        <name>ans3</name>
        <value>ta3</value>
      </testinput>
      <expected>
        <name>prt1</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt1-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt2</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt2-1-T ]</expectedanswernote>
      </expected>
      <expected>
        <name>prt3</name>
        <expectedscore>1.0000000</expectedscore>
        <expectedpenalty>0.0000000</expectedpenalty>
        <expectedanswernote>[ prt3-1-T ]</expectedanswernote>
      </expected>
    </qtest>
  </question>
</quiz>
