<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="stack">
    <name>
      <text>Security cameras in a courtyard</text>
    </name>
    <questiontext format="moodle_auto_format">
      <text><![CDATA[<p>A circular courtyard has its centre at \(O\). Two gateposts, \(A\) and \(B\), stand on its edge.
A security camera at \(P\) on the edge sees both gateposts, and angle \(APB = {@v@}^\circ\).
A second camera is fixed at \(Q\), also on the edge and on the same side of \(AB\) as \(P\).</p>

[[jsxgraph width="380px" height="400px"]]
var board = JXG.JSXGraph.initBoard(divid, {boundingbox:[-1.6,1.6,1.6,-1.6], axis:false, showCopyright:false, showNavigation:false, keepAspectRatio:true});
var r = Math.PI/180;
var v = {#v#}, pa = {#pa#}, qa = {#qa#};
function pt(nm, deg, col) {
  var t = deg*r;
  return board.create('point',[Math.cos(t),Math.sin(t)],{name:nm,fixed:true,size:5,strokeColor:col,fillColor:col,
    label:{offset:[16*Math.cos(t)-4,16*Math.sin(t)-4],strokeColor:col}});
}
board.create('circle',[[0,0],1],{fixed:true,strokeColor:'#b45309',strokeWidth:4,fillColor:'#fde68a',fillOpacity:0.6,highlight:false});
var O = board.create('point',[0,0],{name:'O',fixed:true,size:4,strokeColor:'#1e293b',fillColor:'#1e293b',label:{offset:[8,-18]}});
var A = pt('A', 90+v, '#dc2626');
var B = pt('B', 90-v, '#dc2626');
var P = pt('P', pa, '#2563eb');
var Q = pt('Q', qa, '#16a34a');
board.create('segment',[A,B],{fixed:true,strokeColor:'#7c3aed',strokeWidth:3,dash:2});
board.create('segment',[P,A],{fixed:true,strokeColor:'#2563eb',strokeWidth:2});
board.create('segment',[P,B],{fixed:true,strokeColor:'#2563eb',strokeWidth:2});
board.create('segment',[Q,A],{fixed:true,strokeColor:'#16a34a',strokeWidth:2});
board.create('segment',[Q,B],{fixed:true,strokeColor:'#16a34a',strokeWidth:2});
board.create('segment',[O,A],{fixed:true,strokeColor:'#db2777',strokeWidth:2});
board.create('segment',[O,B],{fixed:true,strokeColor:'#db2777',strokeWidth:2});
board.create('angle',[B,P,A],{radius:0.3,fillColor:'#93c5fd',fillOpacity:0.9,strokeColor:'#2563eb',withLabel:false});
board.create('angle',[B,Q,A],{radius:0.3,fillColor:'#86efac',fillOpacity:0.9,strokeColor:'#16a34a',withLabel:false});
board.create('angle',[B,O,A],{radius:0.25,fillColor:'#f9a8d4',fillOpacity:0.9,strokeColor:'#db2777',withLabel:false});
[[/jsxgraph]]

<p>(a) Find angle \(AQB\). [[input:ans1]] [[validation:ans1]]</p>
<div>[[feedback:prt1]]</div>

<p>(b) Find angle \(AOB\). [[input:ans2]] [[validation:ans2]]</p>
<div>[[feedback:prt2]]</div>

<p>(c) Find angle \(OAB\). [[input:ans3]] [[validation:ans3]]</p>
<div>[[feedback:prt3]]</div>]]></text>
    </questiontext>
    <generalfeedback format="moodle_auto_format">
      <text><![CDATA[<h4>Worked solution</h4>
<p><strong>(a)</strong> \(P\) and \(Q\) are on the same side of \(AB\), so they stand on the same arc and the angles are equal:
\(AQB = {@v@}^\circ\).</p>
<p><strong>(b)</strong> The angle at the centre is twice the angle at the circumference:
\(AOB = 2 \times {@v@} = {@ta2@}^\circ\).</p>
<p><strong>(c)</strong> \(OA = OB\) because they are radii, so triangle \(OAB\) is isosceles:
\(OAB = (180 - {@ta2@}) \div 2 = {@ta3@}^\circ\).</p>]]></text>
    </generalfeedback>
    <defaultgrade>1</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber/>
    <stackversion>
      <text>2026063000</text>
    </stackversion>
    <questionvariables>
      <text>v: rand_with_step(30,60,5);
pa: rand_with_step(190,220,10);
qa: rand_with_step(250,290,10);
ta1: v;
ta2: 2*v;
ta3: 90-v;</text>
    </questionvariables>
    <specificfeedback format="html">
      <text/>
    </specificfeedback>
    <questionnote format="html">
      <text>{@ta1@},{@ta2@},{@ta3@},v = {@v@}, pa = {@pa@}, qa = {@qa@}</text>
    </questionnote>
    <questiondescription format="html">
      <text/>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done.]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
      <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>none</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <isbroken>0</isbroken>
    <variantsselectionseed/>
    <input>
      <name>ans1</name>
      <type>algebraic</type>
      <tans>ta1</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans2</name>
      <type>algebraic</type>
      <tans>ta2</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <input>
      <name>ans3</name>
      <type>algebraic</type>
      <tans>ta3</tans>
      <boxsize>15</boxsize>
      <strictsyntax>1</strictsyntax>
      <insertstars>0</insertstars>
      <syntaxhint/>
      <syntaxattribute>0</syntaxattribute>
      <forbidwords/>
      <allowwords/>
      <forbidfloat>1</forbidfloat>
      <requirelowestterms>0</requirelowestterms>
      <checkanswertype>0</checkanswertype>
      <mustverify>1</mustverify>
      <showvalidation>1</showvalidation>
      <options/>
    </input>
    <prt>
      <name>prt1</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans1</sans>
        <tans>ta1</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt1-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt1-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt2</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans2</sans>
        <tans>ta2</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt2-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt2-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <prt>
      <name>prt3</name>
      <value>1.0000000</value>
      <autosimplify>1</autosimplify>
      <feedbackstyle>1</feedbackstyle>
      <feedbackvariables>
        <text/>
      </feedbackvariables>
      <node>
        <name>0</name>
        <description/>
        <answertest>AlgEquiv</answertest>
        <sans>ans3</sans>
        <tans>ta3</tans>
        <testoptions/>
        <quiet>0</quiet>
        <truescoremode>=</truescoremode>
        <truescore>1</truescore>
        <truepenalty/>
        <truenextnode>-1</truenextnode>
        <trueanswernote>prt3-1-T</trueanswernote>
        <truefeedback format="html">
          <text/>
        </truefeedback>
        <falsescoremode>=</falsescoremode>
        <falsescore>0</falsescore>
        <falsepenalty/>
        <falsenextnode>-1</falsenextnode>
        <falseanswernote>prt3-1-F</falseanswernote>
        <falsefeedback format="html">
          <text/>
        </falsefeedback>
      </node>
    </prt>
    <deployedseed>1437839624</deployedseed>
    <deployedseed>654172441</deployedseed>
    <deployedseed>1525871001</deployedseed>
    <deployedseed>967784283</deployedseed>
    <deployedseed>1315828929</deployedseed>
    <deployedseed>735861519</deployedseed>
    <deployedseed>173805838</deployedseed>
    <deployedseed>1321177062</deployedseed>
    <deployedseed>945491152</deployedseed>
    <deployedseed>1788156776</deployedseed>
    <deployedseed>1170205691</deployedseed>
    <deployedseed>674345162</deployedseed>
    <deployedseed>1004530189</deployedseed>
    <deployedseed>112805215</deployedseed>
    <deployedseed>1708239836</deployedseed>
    <deployedseed>1261822741</deployedseed>
    <deployedseed>770353753</deployedseed>
    <deployedseed>2086606034</deployedseed>
    <deployedseed>2013510071</deployedseed>
  </question>
</quiz>
